```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("10_inheritance_q", "week_7", "readings", assignment_points = 24.0, assignment_tag = 'week7-readings')

# Initialize Otter
import otter
grader = otter.Notebook("10_inheritance_q.ipynb")
```

# ❓🏈 Inheritance in Python: Building NFL Madden Player Classes

![](./assets/figures/justinJ.jpeg)

The **Madden NFL** video game franchise has been a staple for football fans for decades. Each player in the game has unique **attributes and abilities**, but all players share common **characteristics** like **name, position, and rating**.  

In Python, **inheritance** allows us to model this by creating a **base Player class** and extending it into **specialized player types** like **Quarterbacks, Wide Receivers, and Defensive Linemen**.  

Let’s explore **how inheritance helps us organize player attributes and skills in Madden NFL** using Python classes!

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._10_inheritance_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._10_inheritance_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._10_inheritance_q import Question3
Question3().show()

```

## Question 1 (Points: 15.0): 🏈 Free Response: Create an Inheritance Structure for Madden NFL Players

Implement a class hierarchy for Madden NFL players:

- Create a **base class** `Player` with:
    - `name` (str) and `rating` (int).
    - A method `describe()` that prints player details. It should print `Player: <name>, Rating: <rating>`, where the `<name>` and `<rating>` are the player's attributes.
- Create a subclass `Quarterback` that:
    - Inherits from `Player`.
    - Adds a `throwing_power` attribute.
    - Overrides `describe()` to include `throwing_power`. To do this, you need to call super on the `describe()` method of the parent class.

### **Example Usage**
```python
qb = Quarterback("Lamar Jackson", 91, 96)
qb.describe()
```
**Expected Output:**
```
Player: Lamar Jackson, Rating: 91
Throwing Power: 96
```

```python
...

# Example usage
qb = Quarterback("Lamar Jackson", 91, 96)
qb.describe()
```

```python
grader.check("Madden-Inheritance-Classes")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week7-readings", "10_inheritance_q")
```

Tests for Madden-Inheritance-Classes (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(15.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(15.0)
>>> log_variable('total-points', f'week7-readings, 10_inheritance_q', 15.0)
>>> question_id = 'Madden-Inheritance-Classes-1'
>>> max_score = 1.0
>>> score = 0
>>> class_exists = 'Player' in globals() and 'Quarterback' in globals()
>>> assert class_exists, 'Both `Player` and `Quarterback` classes must be defined.'
>>> if class_exists:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_inheritance_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(class_exists))
```
failure_message: Failed: One or both classes (`Player`, `Quarterback`) are missing.
success_message: Success: `Player` and `Quarterback` classes are defined!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Madden-Inheritance-Classes-2'
>>> max_score = 2.0
>>> score = 0
>>> test_player = Player('Patrick Mahomes', 97)
>>> player_attributes = {'name': test_player.name, 'rating': test_player.rating}
>>> expected_attributes = {'name': 'Patrick Mahomes', 'rating': 97}
>>> condition = player_attributes == expected_attributes
>>> assert condition, f'Expected {expected_attributes}, but got {player_attributes}.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_inheritance_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(player_attributes))
```
failure_message: Failed: `Player` does not initialize attributes correctly.
success_message: Success: `Player` initializes attributes correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Madden-Inheritance-Classes-3'
>>> max_score = 2.0
>>> score = 0
>>> test_qb = Quarterback('Josh Allen', 92, 98)
>>> qb_attributes = {'name': test_qb.name, 'rating': test_qb.rating, 'throwing_power': test_qb.throwing_power}
>>> expected_attributes = {'name': 'Josh Allen', 'rating': 92, 'throwing_power': 98}
>>> condition = qb_attributes == expected_attributes
>>> assert condition, f'Expected {expected_attributes}, but got {qb_attributes}.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_inheritance_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(qb_attributes))
```
failure_message: Failed: `Quarterback` does not initialize attributes correctly.
success_message: Success: `Quarterback` initializes attributes correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Madden-Inheritance-Classes-4'
>>> max_score = 2.0
>>> score = 0
>>> is_subclass = issubclass(Quarterback, Player)
>>> assert is_subclass, '`Quarterback` must inherit from `Player`.'
>>> if is_subclass:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_inheritance_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(is_subclass))
```
failure_message: Failed: `Quarterback` does not inherit from `Player`.
success_message: Success: `Quarterback` correctly inherits from `Player`!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Madden-Inheritance-Classes-5'
>>> max_score = 2.0
>>> score = 0
>>> is_subclass = issubclass(Quarterback, Player)
>>> assert is_subclass, '`Quarterback` must inherit from `Player`.'
>>> if is_subclass:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_inheritance_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(is_subclass))
```
failure_message: Failed: `Quarterback` does not inherit from `Player`.
success_message: Success: `Quarterback` correctly inherits from `Player`!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import inspect
>>> question_id = 'Madden-Inheritance-Classes-6'
>>> max_score = 2.0
>>> score = 0
>>> qb_describe_code = inspect.getsource(Quarterback.describe)
>>> method_overridden = 'super().describe()' in qb_describe_code
>>> assert method_overridden, 'The `describe()` method in `Quarterback` must call `super().describe()`.'
>>> if method_overridden:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_inheritance_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(method_overridden))
```
failure_message: Failed: `Quarterback.describe()` does not override `Player.describe()`.
success_message: Success: `Quarterback.describe()` overrides `Player.describe()` correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import io
>>> from contextlib import redirect_stdout
>>> question_id = 'Madden-Inheritance-Classes-7'
>>> max_score = 2.0
>>> score = 0
>>> output = io.StringIO()
>>> test_player = Player('Joe Burrow', 95)
>>> with redirect_stdout(output):
...     test_player.describe()
>>> stdout_content = output.getvalue().strip()
>>> expected_output = 'Player: Joe Burrow, Rating: 95'
>>> condition = stdout_content == expected_output
>>> assert condition, f"Expected '{expected_output}', but got '{stdout_content}'."
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_inheritance_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(stdout_content))
```
failure_message: Failed: `describe()` does not print correctly for `Player`.
success_message: Success: `describe()` correctly prints for `Player`!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import io
>>> from contextlib import redirect_stdout
>>> question_id = 'Madden-Inheritance-Classes-8'
>>> max_score = 2.0
>>> score = 0
>>> output = io.StringIO()
>>> test_qb = Quarterback('Aaron Rodgers', 94, 97)
>>> with redirect_stdout(output):
...     test_qb.describe()
>>> stdout_content = output.getvalue().strip()
>>> expected_output = 'Player: Aaron Rodgers, Rating: 94\nThrowing Power: 97'
>>> condition = stdout_content == expected_output
>>> assert condition, f"Expected:\n'{expected_output}'\nbut got:\n'{stdout_content}'."
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_inheritance_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(stdout_content))
```
failure_message: Failed: `describe()` does not print correctly for `Quarterback`.
success_message: Success: `describe()` correctly prints for `Quarterback`!
