```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("12_pokemon_polymorphism_q", "week_7", "readings", assignment_points = 18.0, assignment_tag = 'week7-readings')

# Initialize Otter
import otter
grader = otter.Notebook("12_pokemon_polymorphism_q.ipynb")
```

# ❓🔥 Polymorphism in Python: Pokémon Battle System

![](./assets/figures/pokemon_go.jpg)

In the **Pokémon** universe, different Pokémon species share common actions like **attacking, using abilities, and evolving**, but they execute these actions in **different ways** depending on their type.  

This is similar to **polymorphism** in Python, where **multiple classes can share the same method names but implement them differently**. For example, both **Charizard** and **Pikachu** can use the **attack()** method, but Charizard might use **Flamethrower**, while Pikachu uses **Thunderbolt**.  

Let’s explore how **polymorphism** allows us to create flexible and reusable Pokémon battle systems in Python!

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._12_pokemon_polymorphism_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._12_pokemon_polymorphism_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._12_pokemon_polymorphism_q import Question3
Question3().show()

```

## Question 1 (Points: 11.0): ⚡ Free Response: Implement Polymorphism in a Pokémon Battle System

Implement a **Pokémon battle system** using polymorphism:

- Create a **base class** `Pokemon` with:
    - A method `attack()` that returns `"A Pokémon attacks!"`
- Create subclasses `Pikachu` and `Charizard` that:
    - Override `attack()` with `"Pikachu uses Thunderbolt!"` and `"Charizard uses Flamethrower!"`
- Create a function `battle(pokemon1, pokemon2)` that:
    - Calls `.attack()` on both Pokémon and prints the results.

### **Example Usage**
```python
p1 = Pikachu()
p2 = Charizard()
battle(p1, p2)
```
**Expected Output:**
```
Pikachu uses Thunderbolt!
Charizard uses Flamethrower!
```

```python
...


# Example usage
p1 = Pikachu()
p2 = Charizard()
battle(p1, p2)
```

```python
grader.check("Pokemon-Polymorphism-Classes")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week7-readings", "12_pokemon_polymorphism_q")
```

Tests for Pokemon-Polymorphism-Classes (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(11.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(11.0)
>>> log_variable('total-points', f'week7-readings, 12_pokemon_polymorphism_q', 11.0)
>>> question_id = 'Pokemon-Polymorphism-Classes-1'
>>> max_score = 2.0
>>> score = 0
>>> p = Pokemon()
>>> assert p.attack() == 'A Pokémon attacks!', 'Incorrect base class attack method.'
>>> if p.attack() == 'A Pokémon attacks!':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('12_pokemon_polymorphism_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The base Pokemon class is missing attributes.
success_message: Success: The base Pokemon class initializes correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Pokemon-Polymorphism-Classes-2'
>>> max_score = 3.0
>>> score = 0
>>> pikachu = Pikachu()
>>> condition = pikachu.attack() == 'Pikachu uses Thunderbolt!'
>>> assert condition, 'Incorrect Pikachu attack method.'
>>> if condition:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('12_pokemon_polymorphism_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Pikachu's attack() does not return the expected string.
success_message: Success: Pikachu subclass overrides attack() correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Pokemon-Polymorphism-Classes-3'
>>> max_score = 3.0
>>> score = 0
>>> charizard = Charizard()
>>> condition = charizard.attack() == 'Charizard uses Flamethrower!'
>>> assert condition, 'Incorrect Charizard attack method.'
>>> if condition:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('12_pokemon_polymorphism_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Charizard's attack() does not return the expected string.
success_message: Success: Charizard subclass overrides attack() correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import io
>>> import sys
>>> question_id = 'Pokemon-Polymorphism-Classes-4'
>>> max_score = 3.0
>>> score = 0
>>> output_capture = io.StringIO()
>>> sys.stdout = output_capture
>>> battle(Pikachu(), Charizard())
>>> sys.stdout = sys.__stdout__
>>> expected_output = 'Pikachu uses Thunderbolt!\nCharizard uses Flamethrower!'
>>> condition = output_capture.getvalue().strip() == expected_output
>>> assert condition, 'battle() does not return expected results.'
>>> if condition:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('12_pokemon_polymorphism_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The battle function does not properly call attack() methods.
success_message: Success: The battle function correctly calls attack() on both Pokémon!
