```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("16_GMT_q", "week_7", "readings", assignment_points = 21.0, assignment_tag = 'week7-readings')

# Initialize Otter
import otter
grader = otter.Notebook("16_GMT_q.ipynb")
```

# ⌚ Multiple Inheritance in Python: The Rolex GMT Master

![](./assets/figures/GMT-rolex.jpg)

The **Rolex GMT-Master** is one of the most iconic watches ever created, designed for **pilots and travelers** who needed to track multiple time zones. The watch combines **precise movement**, **durability**, and **aesthetic elegance**.  

Similarly, in Python, **multiple inheritance** allows a class to combine **features from multiple parent classes**, just like the **GMT-Master blends aviation functionality with luxury craftsmanship**.  

Let’s explore how **multiple inheritance** enables us to **design complex and versatile class structures**—just like Rolex engineers do when crafting legendary watches.

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._16_GMT_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._16_GMT_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._16_GMT_q import Question3
Question3().show()

```

## Question 1 (Points: 12.0): ⌚ Free Response: Implement Multiple Inheritance for a Rolex GMT Watch

Design a **Python class structure** that models a **Rolex GMT-Master watch** using **multiple inheritance**:

- Create a **base class** `Timekeeping` with:
    - A method `show_time()` that returns `"Showing local time."`
- Create another **base class** `GMTFunctionality` with:
    - A method `show_gmt()` that returns `"Showing GMT time."`
- Create a subclass `RolexGMT` that:
    - Inherits from **both** `Timekeeping` and `GMTFunctionality`.
    - Overrides `show_time()` to include **both local and GMT time**.

### **Example Usage**
```python
watch = RolexGMT()
print(watch.show_time())
print(watch.show_gmt())
```
**Expected Output:**
```
Local Time: Showing local time.
GMT Time: Showing GMT time.
```

```python
...

# Example usage
watch = RolexGMT()
print(watch.show_time())
print(watch.show_gmt())
```

```python
grader.check("Rolex-GMT-Multiple-Inheritance")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week7-readings", "16_GMT_q")
```

Tests for Rolex-GMT-Multiple-Inheritance (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(12.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(12.0)
>>> log_variable('total-points', f'week7-readings, 16_GMT_q', 12.0)
>>> question_id = 'Rolex-GMT-Multiple-Inheritance-1'
>>> max_score = 3.0
>>> score = 0
>>> t = Timekeeping()
>>> assert t.show_time() == 'Showing local time.', 'Incorrect Timekeeping method output.'
>>> if t.show_time() == 'Showing local time.':
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('16_GMT_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The base Timekeeping class is missing required methods.
success_message: Success: The base Timekeeping class initializes correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Rolex-GMT-Multiple-Inheritance-2'
>>> max_score = 3.0
>>> score = 0
>>> g = GMTFunctionality()
>>> condition = g.show_gmt() == 'Showing GMT time.'
>>> assert condition, 'Incorrect GMTFunctionality method output.'
>>> if condition:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('16_GMT_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The GMTFunctionality class is missing required methods.
success_message: Success: The GMTFunctionality class initializes correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Rolex-GMT-Multiple-Inheritance-3'
>>> max_score = 3.0
>>> score = 0
>>> r = RolexGMT()
>>> condition = isinstance(r, Timekeeping) and isinstance(r, GMTFunctionality)
>>> assert condition, 'RolexGMT does not inherit properly.'
>>> if condition:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('16_GMT_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: RolexGMT does not inherit properly from both base classes.
success_message: Success: The RolexGMT class correctly inherits from both base classes!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Rolex-GMT-Multiple-Inheritance-4'
>>> max_score = 3.0
>>> score = 0
>>> r = RolexGMT()
>>> output = r.show_time()
>>> expected_local = 'Local Time: Showing local time.'
>>> expected_gmt = 'GMT Time: Showing GMT time.'
>>> condition = expected_local in output and expected_gmt in output
>>> assert condition, 'RolexGMT show_time() is missing required outputs.'
>>> if condition:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('16_GMT_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: RolexGMT show_time() does not include both local and GMT time.
success_message: Success: The RolexGMT class correctly overrides show_time()!
