```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("3_headis_q", "week_7", "readings", assignment_points = 19.0, assignment_tag = 'week7-readings')

# Initialize Otter
import otter
grader = otter.Notebook("3_headis_q.ipynb")
```

# 🏓 Classes and Instances in Headis

<iframe width="560" height="315" src="https://www.youtube.com/embed/DjFXXhJ-cqA?si=4Y96DeKaETRPL_7t" title="YouTube video player" frameborder="0" allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share" referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>

Headis is a fast-paced sport that blends table tennis and soccer, where players use their heads instead of paddles to hit the ball. Just like in programming, where we define **classes** to structure objects, Headis players have their own attributes and behaviors—speed, agility, and skill levels. Let’s explore how we can represent this sport using Python classes!

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._3_headis_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._3_headis_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._3_headis_q import Question3
Question3().show()

```

## Question 1 (Points: 12.0): 🎾 Free Response: Design a HeadisPlayer Class

Imagine you're creating a **HeadisPlayer** class in Python. Your class should include:

- Attributes for the player's **name**, **speed**, and **accuracy**.
- A method called `hit_ball()` that prints `"Player {name} hits the ball!"`
- A method `improve_skill()` that increases **accuracy** by a given `amount`.

**Write your class definition below.**

```python
...
```

```python
grader.check("headis-player-class-design")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week7-readings", "3_headis_q")
```

Tests for headis-player-class-design (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(12.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(12.0)
>>> log_variable('total-points', f'week7-readings, 3_headis_q', 12.0)
>>> question_id = 'headis-player-class-design-1'
>>> max_score = 1.0
>>> score = 0
>>> class_exists = 'HeadisPlayer' in globals()
>>> assert class_exists, 'The class `HeadisPlayer` must be defined in the script.'
>>> if class_exists:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_headis_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(class_exists))
```
failure_message: Failed: `HeadisPlayer` class is not defined.
success_message: Success: `HeadisPlayer` class is defined!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'headis-player-class-design-2'
>>> max_score = 2.0
>>> score = 0
>>> test_player = HeadisPlayer('Jordan', 7, 80)
>>> player_attributes = {'name': test_player.name, 'speed': test_player.speed, 'accuracy': test_player.accuracy}
>>> expected_attributes = {'name': 'Jordan', 'speed': 7, 'accuracy': 80}
>>> condition = player_attributes == expected_attributes
>>> assert condition, f'Expected {expected_attributes}, but got {player_attributes}.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_headis_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(player_attributes))
```
failure_message: Failed: `HeadisPlayer` does not initialize attributes correctly.
success_message: Success: `HeadisPlayer` initializes correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import io
>>> from contextlib import redirect_stdout
>>> question_id = 'headis-player-class-design-3'
>>> max_score = 2.0
>>> score = 0
>>> output = io.StringIO()
>>> test_player = HeadisPlayer('Serena', 9, 85)
>>> with redirect_stdout(output):
...     test_player.hit_ball()
>>> stdout_content = output.getvalue().strip()
>>> expected_output = 'Player Serena hits the ball!'
>>> condition = stdout_content == expected_output
>>> assert condition, f"Expected '{expected_output}', but got '{stdout_content}'."
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_headis_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(stdout_content))
```
failure_message: Failed: `hit_ball()` does not print correctly.
success_message: Success: `hit_ball()` prints the correct message!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import io
>>> from contextlib import redirect_stdout
>>> question_id = 'headis-player-class-design-4'
>>> max_score = 2.0
>>> score = 0
>>> output = io.StringIO()
>>> test_player = HeadisPlayer('Serena', 9, 85)
>>> with redirect_stdout(output):
...     test_player.hit_ball()
>>> stdout_content = output.getvalue().strip()
>>> expected_output = 'Player Serena hits the ball!'
>>> condition = stdout_content == expected_output
>>> assert condition, f"Expected '{expected_output}', but got '{stdout_content}'."
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_headis_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(stdout_content))
```
failure_message: Failed: `hit_ball()` does not print correctly.
success_message: Success: `hit_ball()` prints the correct message!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import io
>>> from contextlib import redirect_stdout
>>> question_id = 'headis-player-class-design-5'
>>> max_score = 3.0
>>> score = 0
>>> test_player = HeadisPlayer('Nadal', 6, 70)
>>> output = io.StringIO()
>>> with redirect_stdout(output):
...     test_player.improve_skill(10)
>>> stdout_content = output.getvalue().strip()
>>> updated_accuracy = test_player.accuracy
>>> expected_accuracy = 80
>>> expected_output = "Nadal's accuracy increased to 80!"
>>> condition1 = updated_accuracy == expected_accuracy
>>> assert condition1, f'Expected accuracy {expected_accuracy}, but got {updated_accuracy}.'
>>> condition2 = stdout_content == expected_output
>>> assert condition2, f"Expected '{expected_output}', but got '{stdout_content}'."
>>> if condition1 and condition2:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_headis_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(updated_accuracy))
```
failure_message: Failed: `improve_skill()` does not update accuracy correctly.
success_message: Success: `improve_skill()` correctly increases accuracy!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import io
>>> from contextlib import redirect_stdout
>>> question_id = 'headis-player-class-design-6'
>>> max_score = 2.0
>>> score = 0
>>> test_player = HeadisPlayer('Federer', 8, 90)
>>> output = io.StringIO()
>>> with redirect_stdout(output):
...     test_player.improve_skill(-20)
>>> updated_accuracy = test_player.accuracy
>>> expected_accuracy = 70
>>> condition = updated_accuracy == expected_accuracy
>>> assert condition, f'Expected accuracy {expected_accuracy}, but got {updated_accuracy}.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_headis_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(updated_accuracy))
```
failure_message: Failed: `improve_skill()` does not handle negative values correctly.
success_message: Success: `improve_skill()` correctly handles negative values!
