```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("5_change_methods_q", "week_7", "readings", assignment_points = 26.0, assignment_tag = 'week7-readings')

# Initialize Otter
import otter
grader = otter.Notebook("5_change_methods_q.ipynb")
```

# ❓✍️ Implementing Change Methods in Classes: The Evolution of Helvetica

<iframe width="560" height="315" src="https://www.youtube.com/embed/Ayg_K_CR2tc?si=peNGIzO6iducMzPX" title="YouTube video player" frameborder="0" allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share" referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>


Helvetica, one of the most iconic typefaces, has undergone several refinements since its creation in **1957** by Swiss type designer **Max Miedinger**. Originally named **Neue Haas Grotesk**, it was later adapted and rebranded as **Helvetica** to appeal to international audiences.  

Just as Helvetica was updated and refined over time, we often need to **modify objects in Python classes** by implementing **change methods** that update attributes dynamically. Let’s explore how Python **change methods** can help model **the evolution of Helvetica** in an object-oriented way!

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._5_change_methods_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._5_change_methods_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._5_change_methods_q import Question3
Question3().show()

```

## Question 1 (Points: 17.0): 🖋️ Free Response: Implement Change Methods for a `Font` Class

Design a Python class **`Font`** that models the **evolution of Helvetica**. Your class should:

- Have attributes for **name**, **size**, and **weight** (e.g., `"Helvetica"`, `12`, `"Regular"`).
- Include the following **change methods**:
    - `update_name(self, new_name)`: Changes the font name.
    - `set_size(self, new_size)`: Updates the font size.
    - `change_weight(self, new_weight)`: Adjusts the font weight.
    - **(Bonus)** Return `self` in each method to allow **method chaining**.

### **Example Usage**
```python
f = Font("Helvetica", 12, "Regular")
f.update_name("Helvetica Neue").set_size(14).change_weight("Bold")
print(f.name, f.size, f.weight)  # Output: Helvetica Neue 14 Bold
```

```python
...

# Example usage
f = Font("Helvetica", 12, "Regular")
f.update_name("Helvetica Neue").set_size(14).change_weight("Bold")
print(f.name, f.size, f.weight)  # Expected: Helvetica Neue 14 Bold
```

```python
grader.check("Helvetica-Class-Methods")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week7-readings", "5_change_methods_q")
```

Tests for Helvetica-Class-Methods (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(17.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(17.0)
>>> log_variable('total-points', f'week7-readings, 5_change_methods_q', 17.0)
>>> question_id = 'Helvetica-Class-Methods-1'
>>> max_score = 1.0
>>> score = 0
>>> class_exists = 'Font' in globals()
>>> assert class_exists, 'The class `Font` must be defined in the script.'
>>> if class_exists:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_change_methods_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(class_exists))
```
failure_message: Failed: `Font` class is not defined.
success_message: Success: `Font` class is defined!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Helvetica-Class-Methods-2'
>>> max_score = 2.0
>>> score = 0
>>> test_font = Font('Arial', 16, 'Medium')
>>> font_attributes = {'name': test_font.name, 'size': test_font.size, 'weight': test_font.weight}
>>> expected_attributes = {'name': 'Arial', 'size': 16, 'weight': 'Medium'}
>>> condition = font_attributes == expected_attributes
>>> assert condition, f'Expected {expected_attributes}, but got {font_attributes}.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_change_methods_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(font_attributes))
```
failure_message: Failed: `Font` does not initialize attributes correctly.
success_message: Success: `Font` initializes attributes correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Helvetica-Class-Methods-3'
>>> max_score = 2.0
>>> score = 0
>>> test_font = Font('Times', 14, 'Bold')
>>> out = test_font.update_name('Times New Roman')
>>> updated_name = test_font.name
>>> expected_name = 'Times New Roman'
>>> condition = updated_name == expected_name
>>> assert condition, f'Expected name to be {expected_name}, but got {updated_name}.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_change_methods_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(updated_name))
```
failure_message: Failed: `update_name()` does not update the font name correctly.
success_message: Success: `update_name()` correctly updates the font name!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Helvetica-Class-Methods-4'
>>> max_score = 2.0
>>> score = 0
>>> test_font = Font('Times', 14, 'Bold')
>>> out = test_font.update_name('Times New Roman')
>>> updated_name = test_font.name
>>> expected_name = 'Times New Roman'
>>> condition = updated_name == expected_name
>>> assert condition, f'Expected name to be {expected_name}, but got {updated_name}.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_change_methods_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(updated_name))
```
failure_message: Failed: `update_name()` does not update the font name correctly.
success_message: Success: `update_name()` correctly updates the font name!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Helvetica-Class-Methods-5'
>>> max_score = 2.0
>>> score = 0
>>> test_font = Font('Verdana', 10, 'Light')
>>> out = test_font.set_size(18)
>>> updated_size = test_font.size
>>> expected_size = 18
>>> condition = updated_size == expected_size
>>> assert condition, f'Expected size to be {expected_size}, but got {updated_size}.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_change_methods_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(updated_size))
```
failure_message: Failed: `set_size()` does not update the font size correctly.
success_message: Success: `set_size()` correctly updates the font size!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Helvetica-Class-Methods-6'
>>> max_score = 2.0
>>> score = 0
>>> test_font = Font('Courier', 12, 'Thin')
>>> out = test_font.change_weight('Black')
>>> updated_weight = test_font.weight
>>> expected_weight = 'Black'
>>> condition = updated_weight == expected_weight
>>> assert condition, f'Expected weight to be {expected_weight}, but got {updated_weight}.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_change_methods_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(updated_weight))
```
failure_message: Failed: `change_weight()` does not update the font weight correctly.
success_message: Success: `change_weight()` correctly updates the font weight!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Helvetica-Class-Methods-7'
>>> max_score = 3.0
>>> score = 0
>>> test_font = Font('Tahoma', 14, 'Normal')
>>> chaining_result = test_font.update_name('Tahoma Pro').set_size(20).change_weight('Extra Bold')
>>> chaining_success = isinstance(chaining_result, Font)
>>> assert chaining_success, 'Methods must return `self` to support method chaining.'
>>> if chaining_success:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_change_methods_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(chaining_success))
```
failure_message: Failed: Methods do not return `self` for chaining.
success_message: Success: Method chaining works correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Helvetica-Class-Methods-8'
>>> max_score = 3.0
>>> score = 0
>>> test_font = Font('Comic Sans', 12, 'Regular')
>>> out = test_font.update_name('Comic Sans Pro').set_size(16).change_weight('SemiBold')
>>> final_values = {'name': test_font.name, 'size': test_font.size, 'weight': test_font.weight}
>>> expected_values = {'name': 'Comic Sans Pro', 'size': 16, 'weight': 'SemiBold'}
>>> condition = final_values == expected_values
>>> assert condition, f'Expected {expected_values}, but got {final_values}.'
>>> if condition:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_change_methods_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(final_values))
```
failure_message: Failed: BONUS Sequential updates do not modify the `Font` instance as expected.
success_message: Success: BONUS Multiple updates correctly modify the `Font` instance!
