```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("7_magic_methods_horseshoe_q", "week_7", "readings", assignment_points = 22.0, assignment_tag = 'week7-readings')

# Initialize Otter
import otter
grader = otter.Notebook("7_magic_methods_horseshoe_q.ipynb")
```

# 🦀 Magic Methods and the "Blue Blood" of Horseshoe Crabs

![](./assets/figures/blue-blood-horseshoe.jpg)

Horseshoe crabs have survived for **hundreds of millions of years**, thanks in part to their unique **blue blood**. This blood contains **hemocyanin**, which transports oxygen and has remarkable antibacterial properties, making it vital for **modern medicine**.  

In Python, **magic methods** (also called **dunder methods**) give classes special abilities, just as hemocyanin grants horseshoe crabs extraordinary resilience. These methods allow objects to support **addition, string representation, comparisons**, and more.

Let’s explore how **magic methods** empower Python classes, just as the blue blood of horseshoe crabs has ensured their survival through the ages!

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._7_magic_methods_horseshoe_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._7_magic_methods_horseshoe_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._7_magic_methods_horseshoe_q import Question3
Question3().show()

```

## Question 1 (Points: 13.0): 🦀 Free Response: Implement Magic Methods in a `HorseshoeCrab` Class

Design a Python class `HorseshoeCrab` that includes:

- **`__init__`**: Initializes attributes for `species` (str) and `hemocyanin_level` (int).
- **`__str__`**: Returns a string in the format `"Horseshoe Crab of species <species>, hemocyanin level: <hemocyanin_level>"`.
- **`__add__`**: Combines two crabs' hemocyanin levels when using `+`. Add must make sure that the species of both crabs is the same, and should return a new `HorseshoeCrab` instance.
                            if the crabs are not the same species return `ValueError("Crabs must be of the same species to combine hemocyanin levels.")`.
- **`__eq__`**: Returns `True` if two crabs have the same `species` and `hemocyanin_level`.

**Example Usage**
```python
h1 = HorseshoeCrab("Limulus polyphemus", 100)
h2 = HorseshoeCrab("Limulus polyphemus", 150)
print(h1)  # Output: "Horseshoe Crab of species Limulus polyphemus, hemocyanin level: 100"
h3 = h1 + h2
print(h3)  # Output: "Horseshoe Crab of species Limulus polyphemus, hemocyanin level: 250"
print(h1 == h2)  # Output: False
```

```python
...

# Example usage
h1 = HorseshoeCrab("Limulus polyphemus", 100)
h2 = HorseshoeCrab("Limulus polyphemus", 150)
h3 = h1 + h2
print(h3)  # Expected: "Horseshoe Crab of species Limulus polyphemus, hemocyanin level: 250"
print(h1 == h2)  # Expected: False
```

```python
grader.check("HorseshoeCrab-Magic-Methods")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week7-readings", "7_magic_methods_horseshoe_q")
```

Tests for HorseshoeCrab-Magic-Methods (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(13.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(13.0)
>>> log_variable('total-points', f'week7-readings, 7_magic_methods_horseshoe_q', 13.0)
>>> question_id = 'HorseshoeCrab-Magic-Methods-1'
>>> max_score = 1.0
>>> score = 0
>>> class_exists = 'HorseshoeCrab' in globals()
>>> assert class_exists, 'The class `HorseshoeCrab` must be defined in the script.'
>>> if class_exists:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('7_magic_methods_horseshoe_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(class_exists))
```
failure_message: Failed: `HorseshoeCrab` class is not defined.
success_message: Success: `HorseshoeCrab` class is defined!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'HorseshoeCrab-Magic-Methods-2'
>>> max_score = 2.0
>>> score = 0
>>> test_crab = HorseshoeCrab('Tachypleus tridentatus', 120)
>>> crab_attributes = {'species': test_crab.species, 'hemocyanin_level': test_crab.hemocyanin_level}
>>> expected_attributes = {'species': 'Tachypleus tridentatus', 'hemocyanin_level': 120}
>>> condition = crab_attributes == expected_attributes
>>> assert condition, f'Expected {expected_attributes}, but got {crab_attributes}.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('7_magic_methods_horseshoe_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(crab_attributes))
```
failure_message: Failed: `HorseshoeCrab` does not initialize attributes correctly.
success_message: Success: `HorseshoeCrab` initializes attributes correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'HorseshoeCrab-Magic-Methods-3'
>>> max_score = 2.0
>>> score = 0
>>> test_crab = HorseshoeCrab('Carcinoscorpius rotundicauda', 85)
>>> crab_string = str(test_crab)
>>> expected_string = 'Horseshoe Crab of species Carcinoscorpius rotundicauda, hemocyanin level: 85'
>>> condition = crab_string == expected_string
>>> assert condition, f"Expected '{expected_string}', but got '{crab_string}'."
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('7_magic_methods_horseshoe_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(crab_string))
```
failure_message: Failed: `__str__()` does not return the correct string.
success_message: Success: `__str__()` correctly represents the crab!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'HorseshoeCrab-Magic-Methods-4'
>>> max_score = 3.0
>>> score = 0
>>> crab1 = HorseshoeCrab('Limulus polyphemus', 200)
>>> crab2 = HorseshoeCrab('Limulus polyphemus', 300)
>>> combined_crab = crab1 + crab2
>>> combined_hemocyanin = combined_crab.hemocyanin_level
>>> expected_hemocyanin = 500
>>> condition = combined_hemocyanin == expected_hemocyanin
>>> assert condition, f'Expected hemocyanin level {expected_hemocyanin}, but got {combined_hemocyanin}.'
>>> if condition:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('7_magic_methods_horseshoe_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(combined_hemocyanin))
```
failure_message: Failed: `__add__()` does not add hemocyanin levels correctly.
success_message: Success: `__add__()` correctly adds hemocyanin levels!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'HorseshoeCrab-Magic-Methods-5'
>>> max_score = 3.0
>>> score = 0
>>> crab1 = HorseshoeCrab('Limulus polyphemus', 200)
>>> crab2 = HorseshoeCrab('Limulus polyphemus', 300)
>>> combined_crab = crab1 + crab2
>>> combined_hemocyanin = combined_crab.hemocyanin_level
>>> expected_hemocyanin = 500
>>> condition = combined_hemocyanin == expected_hemocyanin
>>> assert condition, f'Expected hemocyanin level {expected_hemocyanin}, but got {combined_hemocyanin}.'
>>> if condition:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('7_magic_methods_horseshoe_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(combined_hemocyanin))
```
failure_message: Failed: `__add__()` does not add hemocyanin levels correctly.
success_message: Success: `__add__()` correctly adds hemocyanin levels!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'HorseshoeCrab-Magic-Methods-6'
>>> max_score = 2.0
>>> score = 0
>>> crab1 = HorseshoeCrab('Limulus polyphemus', 250)
>>> crab2 = HorseshoeCrab('Limulus polyphemus', 250)
>>> crab3 = HorseshoeCrab('Limulus polyphemus', 300)
>>> equality_check = crab1 == crab2 and (not crab1 == crab3)
>>> assert equality_check
>>> if equality_check:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('7_magic_methods_horseshoe_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(equality_check))
```
failure_message: Failed: `__eq__()` does not compare crabs correctly.
success_message: Success: `__eq__()` correctly compares crabs!
