```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_practice_midterm_q", "week_6", "PracticeMidterm", assignment_points = 86.0, assignment_tag = 'week6-PracticeMidterm')

# Initialize Otter
import otter
grader = otter.Notebook("1_practice_midterm_q.ipynb")
```

# â“ ENGR131: Introduction to Programming for Engineers
## Practice Midterm Exam

The practice midterm is designed to help you prepare for the actual midterm exam. It is designed to be similar in format and difficulty to the actual exam. The practice exam is graded and counts as 10% bonus points towards your  midterm exam grade.

For the multiple choice questions it is expected that you have access to a python interpreter. Do not go at it alone, use the resources available to you to help you answer the questions. Without the interpreter some of the questions may be difficult to answer.

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_practice_midterm_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_practice_midterm_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_practice_midterm_q import Question3
Question3().show()

```

## Engineering Component Tracker: Loops, Continue, Break, Dictionaries, and Lists

### Instructions:
You are managing the assembly of a robotic arm. Each component of the arm is represented as a dictionary, where the keys are component IDs (e.g., "A1") and the values are dictionaries containing component details. Your task is to write a function `process_components()` that processes a list of components using loops, `continue`, and `break`. The task involves multiple steps that can be graded individually.

## Engineering Component Tracker: Loops, Continue, Break, Dictionaries, and Lists

### Instructions:
You are managing the assembly of a robotic arm. Each component of the arm is represented as a dictionary, where the keys are component IDs (e.g., "A1") and the values are dictionaries containing component details. Your task is to write a function `process_components()` that processes a list of components using loops, `continue`, and `break`. The task involves multiple steps that can be graded individually.
 
 ### Requirements:
 1. **Input:** The function receives a list of dictionaries, each representing a component:
 
            | id  | status | weight |
            |-----|--------|--------|
            | A1  | ok     | 3.5    |
            | B2  | ok     | 2.0    |
            | C3  | error  | 4.1    |
            | D4  | ok     | 1.2    |
 2. **Part 1:** Count the components with status "ok" using a loop and return the count.
 3. **Part 2:** Skip components with status "error", and create a list of IDs for valid components.
 4. **Part 3:** If a component's weight exceeds 4.0, break the loop and return "Overweight component detected!".
 5. **Part 4:** Return a dictionary summarizing the count of "ok" components, the valid IDs list, and the break message (if triggered).
 
 ### Helpful Information:
 - Use a loop to iterate through the `components` list.
 - Use `continue` to skip certain iterations based on a condition.
 - Use `break` to exit the loop early when another condition is met.
 - Remember that you can access dictionary values using the syntax `dictionary[key]`.
 - Example methods:
   - `append`: Add an element to a list.

We provide detailed inline instructions for each part of the task. Please read the instructions carefully and complete each part accordingly.

```python
# define a function `process_components` that takes a dictionary of components as input
...
    
    # Initialize variables
    # create a variable `ok_count` and set it to 0
    # create a variable `valid_ids` and set it to an empty list
    # create a variable `break_message` and set it to None
    ...

    # Loop through components
    ...
    
        # Part 3: Check for overweight components
        # if the component's weight is greater than 4.0, set `break_message` to "Overweight component detected!" and break the loop
        ...
        
        # Part 2: Skip components with status "error"
        # if the component's status is "error", continue to the next iteration
        ...

        # Part 1: Count components with status "ok"
        # if the component's status is "ok", increment `ok_count` by 1 and append the component's id to `valid_ids`
        ...
        
    # Part 4: Return the results as a dictionary
    # return a dictionary with keys "ok_count", "valid_ids", and "break_message" and their corresponding values
    ...

# Example usage
# write the component dictionary based on the table provided
...

# Call the function with the components dictionary and store the result in a variable `result`
# You can print the results function to check if it's working correctly
...

```

```python
grader.check("loops-and-dictionaries-engineering")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week6-PracticeMidterm", "1_practice_midterm_q")
```

Tests for loops-and-dictionaries-engineering (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(22.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(22.0)
>>> log_variable('total-points', f'Reading-Week-X, 1_practice_midterm_q', 22.0)
>>> question_id = 'loops-and-dictionaries-engineering-1'
>>> max_score = 1.0
>>> score = 0
>>> function_exists = callable(process_components)
>>> assert function_exists, 'The function `process_components` must be defined and callable.'
>>> if (function_exists, 'The function `process_components` must be defined and callable.'):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_midterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(function_exists))
```
failure_message: Failed: The function `process_components` is not defined correctly.
success_message: Success: The function `process_components` exists and is callable!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import inspect
>>> import re
>>> question_id = 'loops-and-dictionaries-engineering-2'
>>> max_score = 2.0
>>> score = 0
>>> code = inspect.getsource(process_components)
>>> ok_count_init = re.search('ok_count\\s*=\\s*0', code)
>>> valid_ids_init = re.search('valid_ids\\s*=\\s*\\[\\]', code)
>>> break_message_init = re.search('break_message\\s*=\\s*None', code)
>>> initial_values = all([ok_count_init, valid_ids_init, break_message_init])
>>> assert initial_values, 'Expected variables `ok_count = 0`, `valid_ids = []`, and `break_message = None` in the function.'
>>> if (initial_values, 'Expected variables `ok_count = 0`, `valid_ids = []`, and `break_message = None` in the function.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_midterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(initial_values))
```
failure_message: Failed: One or more of `ok_count`, `valid_ids`, or `break_message` are not initialized properly.
success_message: Success: `ok_count`, `valid_ids`, and `break_message` are initialized correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import re
>>> import inspect
>>> question_id = 'loops-and-dictionaries-engineering-3'
>>> max_score = 2.0
>>> score = 0
>>> code = inspect.getsource(process_components)
>>> loop_check = re.search("\n    for\\s+          # 'for' keyword followed by one or more spaces\n    component\\s+    # 'component' variable followed by one or more spaces\n    in\\s+           # 'in' keyword followed by one or more spaces\n    components\\s*   # 'components' variable followed by zero or more spaces\n    :               # colon\n    ", code, re.IGNORECASE | re.VERBOSE)
>>> assert loop_check, 'Expected a loop with `for component in components:` to iterate over the dictionary.'
>>> if (loop_check, 'Expected a loop with `for component in components:` to iterate over the dictionary.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_midterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not correctly iterate over `components`.
success_message: Success: The function iterates over `components` correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-engineering-4'
>>> max_score = 3.0
>>> score = 0
>>> test_components = [{'id': 'X1', 'status': 'ok', 'weight': 3.0}, {'id': 'X2', 'status': 'ok', 'weight': 5.1}, {'id': 'X3', 'status': 'ok', 'weight': 2.5}]
>>> result = process_components(test_components)
>>> expected_break_message = 'Overweight component detected!'
>>> break_message = result['break_message']
>>> assert break_message == expected_break_message, f"Expected `break_message` to be '{expected_break_message}', but got '{break_message}'."
>>> if (break_message == expected_break_message, f"Expected `break_message` to be '{expected_break_message}', but got '{break_message}'."):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_midterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(break_message))
```
failure_message: Failed: Overweight components are not handled correctly.
success_message: Success: The function correctly detects overweight components and stops processing!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-engineering-5'
>>> max_score = 3.0
>>> score = 0
>>> test_components = [{'id': 'Y1', 'status': 'ok', 'weight': 2.0}, {'id': 'Y2', 'status': 'error', 'weight': 3.0}, {'id': 'Y3', 'status': 'ok', 'weight': 1.5}]
>>> result = process_components(test_components)
>>> expected_ok_count = 2
>>> expected_valid_ids = ['Y1', 'Y3']
>>> ok_count = result['ok_count']
>>> valid_ids = result['valid_ids']
>>> assert result['ok_count'] == expected_ok_count, f'Expected `ok_count` to be {expected_ok_count}, but got {result['ok_count']}.'
>>> assert result['valid_ids'] == expected_valid_ids, f'Expected `valid_ids` to be {expected_valid_ids}, but got {result['valid_ids']}.'
>>> if (result['ok_count'] == expected_ok_count, f'Expected `ok_count` to be {expected_ok_count}, but got {result['ok_count']}.') and (result['valid_ids'] == expected_valid_ids, f'Expected `valid_ids` to be {expected_valid_ids}, but got {result['valid_ids']}.'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_midterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(ok_count))
>>> responses = update_responses(question_id, str(valid_ids))
```
failure_message: Failed: The function does not skip `error` components correctly.
success_message: Success: The function correctly skips components with `status='error'`!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-engineering-6'
>>> max_score = 3.0
>>> score = 0
>>> test_components = [{'id': 'Z1', 'status': 'ok', 'weight': 2.5}, {'id': 'Z2', 'status': 'ok', 'weight': 3.8}, {'id': 'Z3', 'status': 'ok', 'weight': 2.1}]
>>> result = process_components(test_components)
>>> expected_ok_count = 3
>>> expected_valid_ids = ['Z1', 'Z2', 'Z3']
>>> ok_count = result['ok_count']
>>> valid_ids = result['valid_ids']
>>> assert result['ok_count'] == expected_ok_count, f'Expected `ok_count` to be {expected_ok_count}, but got {result['ok_count']}.'
>>> assert result['valid_ids'] == expected_valid_ids, f'Expected `valid_ids` to be {expected_valid_ids}, but got {result['valid_ids']}.'
>>> if (result['ok_count'] == expected_ok_count, f'Expected `ok_count` to be {expected_ok_count}, but got {result['ok_count']}.') and (result['valid_ids'] == expected_valid_ids, f'Expected `valid_ids` to be {expected_valid_ids}, but got {result['valid_ids']}.'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_midterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(ok_count))
>>> responses = update_responses(question_id, str(valid_ids))
```
failure_message: Failed: The function does not correctly count `ok` components or append their IDs.
success_message: Success: The function correctly counts `ok` components and stores their IDs!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-engineering-7'
>>> max_score = 2.0
>>> score = 0
>>> test_components = [{'id': 'A1', 'status': 'ok', 'weight': 2.0}]
>>> result = process_components(test_components)
>>> expected_keys = {'ok_count', 'valid_ids', 'break_message'}
>>> result_keys = set(result.keys())
>>> assert result_keys == expected_keys, f'Expected keys {expected_keys}, but got {result_keys}.'
>>> if (result_keys == expected_keys, f'Expected keys {expected_keys}, but got {result_keys}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_midterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(result_keys))
```
failure_message: Failed: The function does not return a dictionary with the expected keys.
success_message: Success: The function returns a correctly structured dictionary!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-engineering-8'
>>> max_score = 2.0
>>> score = 0
>>> test_components = []
>>> result = process_components(test_components)
>>> ok_count = result['ok_count']
>>> valid_ids = result['valid_ids']
>>> expected_result = {'ok_count': 0, 'valid_ids': [], 'break_message': None}
>>> assert result == expected_result, f'Expected result {expected_result}, but got {result}.'
>>> if (result == expected_result, f'Expected result {expected_result}, but got {result}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_midterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(result))
```
failure_message: Failed: The function does not handle an empty list correctly.
success_message: Success: The function correctly handles an empty components list!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-engineering-9'
>>> max_score = 2.0
>>> score = 0
>>> test_components = [{'id': 'DUP1', 'status': 'ok', 'weight': 2.0}, {'id': 'DUP1', 'status': 'ok', 'weight': 3.0}, {'id': 'DUP2', 'status': 'ok', 'weight': 1.5}]
>>> result = process_components(test_components)
>>> valid_ids = result['valid_ids']
>>> exec(base64.b64decode('ZXhwZWN0ZWRfdmFsaWRfaWRzID0gWyJEVVAxIiwgIkRVUDEiLCAiRFVQMiJdICAjIElEcyBzaG91bGQgYmUgYXBwZW5kZWQgYXMgdGhleSBhcHBlYXI=').decode())
>>> assert result['valid_ids'] == expected_valid_ids, f'Expected `valid_ids` to be {expected_valid_ids}, but got {result['valid_ids']}.'
>>> if (result['valid_ids'] == expected_valid_ids, f'Expected `valid_ids` to be {expected_valid_ids}, but got {result['valid_ids']}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_midterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(valid_ids))
```
failure_message: Failed: The function does not handle duplicate component IDs correctly.
success_message: Success: The function correctly processes duplicate component IDs!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-engineering-10'
>>> max_score = 2.0
>>> score = 0
>>> test_components = [{'id': 'W1', 'status': 'ok', 'weight': 4.0}, {'id': 'W2', 'status': 'ok', 'weight': 3.9}]
>>> result = process_components(test_components)
>>> ok_count = result['ok_count']
>>> valid_ids = result['valid_ids']
>>> expected_ok_count = 2
>>> expected_valid_ids = ['W1', 'W2']
>>> assert result['ok_count'] == expected_ok_count, f'Expected `ok_count` to be {expected_ok_count}, but got {result['ok_count']}.'
>>> assert result['valid_ids'] == expected_valid_ids, f'Expected `valid_ids` to be {expected_valid_ids}, but got {result['valid_ids']}.'
>>> if (result['ok_count'] == expected_ok_count, f'Expected `ok_count` to be {expected_ok_count}, but got {result['ok_count']}.') and (result['valid_ids'] == expected_valid_ids, f'Expected `valid_ids` to be {expected_valid_ids}, but got {result['valid_ids']}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_midterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(ok_count))
>>> responses = update_responses(question_id, str(valid_ids))
```
failure_message: Failed: The function incorrectly excludes components at exactly 4.0 weight.
success_message: Success: The function correctly allows components at the weight limit!
