```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_lab_scrabble", "week_2", "lab", assignment_points = 14.0, assignment_tag = 'week2-lab')

# Initialize Otter
import otter
grader = otter.Notebook("1_lab_scrabble.ipynb")
```

# 🧪 🎮 Lab 2: Scrabble Game Calculator

In this laboratory assignment, you will develop a scoring system for competitive word games while learning essential functional programming concepts including map and reduce functions. The scoring system you create will be similar to those used in professional tournament play.

## Professional Scrabble Gaming Context

Professional players utilize sophisticated scoring systems that require rapid calculation while considering multiple factors. The 2004 documentary "Word Wars" demonstrates how championship players have memorized over 120,000 playable words and can calculate optimal scoring positions instantaneously. Your task will be to develop the foundational components of such a scoring system.

<iframe width="1238" height="696" src="https://www.youtube.com/embed/kL4UfA6wUuw" title="Word Wars - Official Trailer" frameborder="0" allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share" referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>

## Scoring System

The scoring system assigns point values to letters based on their frequency in English language usage. The official scoring system is structured as follows:

| Points | Letters | Strategic Value |
|--------|----------|----------------|
| 1 point | A, E, I, O, U, L, N, S, T, R | Foundation letters for word construction |
| 2 points | D, G | Valuable consonants for word building |
| 3 points | B, C, M, P | Strong word initiators |
| 4 points | F, H, V, W, Y | High-value consonants |
| 5 points | K | Strategic positioning letter |
| 8 points | J, X | Premium scoring opportunities |
| 10 points | Q, Z | Maximum point potential |

## Task 1: Letter score

Your first task is to complete a function which takes in a letter and outputs the score for that letter. A dictionary can be used to store the information from the scoring table above.   

Write python code to do the following:

* Inside the provided `letter_score` function, use the provided scores dictionary `S` to determine the score for the provided letter called `letter`.
* Store the score in a variable called `score`.

Note: you may assume that the input letter will always be uppercase.

Your code replaces the prompt:  `...`

```python
# this line creates a function called letter_score with one argument: letter
def letter_score(letter):
    # create the scoring dictionary of letter scores
    S = {
        "A" : 1,
        "B" : 3,
        "C" : 3,
        "D" : 2,
        "E" : 1,
        "F" : 4,
        "G" : 2,
        "H" : 4,
        "I" : 1,
        "J" : 8,
        "K" : 5,
        "L" : 1,
        "M" : 3,
        "N" : 1,
        "O" : 1,
        "P" : 3,
        "Q" : 10,
        "R" : 1,
        "S" : 1,
        "T" : 1,
        "U" : 1,
        "V" : 4,
        "W" : 4,
        "X" : 8,
        "Y" : 4,
        "Z" : 10,
    }
    ...
    # this line outputs the score from the function
    return score

# this runs the letter score function with input "Q" and prints its output
print(letter_score("Q"))
```

```python
grader.check("task1-letter-score")
```

## Task 2: Letter scores

Now that you have a function that computes the score for a letter, you can use Python's built-in `map` function to apply it to each letter in a word to get a list of letter scores. 

First you will need to convert the string `word` into a list of the individual letters. This can be achieved with `list(word)`. 

Next, use the `map` function to apply the `letter_score` function to each letter in the list. Recall how the `map` function is used: If you have a function `f` and a list `L = [a,b,c]`, you can do `map(f,L)` which results in the mapped list `[f(a), f(b), f(c)]`. 

Write python code to do the following:

* Inside the provided `letter_scores` function, first convert the argument string `word` into a list.
* Map the `letter_score` function onto the list to get a list of letter scores. 
* Store the final list of letter scores in a variable called `score_list`.

Your code replaces the prompt:  `...`

```python
# this line creates a function called letter_scores with one argument: word
def letter_scores(word):
    ...
    # this line outputs the score list from the function
    return list(score_list)

# this runs the letter_scores function with input "EXAMPLE" and prints out the resulting list of letter scores
print(letter_scores("EXAMPLE"))
```

```python
grader.check("task2-letter-scores")
```

**Task 3: Word score** 

The last thing you need to do is sum the list of indivudal letters scores to get the total point value for the entire word.

To do this you should use the `sum` function. This adds all the elements of a list and returns the sum.

Write python code to do the following:

* Inside the provided `word_score` function, first call the `letter_scores` function on the input `word` to get the list of individual letter scores.
* Next use the `sum` function on the list and `return` the final sum. 

Your code replaces the prompt:  `...`

```python
def word_score(word):
    ...

# use this to check your results
word_score("ZEBRA")
```

```python
grader.check("task3-word-score")
```

## Testing Your Implementation

Test your functions with various words to ensure accurate scoring:

```python
# Test cases
test_words = ["PYTHON", "ALGORITHM", "CODE", "QUIZ"]

for word in test_words:
    print(
        f"{word}: {word_score(word)} points (Basic), {word_score(word)} points"
    )
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week2-lab", "1_lab_scrabble")
```

Tests for task1-letter-score (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(5.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(14.0)
>>> log_variable('total-points', f'Reading-Week-X, 1_lab_scrabble', 14.0)
>>> question_id = 'task1-letter-score-1'
>>> max_score = 1.0
>>> score = 0
>>> assert letter_score('A') == 1, 'letter score for A is incorrect'
>>> if letter_score('A') == 1:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_scrabble', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: letter_score("A") is incorrect.
success_message: letter_score("A") is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task1-letter-score-2'
>>> max_score = 1.0
>>> score = 0
>>> assert letter_score('X') == 8, "letter_score('X') is incorrect"
>>> if letter_score('X') == 8:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_scrabble', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: letter_score("X") is incorrect.
success_message: letter_score("X") is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task1-letter-score-3'
>>> max_score = 1.0
>>> score = 0
>>> assert letter_score('Z') == 10, "letter_score('Z') is incorrect."
>>> if letter_score('Z') == 10:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_scrabble', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: letter_score("Z") is incorrect.
success_message: letter_score("Z") is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task1-letter-score-4'
>>> max_score = 1.0
>>> score = 0
>>> assert letter_score('D') == 2, "letter_score('D') should return 2"
>>> if letter_score('D') == 2:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_scrabble', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: letter_score("D") is incorrect.
success_message: letter_score("D") is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task1-letter-score-5'
>>> max_score = 1.0
>>> score = 0
>>> assert letter_score('C') == 3, "letter_score('C') should return 3"
>>> if letter_score('C') == 3:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_scrabble', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: letter_score("C") is incorrect.
success_message: letter_score("C") is correct.


Tests for task2-letter-scores (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(5.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(14.0)
>>> log_variable('total-points', f'Reading-Week-X, 1_lab_scrabble', 14.0)
>>> question_id = 'task2-letter-scores-1'
>>> max_score = 1.0
>>> score = 0
>>> assert letter_scores('AND') == [1, 1, 2], "the letter scores for 'AND' should be [1,1,2]"
>>> if letter_scores('AND') == [1, 1, 2]:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_scrabble', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: letter_scores("AND") is incorrect.
success_message: letter_scores("AND") is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task2-letter-scores-2'
>>> max_score = 1.0
>>> score = 0
>>> assert letter_scores('HYPOTHESIS') == [4, 4, 3, 1, 1, 4, 1, 1, 1, 1], 'The letter scores for HYPOTHESIS should return [4, 4, 3, 1, 1, 4, 1, 1, 1, 1]'
>>> if letter_scores('HYPOTHESIS') == [4, 4, 3, 1, 1, 4, 1, 1, 1, 1]:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_scrabble', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: letter_scores("HYPOTHESIS") is incorrect.
success_message: letter_scores("HYPOTHESIS") is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task2-letter-scores-3'
>>> max_score = 1.0
>>> score = 0
>>> assert letter_scores('CHEMISTRY') == [3, 4, 1, 3, 1, 1, 1, 1, 4], 'The letter scores for CHEMISTRY should return [3, 4, 1, 3, 1, 1, 1, 1, 4]'
>>> if letter_scores('CHEMISTRY') == [3, 4, 1, 3, 1, 1, 1, 1, 4]:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_scrabble', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: letter_scores("CHEMISTRY") is incorrect.
success_message: letter_scores("CHEMISTRY") is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task2-letter-scores-4'
>>> max_score = 1.0
>>> score = 0
>>> assert letter_scores('DOOR') == [2, 1, 1, 1], "letter_scores('DOOR') should return [2,1,1,1]"
>>> if letter_scores('DOOR') == [2, 1, 1, 1]:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_scrabble', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: letter_scores("DOOR") is incorrect.
success_message: letter_scores("DOOR") is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task2-letter-scores-5'
>>> max_score = 1.0
>>> score = 0
>>> assert letter_scores('FIVE') == [4, 1, 4, 1], 'The wrong letter score is being calculated for the word FIVE'
>>> if letter_scores('FIVE') == [4, 1, 4, 1]:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_scrabble', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: letter_scores("FIVE") is incorrect.
success_message: letter_scores("FIVE") is correct.


Tests for task3-word-score (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(4.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(14.0)
>>> log_variable('total-points', f'Reading-Week-X, 1_lab_scrabble', 14.0)
>>> question_id = 'task3-word-score-1'
>>> max_score = 1.0
>>> score = 0
>>> assert word_score('ZEBRA') == 16
>>> if word_score('ZEBRA') == 16:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_scrabble', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: word_score("ZEBRA") is incorrect.
success_message: word_score("ZEBRA") is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task3-word-score-2'
>>> max_score = 1.0
>>> score = 0
>>> assert word_score('MACHINE') == 14
>>> if word_score('MACHINE') == 14:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_scrabble', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: word_score("MACHINE") is incorrect.
success_message: word_score("MACHINE") is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task3-word-score-3'
>>> max_score = 1.0
>>> score = 0
>>> assert word_score('CORRUPT') == 11
>>> if word_score('CORRUPT') == 11:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_scrabble', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: word_score("CORRUPT") is incorrect.
success_message: word_score("CORRUPT") is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task3-word-score-4'
>>> max_score = 1.0
>>> score = 0
>>> assert word_score('MOVIE') == 10
>>> if word_score('MOVIE') == 10:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_scrabble', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: word_score("MOVIE") is incorrect.
success_message: word_score("MOVIE") is correct.
