```python
import os
os.environ["JUPYTERHUB_USER"] = "JCA"

# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("3_southpark_q", "week_4", "readings", assignment_points = 15.0, assignment_tag = 'week4-readings')

# Initialize Otter
import otter
grader = otter.Notebook("3_southpark_q.ipynb")
```

# â“ ðŸžï¸ South Park Adventures ðŸ§¢ðŸŒ²

![southpark](./assets/figures/south-park1.jpg)

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._3_southpark_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._3_southpark_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._3_southpark_q import Question3
Question3().show()

```

## Cartman's Decision

![cartman](./assets/figures/Eric_Cartman.webp)

### Instructions:
Write a Python program where Cartman decides whether to skip school or go to class. Use the following conditions:
- Cartman skips school if it's snowing OR it's a weekend.
- Otherwise, he goes to class.

Your solution must include:
1. A function `cartman_decision(is_snowing, is_weekend)` (we have provided the function for you) that:
   - Takes two boolean arguments `is_snowing` and `is_weekend`.
   - Returns `"Skip School"` if Cartman skips school.
   - Returns `"Go to Class"` otherwise.
2. Test cases to validate the correctness of your implementation. We have provided this for you as well.

```python
# Define the function
def cartman_decision(is_snowing, is_weekend):
    # Use if-else statements to determine Cartman's decision
    ...

# Test the function
print(cartman_decision(True, False))  # Expect "Skip School"
print(cartman_decision(False, True))  # Expect "Skip School"
print(cartman_decision(False, False))  # Expect "Go to Class"
```

```python
grader.check("SouthPark-CartmanDecision")
```

## Side Note ðŸ“

Prof. Agar got suspended from middle school for wearing a Mr. Hankey shirt to class. ðŸ¤£

![](./assets/figures/howdy-ho-mr-hankey.gif)

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week4-readings", "3_southpark_q")
```

Tests for SouthPark-CartmanDecision (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(6.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(6.0)
>>> log_variable('total-points', f'Reading-Week-X, 3_southpark_q', 6.0)
>>> question_id = 'SouthPark-CartmanDecision-1'
>>> max_score = 2.0
>>> score = 0
>>> result = cartman_decision(True, False)
>>> assert result == 'Skip School', f"Expected 'Skip School', but got {result}"
>>> if result == 'Skip School':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_southpark_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(result))
```
failure_message: Failed: 'cartman_decision' should return 'Skip School' when it's snowing.
success_message: Success: 'cartman_decision' correctly returns 'Skip School' when it's snowing.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'SouthPark-CartmanDecision-2'
>>> max_score = 2.0
>>> score = 0
>>> result = cartman_decision(False, True)
>>> assert result == 'Skip School', f"Expected 'Skip School', but got {result}"
>>> if result == 'Skip School':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_southpark_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(result))
```
failure_message: Failed: 'cartman_decision' should return 'Skip School' on a weekend.
success_message: Success: 'cartman_decision' correctly returns 'Skip School' on a weekend.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'SouthPark-CartmanDecision-3'
>>> max_score = 2.0
>>> score = 0
>>> result = cartman_decision(False, False)
>>> assert result == 'Go to Class', f"Expected 'Go to Class', but got {result}"
>>> if result == 'Go to Class':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_southpark_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(result))
```
failure_message: Failed: 'cartman_decision' should return 'Go to Class' when neither condition is met.
success_message: Success: 'cartman_decision' correctly returns 'Go to Class' otherwise.
