```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_practicemidterm_q", "week_6", "practicemidterm", assignment_points = 225.0, assignment_tag = 'week6-practicemidterm')
 
# Initialize Otter
import otter
grader = otter.Notebook("1_practicemidterm_q.ipynb")
```

# ❓ ENGR131: Introduction to Programming for Engineers
## Practice Midterm Exam

The practice midterm is designed to help you prepare for the actual midterm exam. It is designed to be similar in format and difficulty to the actual exam. The practice exam is graded and counts as 10% bonus points towards your  midterm exam grade.

For the multiple choice questions it is expected that you have access to a python interpreter. Do not go at it alone, use the resources available to you to help you answer the questions. Without the interpreter some of the questions may be difficult to answer.

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_practicemidterm_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_practicemidterm_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_practicemidterm_q import Question3
Question3().show()

```

## Question 1 (Points: 22.0): Engineering Component Tracker: Loops, Continue, Break, Dictionaries, and Lists

### Instructions:
You are managing the assembly of a robotic arm. Each component of the arm is represented as a dictionary, where the keys are component IDs (e.g., "A1") and the values are dictionaries containing component details. Your task is to write a function `process_components()` that processes a list of components using loops, `continue`, and `break`. 
 
 ### Requirements:
 1. **Input:** The function receives a list of dictionaries, each representing a component:
 
            | id  | status | weight |
            |-----|--------|--------|
            | A1  | ok     | 3.5    |
            | B2  | ok     | 2.0    |
            | C3  | error  | 4.1    |
            | D4  | ok     | 1.2    |
 2. **Part 1:** Count the components with status "ok" using a loop and return the count.
 3. **Part 2:** Skip components with status "error", and create a list of IDs for valid components.
 4. **Part 3:** If a component's weight exceeds 4.0, break the loop and return "Overweight component detected!".
 5. **Part 4:** Return a dictionary summarizing the count of "ok" components, the valid IDs list, and the break message (if triggered).
 
 ### Helpful Information:
 - Use a loop to iterate through the `components` list.
 - Use `continue` to skip certain iterations based on a condition.
 - Use `break` to exit the loop early when another condition is met.
 - Remember that you can access dictionary values using the syntax `dictionary[key]`.
 - Example methods:
   - `append`: Add an element to a list.

We provide detailed inline instructions for each part of the task. Please read the instructions carefully and complete each part accordingly.

```python
# write the components list of dictionaries based on the table provided, assign it to a variable `components`
...

# define a function `process_components` that takes a dictionary of components as input
...
    
    # Initialize variables
    # create a variable `ok_count` and set it to 0
    # create a variable `valid_ids` and set it to an empty list
    # create a variable `break_message` and set it to None
    ...

    # Loop through components
    ...
    
        # Part 3: Check for overweight components
        # if the component's weight is greater than 4.0, set `break_message` to "Overweight component detected!" and break the loop
        ...
        
        # Part 2: Skip components with status "error"
        # if the component's status is "error", continue to the next iteration
        ...

        # Part 1: Count components with status "ok"
        # if the component's status is "ok", increment `ok_count` by 1 and append the component's id to `valid_ids`
        ...
        
    # Part 4: Return the results as a dictionary
    # return a dictionary with keys "ok_count", "valid_ids", and "break_message" and their corresponding values
    ...

# Call the function with the components dictionary and store the result in a variable `result`
# You can print the results function to check if it's working correctly
...
```

```python
grader.check("loops-and-dictionaries-engineering")
```

## Question 2 (Points: 16.0): Symbolic Computation with SymPy

### Instructions:
Using SymPy, a package for symbolic solving in Python. Your task is to define and visualize an equation symbolically.

### Requirements:
1. **Part 1:** Define symbolic variables \( x \) and \( y \).
2. **Part 2:** Construct the equation:

$$
(x^2 + y^2 - 1)^3 - x^2 y^3 = 0
$$

3. **Part 3:** Use `sympy.plot_implicit` to visualize the shape. We do not expect you to know how to use this function. You should use one of the several methods shown in class to read the docstring. 
4. **Part 4:** Return the equation and the plot object.

### Helpful Information:
- Make sure you import from `sympy` `symbols`, `Eq`, and from `sympy.plotting` `plot_implicit`.
- Use `symbols` from SymPy to define \( x \) and \( y \).
- Use `Eq` to represent the equation.
- Use `plot_implicit` to visualize the shape.

```python
# Import required functions from sympy
...

# Define symbolic variables
...

# Construct the shape equation
# save to the variable eq
...

# Plot the equation
# x and y limits should be between -1.5 and 1.5
# The syntax for plot_implicit is plot_implicit(equation, (x, x_min, x_max), (y, y_min, y_max))
# save the plot to the variable plot
...

# do not modify below this line
# this is used for grading
def get_equation():
    return eq, plot

plot.show()
```

```python
grader.check("symbolic-heart-curve")
```

## Question 3 (Points: 14.0): Debugging Python Functions: Syntax, Runtime, and Semantic Errors

### Instructions:
You are provided with the following code which is full of bugs. Your task is to debug the code and fix all the errors. The function `square_positive_numbers` is intended to take a list of numbers as input and return a new list where each positive number is squared, the negative numbers are excluded.

```python
def square_positive_numbers(numbers):
    new_list = []
    for num in numbers:
        if num =< 0 
            continue
        
        newLsit.append(num ** 2)  
    
    print(new_list) 

# Example usage:
numbers = [1, -2, 3, -4, 5]
print(square_positive_numbers(numbers))  # This should print [1, 9, 25]
```

### Requirements:
1. **Input:** The function receives a list of numbers.
2. **Expected Behavior:** The function should return a new list where each number is squared, but only if it is positive.
3. **Possible Errors to Fix:**
    - Syntax Errors (incorrect operators, missing colons, indentation issues).
    - Runtime Errors (incorrect function calls, undefined variables).
    - Semantic Errors (logical mistakes affecting correct behavior).
4. **Return Value:** The function should return the transformed list, not just print it.

```python
# Define a function `square_positive_numbers` that takes a list of numbers as input
...

numbers = [1, -2, 3, -4, 5]
result = square_positive_numbers(numbers)
print(result)
```

```python
grader.check("loops-and-errors-debugging")
```

## Question 4 (Points: 25.0): Numerical Computing with NumPy: Prime Number Checker 🔢

### Instructions:
You are working with numerical data and need to determine which numbers in a given NumPy array are prime. Your task is to write a function `find_primes()` that takes a NumPy array of integers and returns a boolean array indicating which elements are prime.
 
 ### Requirements:
 1. The function receives a NumPy array of positive integers, or a list of positive integers.
 2.  Implement an efficient prime-checking function using NumPy.
 4. The function should return a boolean NumPy array of the same shape as the input, where `True` represents a prime number and `False` represents a non-prime.
 
 ### Helpful Information:
 - A prime number is a natural number greater than **1** that has only two divisors: **1 and itself**.
 - The first few primes are: **2, 3, 5, 7, 11, 13, ...**
 - The function should handle edge cases like `0` and `1` correctly (they are not prime).

 We have provided inline instructions for each part of the task. Please read the instructions carefully and complete each part accordingly.

 Make sure to include docstring and comments in your code to explain the logic and purpose of the function.

```python
# import numpy use the common alias np
...

# Define the function `find_primes` which accepts and input arr 
...
    """
    ...
    """
    # Ensure input is a NumPy array, to do this convert the input to a NumPy array using `np.asarray`
    # Override the input variable `arr` with the NumPy array
    ...
    
    # Initialize a boolean array `is_prime` with the same shape as the input array
    # Set all values to `True` initially, this generally happens by default with ones_like
    # You can use the `np.ones_like` function to initialize the array, make sure to set the `dtype` to `bool`
    ...

    # Mark 0 and 1 as non-prime
    # Use indexing to set the values at indices less than 2 to `False`
    # hint, you need to build a boolean array that can index is_prime that finds all values in the array that are less than 2
    ...

    # The trick to finding primes is to mark multiples of each number as non-prime
    # Loop through numbers starting from 2 up to the square root of the maximum number in the array -- make sure to use the NumPy function which supports arrays
        # If you use the max() function it will not work for multidiemensional arrays that cannot be converted to a list inside python.
    # Use the `np.sqrt` function to calculate the square root
    # Use the `int` function to convert the result to an integer -- make sure to add 1 to the result
    # Use the `range` function to loop through the numbers
    # Write your for loop here
    ...
        # Create a boolean mask to identify multiples of the current number
        # These are the numbers that are divisible by the current number but not equal to the current number
        # You will want to use the `&`` operator which conducts an element-wise AND operation on two boolean arrays, or tuples.
        # Save the mask to a variable `mask`
        ...
        
        # Mark multiples as non-prime by setting the corresponding values in `is_prime` to `False`
        ...

    # Return the boolean array `is_prime`
    ...

# Example usage:
test_array = np.array([0, 1, 2, 3, 4, 5, 10, 11, 13, 15, 17])
prime_results = find_primes(test_array)
print(prime_results)
```

```python
grader.check("numpy-prime-check")
```

## Question 5 (Points: 24.0): String Processing with Sliding Window: Longest Unique Substring 🔠

### Instructions:
You are given a string `s` containing lowercase and uppercase English letters, digits, and special characters. Your task is to implement a function `length_of_longest_substring()` that returns the **length of the longest substring without repeating characters**.

This question is designed to test your understanding of algorithmic thinking. You can solve this problem using any approach you prefer, but we recommend using the sliding window technique for efficiency.

### Requirements:
1. **Input:** The function receives a string `s`.
2.  Ensure that the function correctly tracks character positions to prevent repetition.
3.  Return an integer representing the maximum length of a substring without repeating characters.

### Helpful Information:
- A **substring** is a contiguous sequence of characters within a string.
- You should might want to use a  **dictionary** to store the last seen index of characters.
- Use the **sliding window** approach to efficiently check substrings.
- Ideally the function should work well for very large strings.

### Viable logic

1. Define a dictionary to store the last seen index of characters. The key is the character, and the value is the index.
2. Store the max length of  all substrings found so far.
3. Store the start index, the left pointer of a window in the string. 
4. Iterate through the string, it is helpful to use the enumerate function to get the index and character.
5. If the character is in the dictionary, and the character's index is greater than or equal to the start index, update the start index to the character's index + 1.
6. Update the character's index in the dictionary, to the last seen index.
7. Use the max function to update the max length of all substrings found so far, this is either the current max value, or the length of the current substring.  Make sure you think about the edge case where the substring is the entire string, and how python indexing works.
8. Return the max length of all substrings after traversing the string.

The goal is to go through character by character, and keep track of the last seen index of each character. If a character is repeated then move the left pointer to the right by one so there are no repeating characters in the substring. If the distance between the left pointer and the current character is greater than the current max length, update the max length.

```python
# define your function `length_of_longest_substring`
...

# Example usage:
test_string = "abcabcbb"
result = length_of_longest_substring(test_string)
print(result)  # Expected output: 3
```

```python
grader.check("longest-unique-substring")
```

## Question 6 (Points: 60.0): Bonus Question: Cracking Agar's Password to Uncover a Secret Message 🕵️‍♂️

**Note:** This question is optional. You can choose to attempt this question if you want to earn bonus points. This question is designed to test your problem-solving skills and creativity -- it is well above the difficulty level of the regular questions and is not within the scope of the exam content. This is for students that are looking for an extra challenge.


## Instructions:

You have intercepted a message from Agar, concealing a deep dark secret about a former life as a pop star - at least of the background vocal kind. You know the message `encrypted_message.txt`  is encrypted using a simple XOR cipher, but you don't know the key. Your task is to write a function `crack_agar_password()` that decrypts the message and reveals the secret. You were able to watch Professor Agar in class as he typed in his password, you saw that the password was a series of letters - which you assume is a word in the english language - all of which are provided in `words.txt`, a single number, and a special character [!@#$%^&*()_+]. 

You make some good assumptions that the message is in English, and that the message might contain a link to a website. 

```python
...
```

```python
...

# Run the function
password = crack_agar_password()
```

```python
grader.check("Bonus-Question-Cracking the Code")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week6-practicemidterm", "1_practicemidterm_q")
```

Tests for Bonus-Question-Cracking the Code (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import requests
>>> import json
>>> max_question_points = str(60.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(161.0)
>>> log_variable('total-points', f'week6-practicemidterm, 1_practicemidterm_q', 161.0)
>>> question_id = 'Bonus-Question-Cracking the Code-1'
>>> max_score = 60.0
>>> score = 0
>>> url_base = 'https://engr-131-api.eastus.cloudapp.azure.com'
>>> url = f'{url_base}/check_pass'
>>> params = {'password': password}
>>> response = requests.get(url, params=params)
>>> if response.json() == True:
...     condition = True
... else:
...     condition = False
>>> assert condition
>>> if condition:
...     score = 60.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: You did not decrypt the message or find the password.
success_message: Success: You correctly decrypted the message and found the password!


Tests for longest-unique-substring (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(24.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(161.0)
>>> log_variable('total-points', f'week6-practicemidterm, 1_practicemidterm_q', 161.0)
>>> question_id = 'longest-unique-substring-1'
>>> max_score = 1.0
>>> score = 0
>>> function_exists = callable(length_of_longest_substring)
>>> assert function_exists, 'The function `length_of_longest_substring` must be defined and callable.'
>>> if (function_exists, 'The function `length_of_longest_substring` must be defined and callable.'):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(function_exists))
```
failure_message: Failed: The function `length_of_longest_substring` is not defined correctly.
success_message: Success: The function `length_of_longest_substring` exists and is callable!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'longest-unique-substring-2'
>>> max_score = 3.0
>>> score = 0
>>> test_input = 'abcabcbb'
>>> expected_output = 3
>>> basic_test_results = length_of_longest_substring(test_input)
>>> assert basic_test_results == expected_output, f'Expected {expected_output}, but got {basic_test_results}.'
>>> if (basic_test_results == expected_output, f'Expected {expected_output}, but got {basic_test_results}.'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(basic_test_results))
```
failure_message: Failed: The function does not correctly compute longest unique substrings.
success_message: Success: The function correctly identifies longest unique substrings!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'longest-unique-substring-3'
>>> max_score = 3.0
>>> score = 0
>>> test_input = 'abcdefg'
>>> expected_output = 7
>>> unique_test_results = length_of_longest_substring(test_input)
>>> assert unique_test_results == expected_output, f'Expected {expected_output}, but got {unique_test_results}.'
>>> if (unique_test_results == expected_output, f'Expected {expected_output}, but got {unique_test_results}.'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(unique_test_results))
```
failure_message: Failed: The function does not correctly handle a string with all unique characters.
success_message: Success: The function correctly handles all unique characters!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'longest-unique-substring-4'
>>> max_score = 3.0
>>> score = 0
>>> test_input = 'bbbbb'
>>> expected_output = 1
>>> repeated_test_results = length_of_longest_substring(test_input)
>>> assert repeated_test_results == expected_output, f'Expected {expected_output}, but got {repeated_test_results}.'
>>> if (repeated_test_results == expected_output, f'Expected {expected_output}, but got {repeated_test_results}.'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(repeated_test_results))
```
failure_message: Failed: The function does not correctly identify longest substrings when characters repeat.
success_message: Success: The function correctly handles repeated characters!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'longest-unique-substring-5'
>>> max_score = 3.0
>>> score = 0
>>> test_input = 'aAabcABC'
>>> expected_output = 6
>>> mixed_case_test_results = length_of_longest_substring(test_input)
>>> assert mixed_case_test_results == expected_output, f'Expected {expected_output}, but got {mixed_case_test_results}.'
>>> if (mixed_case_test_results == expected_output, f'Expected {expected_output}, but got {mixed_case_test_results}.'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(mixed_case_test_results))
```
failure_message: Failed: The function does not correctly handle mixed uppercase and lowercase characters.
success_message: Success: The function correctly handles mixed-case input!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'longest-unique-substring-6'
>>> max_score = 3.0
>>> score = 0
>>> test_input = 'ab!@#ab!@#'
>>> expected_output = 5
>>> special_char_test_results = length_of_longest_substring(test_input)
>>> assert special_char_test_results == expected_output, f'Expected {expected_output}, but got {special_char_test_results}.'
>>> if (special_char_test_results == expected_output, f'Expected {expected_output}, but got {special_char_test_results}.'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(special_char_test_results))
```
failure_message: Failed: The function does not correctly identify substrings with special characters.
success_message: Success: The function correctly handles strings with special characters!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'longest-unique-substring-7'
>>> max_score = 3.0
>>> score = 0
>>> test_input = 'abcdefghijklmnopqrstuvwxyz' * 1000
>>> expected_output = 26
>>> long_test_results = length_of_longest_substring(test_input)
>>> assert long_test_results == expected_output, f'Expected {expected_output}, but got {long_test_results}.'
>>> if (long_test_results == expected_output, f'Expected {expected_output}, but got {long_test_results}.'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(long_test_results))
```
failure_message: Failed: The function is too slow for long inputs.
success_message: Success: The function correctly handles long strings efficiently!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'longest-unique-substring-8'
>>> max_score = 2.0
>>> score = 0
>>> test_input = ''
>>> expected_output = 0
>>> empty_string_results = length_of_longest_substring(test_input)
>>> assert empty_string_results == expected_output, f'Expected {expected_output}, but got {empty_string_results}.'
>>> if (empty_string_results == expected_output, f'Expected {expected_output}, but got {empty_string_results}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(empty_string_results))
```
failure_message: Failed: The function does not handle empty strings correctly.
success_message: Success: The function correctly handles an empty string!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'longest-unique-substring-9'
>>> max_score = 3.0
>>> score = 0
>>> test_input = 'z'
>>> expected_output = 1
>>> single_char_results = length_of_longest_substring(test_input)
>>> assert single_char_results == expected_output, f'Expected {expected_output}, but got {single_char_results}.'
>>> if (single_char_results == expected_output, f'Expected {expected_output}, but got {single_char_results}.'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(single_char_results))
```
failure_message: Failed: The function does not correctly handle single-character strings.
success_message: Success: The function correctly handles single-character strings!


Tests for loops-and-dictionaries-engineering (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(22.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(161.0)
>>> log_variable('total-points', f'week6-practicemidterm, 1_practicemidterm_q', 161.0)
>>> question_id = 'loops-and-dictionaries-engineering-1'
>>> max_score = 1.0
>>> score = 0
>>> function_exists = callable(process_components)
>>> assert function_exists, 'The function `process_components` must be defined and callable.'
>>> if (function_exists, 'The function `process_components` must be defined and callable.'):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(function_exists))
```
failure_message: Failed: The function `process_components` is not defined correctly.
success_message: Success: The function `process_components` exists and is callable!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import inspect
>>> import re
>>> question_id = 'loops-and-dictionaries-engineering-2'
>>> max_score = 2.0
>>> score = 0
>>> code = inspect.getsource(process_components)
>>> ok_count_init = re.search('ok_count\\s*=\\s*0', code)
>>> valid_ids_init = re.search('valid_ids\\s*=\\s*\\[\\]', code)
>>> break_message_init = re.search('break_message\\s*=\\s*None', code)
>>> initial_values = all([ok_count_init, valid_ids_init, break_message_init])
>>> assert initial_values, 'Expected variables `ok_count = 0`, `valid_ids = []`, and `break_message = None` in the function.'
>>> if (initial_values, 'Expected variables `ok_count = 0`, `valid_ids = []`, and `break_message = None` in the function.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(initial_values))
```
failure_message: Failed: One or more of `ok_count`, `valid_ids`, or `break_message` are not initialized properly.
success_message: Success: `ok_count`, `valid_ids`, and `break_message` are initialized correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import re
>>> import inspect
>>> question_id = 'loops-and-dictionaries-engineering-3'
>>> max_score = 2.0
>>> score = 0
>>> code = inspect.getsource(process_components)
>>> loop_check = re.search("\n    for\\s+          # 'for' keyword followed by one or more spaces\n    component\\s+    # 'component' variable followed by one or more spaces\n    in\\s+           # 'in' keyword followed by one or more spaces\n    components\\s*   # 'components' variable followed by zero or more spaces\n    :               # colon\n    ", code, re.IGNORECASE | re.VERBOSE)
>>> assert loop_check, 'Expected a loop with `for component in components:` to iterate over the dictionary.'
>>> if (loop_check, 'Expected a loop with `for component in components:` to iterate over the dictionary.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not correctly iterate over `components`.
success_message: Success: The function iterates over `components` correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-engineering-4'
>>> max_score = 3.0
>>> score = 0
>>> test_components = [{'id': 'X1', 'status': 'ok', 'weight': 3.0}, {'id': 'X2', 'status': 'ok', 'weight': 5.1}, {'id': 'X3', 'status': 'ok', 'weight': 2.5}]
>>> result = process_components(test_components)
>>> exec(base64.b64decode('ZXhwZWN0ZWRfYnJlYWtfbWVzc2FnZSA9ICJPdmVyd2VpZ2h0IGNvbXBvbmVudCBkZXRlY3RlZCEi').decode())
>>> exec(base64.b64decode('YnJlYWtfbWVzc2FnZSA9IHJlc3VsdFsiYnJlYWtfbWVzc2FnZSJd').decode())
>>> assert break_message == expected_break_message
>>> if break_message == expected_break_message:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(break_message))
```
failure_message: Failed: Overweight components are not handled correctly.
success_message: Success: The function correctly detects overweight components and stops processing!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-engineering-5'
>>> max_score = 3.0
>>> score = 0
>>> test_components = [{'id': 'Y1', 'status': 'ok', 'weight': 2.0}, {'id': 'Y2', 'status': 'error', 'weight': 3.0}, {'id': 'Y3', 'status': 'ok', 'weight': 1.5}]
>>> result = process_components(test_components)
>>> exec(base64.b64decode('ZXhwZWN0ZWRfb2tfY291bnQgPSAy').decode())
>>> exec(base64.b64decode('ZXhwZWN0ZWRfdmFsaWRfaWRzID0gWyJZMSIsICJZMyJd').decode())
>>> ok_count = result['ok_count']
>>> valid_ids = result['valid_ids']
>>> assert result['ok_count'] == expected_ok_count
>>> assert result['valid_ids'] == expected_valid_ids
>>> if result['ok_count'] == expected_ok_count and result['valid_ids'] == expected_valid_ids:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(ok_count))
>>> responses = update_responses(question_id, str(valid_ids))
```
failure_message: Failed: The function does not skip `error` components correctly.
success_message: Success: The function correctly skips components with `status='error'`!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-engineering-6'
>>> max_score = 3.0
>>> score = 0
>>> test_components = [{'id': 'Z1', 'status': 'ok', 'weight': 2.5}, {'id': 'Z2', 'status': 'ok', 'weight': 3.8}, {'id': 'Z3', 'status': 'ok', 'weight': 2.1}]
>>> result = process_components(test_components)
>>> exec(base64.b64decode('ZXhwZWN0ZWRfb2tfY291bnQgPSAz').decode())
>>> exec(base64.b64decode('ZXhwZWN0ZWRfdmFsaWRfaWRzID0gWyJaMSIsICJaMiIsICJaMyJd').decode())
>>> ok_count = result['ok_count']
>>> valid_ids = result['valid_ids']
>>> assert result['ok_count'] == expected_ok_count
>>> assert result['valid_ids'] == expected_valid_ids
>>> if result['ok_count'] == expected_ok_count and result['valid_ids'] == expected_valid_ids:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(ok_count))
>>> responses = update_responses(question_id, str(valid_ids))
```
failure_message: Failed: The function does not correctly count `ok` components or append their IDs.
success_message: Success: The function correctly counts `ok` components and stores their IDs!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-engineering-7'
>>> max_score = 2.0
>>> score = 0
>>> test_components = [{'id': 'A1', 'status': 'ok', 'weight': 2.0}]
>>> result = process_components(test_components)
>>> exec(base64.b64decode('ZXhwZWN0ZWRfa2V5cyA9IHsib2tfY291bnQiLCAidmFsaWRfaWRzIiwgImJyZWFrX21lc3NhZ2UifQ==').decode())
>>> result_keys = set(result.keys())
>>> assert result_keys == expected_keys
>>> if result_keys == expected_keys:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(result_keys))
```
failure_message: Failed: The function does not return a dictionary with the expected keys.
success_message: Success: The function returns a correctly structured dictionary!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-engineering-8'
>>> max_score = 2.0
>>> score = 0
>>> test_components = []
>>> result = process_components(test_components)
>>> ok_count = result['ok_count']
>>> valid_ids = result['valid_ids']
>>> exec(base64.b64decode('ZXhwZWN0ZWRfcmVzdWx0ID0gewoib2tfY291bnQiOiAwLAoidmFsaWRfaWRzIjogW10sCiJicmVha19tZXNzYWdlIjogTm9uZSwKfQ==').decode())
>>> assert result == expected_result
>>> if result == expected_result:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(result))
```
failure_message: Failed: The function does not handle an empty list correctly.
success_message: Success: The function correctly handles an empty components list!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-engineering-9'
>>> max_score = 2.0
>>> score = 0
>>> test_components = [{'id': 'DUP1', 'status': 'ok', 'weight': 2.0}, {'id': 'DUP1', 'status': 'ok', 'weight': 3.0}, {'id': 'DUP2', 'status': 'ok', 'weight': 1.5}]
>>> result = process_components(test_components)
>>> valid_ids = result['valid_ids']
>>> exec(base64.b64decode('ZXhwZWN0ZWRfdmFsaWRfaWRzID0gWyJEVVAxIiwgIkRVUDEiLCAiRFVQMiJdICAjIElEcyBzaG91bGQgYmUgYXBwZW5kZWQgYXMgdGhleSBhcHBlYXI=').decode())
>>> assert result['valid_ids'] == expected_valid_ids
>>> if result['valid_ids'] == expected_valid_ids:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(valid_ids))
```
failure_message: Failed: The function does not handle duplicate component IDs correctly.
success_message: Success: The function correctly processes duplicate component IDs!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-engineering-10'
>>> max_score = 2.0
>>> score = 0
>>> test_components = [{'id': 'W1', 'status': 'ok', 'weight': 4.0}, {'id': 'W2', 'status': 'ok', 'weight': 3.9}]
>>> result = process_components(test_components)
>>> ok_count = result['ok_count']
>>> valid_ids = result['valid_ids']
>>> exec(base64.b64decode('ZXhwZWN0ZWRfb2tfY291bnQgPSAy').decode())
>>> exec(base64.b64decode('ZXhwZWN0ZWRfdmFsaWRfaWRzID0gWyJXMSIsICJXMiJd').decode())
>>> assert result['ok_count'] == expected_ok_count
>>> assert result['valid_ids'] == expected_valid_ids
>>> if result['ok_count'] == expected_ok_count and result['valid_ids'] == expected_valid_ids:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(ok_count))
>>> responses = update_responses(question_id, str(valid_ids))
```
failure_message: Failed: The function incorrectly excludes components at exactly 4.0 weight.
success_message: Success: The function correctly allows components at the weight limit!


Tests for loops-and-errors-debugging (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(14.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(161.0)
>>> log_variable('total-points', f'week6-practicemidterm, 1_practicemidterm_q', 161.0)
>>> question_id = 'loops-and-errors-debugging-1'
>>> max_score = 2.0
>>> score = 0
>>> function_exists = callable(square_positive_numbers)
>>> assert function_exists, 'The function `square_positive_numbers` must be defined and callable.'
>>> if (function_exists, 'The function `square_positive_numbers` must be defined and callable.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(function_exists))
```
failure_message: Failed: The function `square_positive_numbers` is not defined correctly.
success_message: Success: The function `square_positive_numbers` exists and is callable!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-errors-debugging-2'
>>> max_score = 4.0
>>> score = 0
>>> condition1 = square_positive_numbers([1, -3, 4, -2]) == [1, 16]
>>> condition2 = square_positive_numbers([-1, -2, -3]) == []
>>> condition3 = square_positive_numbers([0, 2, 3]) == [4, 9]
>>> assert condition1
>>> assert condition2
>>> assert condition3
>>> if condition1 and condition2 and condition3:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(result))
```
failure_message: Failed: The function does not correctly filter and square positive numbers.
success_message: Success: The function correctly filters and squares positive numbers!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-errors-debugging-3'
>>> max_score = 4.0
>>> score = 0
>>> condition = square_positive_numbers([]) == []
>>> assert square_positive_numbers([]) == []
>>> if square_positive_numbers([]) == []:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(result))
```
failure_message: Failed: The function does not handle an empty list correctly.
success_message: Success: The function correctly handles an empty input list!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-errors-debugging-4'
>>> max_score = 4.0
>>> score = 0
>>> condition1 = square_positive_numbers([3, 1, 2]) == [9, 1, 4]
>>> condition2 = square_positive_numbers([7, -1, 5, 0, 3]) == [49, 25, 9]
>>> assert condition1
>>> assert condition2
>>> if condition1 and condition2:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(result))
```
failure_message: Failed: The function does not preserve order of elements.
success_message: Success: The function preserves order of elements!


Tests for numpy-prime-check (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(25.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(161.0)
>>> log_variable('total-points', f'week6-practicemidterm, 1_practicemidterm_q', 161.0)
>>> question_id = 'numpy-prime-check-1'
>>> max_score = 1.0
>>> score = 0
>>> function_exists = callable(find_primes)
>>> assert function_exists, 'The function `find_primes` must be defined and callable.'
>>> if (function_exists, 'The function `find_primes` must be defined and callable.'):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(function_exists))
```
failure_message: Failed: The function `find_primes` is not defined correctly.
success_message: Success: The function `find_primes` exists and is callable!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'numpy-prime-check-2'
>>> max_score = 1.0
>>> score = 0
>>> docstring_exists = find_primes.__doc__ is not None
>>> docstring_word_count = len(find_primes.__doc__.split()) if docstring_exists else 0
>>> assert docstring_exists and docstring_word_count >= 8, f'The function `find_primes` must have a docstring with at least 8 words. Current docstring word count: {docstring_word_count}'
>>> if (docstring_exists and docstring_word_count >= 8, f'The function `find_primes` must have a docstring with at least 8 words. Current docstring word count: {docstring_word_count}'):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(docstring_exists))
>>> responses = update_responses(question_id, str(docstring_word_count))
```
failure_message: Failed: The function `find_primes` does not have a docstring with at least 8 words.
success_message: Success: The function `find_primes` has a docstring with at least 8 words.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'numpy-prime-check-3'
>>> max_score = 2.0
>>> score = 0
>>> input_data = [0, 1, 2, 3, 4, 5, 10, 11, 13, 15, 17]
>>> output_results = find_primes(input_data)
>>> exec(base64.b64decode('ZXhwZWN0ZWRfcmVzdWx0cyA9IG5wLmFycmF5KApbRmFsc2UsIEZhbHNlLCBUcnVlLCBUcnVlLCBGYWxzZSwgVHJ1ZSwgRmFsc2UsIFRydWUsIFRydWUsIEZhbHNlLCBUcnVlXQop').decode())
>>> assert np.array_equal(output_results, expected_results), f'Expected {expected_results}, but got {output_results}'
>>> if np.array_equal(output_results, expected_results):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(input_data))
>>> responses = update_responses(question_id, str(output_results))
```
failure_message: Failed: The function `find_primes` does not work correctly with non-array input.
success_message: Success: The function `find_primes` works correctly with non-array input.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'numpy-prime-check-4'
>>> max_score = 2.0
>>> score = 0
>>> test_input = np.array([2, 3, 5, 7, 11])
>>> exec(base64.b64decode('ZXhwZWN0ZWRfb3V0cHV0ID0gbnAuYXJyYXkoW1RydWUsIFRydWUsIFRydWUsIFRydWUsIFRydWVdKQ==').decode())
>>> prime_test_results = find_primes(test_input)
>>> assert np.array_equal(prime_test_results, expected_output), f'Expected {expected_output}, but got {prime_test_results}.'
>>> if (np.array_equal(prime_test_results, expected_output), f'Expected {expected_output}, but got {prime_test_results}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(prime_test_results))
```
failure_message: Failed: The function does not correctly identify prime numbers.
success_message: Success: The function correctly identifies small prime numbers!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'numpy-prime-check-5'
>>> max_score = 1.0
>>> score = 0
>>> test_array = np.array([20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30])
>>> prime_results = find_primes(test_array)
>>> exec(base64.b64decode('ZXhwZWN0ZWRfcmVzdWx0cyA9IG5wLmFycmF5KApbRmFsc2UsIEZhbHNlLCBGYWxzZSwgVHJ1ZSwgRmFsc2UsIEZhbHNlLCBGYWxzZSwgRmFsc2UsIEZhbHNlLCBUcnVlLCBGYWxzZV0KKQ==').decode())
>>> assert np.array_equal(prime_results, expected_results), f'Expected {expected_results}, but got {prime_results}'
>>> if np.array_equal(prime_results, expected_results):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(test_array))
>>> responses = update_responses(question_id, str(prime_results))
```
failure_message: Failed: The function `find_primes` does not correctly identify prime numbers in a larger array.
success_message: Success: The function `find_primes` correctly identifies prime numbers in a larger array.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'numpy-prime-check-6'
>>> max_score = 1.0
>>> score = 0
>>> test_array = np.array([0, 1, 4, 6, 8, 9, 10])
>>> prime_results = find_primes(test_array)
>>> exec(base64.b64decode('ZXhwZWN0ZWRfcmVzdWx0cyA9IG5wLmFycmF5KFtGYWxzZSwgRmFsc2UsIEZhbHNlLCBGYWxzZSwgRmFsc2UsIEZhbHNlLCBGYWxzZV0p').decode())
>>> assert np.array_equal(prime_results, expected_results), f'Expected {expected_results}, but got {prime_results}'
>>> if np.array_equal(prime_results, expected_results):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(test_array))
>>> responses = update_responses(question_id, str(prime_results))
```
failure_message: Failed: The function `find_primes` does not handle edge cases correctly.
success_message: Success: The function `find_primes` handles edge cases correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'numpy-prime-check-7'
>>> max_score = 2.0
>>> score = 0
>>> test_input = np.array([4, 6, 8, 9, 10, 12])
>>> exec(base64.b64decode('ZXhwZWN0ZWRfb3V0cHV0ID0gbnAuYXJyYXkoW0ZhbHNlLCBGYWxzZSwgRmFsc2UsIEZhbHNlLCBGYWxzZSwgRmFsc2VdKQ==').decode())
>>> composite_test_results = find_primes(test_input)
>>> assert np.array_equal(composite_test_results, expected_output), f'Expected {expected_output}, but got {composite_test_results}.'
>>> if (np.array_equal(composite_test_results, expected_output), f'Expected {expected_output}, but got {composite_test_results}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(composite_test_results))
```
failure_message: Failed: The function incorrectly classifies composite numbers as prime.
success_message: Success: The function correctly identifies composite numbers!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'numpy-prime-check-8'
>>> max_score = 2.0
>>> score = 0
>>> test_input = np.array([0, 1, -1, -5, -10])
>>> exec(base64.b64decode('ZXhwZWN0ZWRfb3V0cHV0ID0gbnAuYXJyYXkoW0ZhbHNlLCBGYWxzZSwgRmFsc2UsIEZhbHNlLCBGYWxzZV0p').decode())
>>> edge_case_results = find_primes(test_input)
>>> assert np.array_equal(edge_case_results, expected_output), f'Expected {expected_output}, but got {edge_case_results}.'
>>> if (np.array_equal(edge_case_results, expected_output), f'Expected {expected_output}, but got {edge_case_results}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(edge_case_results))
```
failure_message: Failed: The function does not correctly handle edge cases.
success_message: Success: The function correctly handles edge cases (0, 1, negatives)!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'numpy-prime-check-9'
>>> max_score = 3.0
>>> score = 0
>>> test_input = np.array([101, 103, 107, 109, 113, 200, 201, 202, 203])
>>> exec(base64.b64decode('ZXhwZWN0ZWRfb3V0cHV0ID0gbnAuYXJyYXkoW1RydWUsIFRydWUsIFRydWUsIFRydWUsIFRydWUsIEZhbHNlLCBGYWxzZSwgRmFsc2UsIEZhbHNlXSk=').decode())
>>> large_number_results = find_primes(test_input)
>>> assert np.array_equal(large_number_results, expected_output), f'Expected {expected_output}, but got {large_number_results}.'
>>> if (np.array_equal(large_number_results, expected_output), f'Expected {expected_output}, but got {large_number_results}.'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(large_number_results))
```
failure_message: Failed: The function does not efficiently handle large numbers.
success_message: Success: The function correctly handles large numbers!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'numpy-prime-check-10'
>>> max_score = 3.0
>>> score = 0
>>> test_input = np.array([[2, 3, 4], [5, 6, 7]])
>>> exec(base64.b64decode('ZXhwZWN0ZWRfc2hhcGUgPSB0ZXN0X2lucHV0LnNoYXBl').decode())
>>> shape_test_results = find_primes(test_input).shape
>>> assert shape_test_results == expected_shape, f'Expected shape {expected_shape}, but got {shape_test_results}.'
>>> if (shape_test_results == expected_shape, f'Expected shape {expected_shape}, but got {shape_test_results}.'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(shape_test_results))
```
failure_message: Failed: The function does not preserve the shape of the input array.
success_message: Success: The function maintains NumPy array shape!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'numpy-prime-check-11'
>>> max_score = 3.0
>>> score = 0
>>> test_input = np.array([[29, 30, 31], [32, 33, 34]])
>>> exec(base64.b64decode('ZXhwZWN0ZWRfb3V0cHV0ID0gbnAuYXJyYXkoW1tUcnVlLCBGYWxzZSwgVHJ1ZV0sIFtGYWxzZSwgRmFsc2UsIEZhbHNlXV0p').decode())
>>> large_2d_results = find_primes(test_input)
>>> assert np.array_equal(large_2d_results, expected_output), f'Expected {expected_output}, but got {large_2d_results}.'
>>> if (np.array_equal(large_2d_results, expected_output), f'Expected {expected_output}, but got {large_2d_results}.'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(large_2d_results))
```
failure_message: Failed: The function does not correctly process 2D arrays.
success_message: Success: The function correctly handles large 2D arrays!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'numpy-prime-check-12'
>>> max_score = 1.0
>>> score = 0
>>> test_array = np.array([7])
>>> prime_results = find_primes(test_array)
>>> exec(base64.b64decode('ZXhwZWN0ZWRfcmVzdWx0cyA9IG5wLmFycmF5KFtUcnVlXSk=').decode())
>>> assert np.array_equal(prime_results, expected_results), f'Expected {expected_results}, but got {prime_results}'
>>> if np.array_equal(prime_results, expected_results):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(test_array))
>>> responses = update_responses(question_id, str(prime_results))
```
failure_message: Failed: The function `find_primes` does not correctly identify a single prime number.
success_message: Success: The function `find_primes` correctly identifies a single prime number.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'numpy-prime-check-13'
>>> max_score = 3.0
>>> score = 0
>>> test_input = np.array([1000003, 1000033, 1000037, 1000039, 1000000, 1000001])
>>> exec(base64.b64decode('ZXhwZWN0ZWRfb3V0cHV0ID0gbnAuYXJyYXkoW1RydWUsIFRydWUsIFRydWUsIFRydWUsIEZhbHNlLCBGYWxzZV0p').decode())
>>> huge_number_results = find_primes(test_input)
>>> assert np.array_equal(huge_number_results, expected_output), f'Expected {expected_output}, but got {huge_number_results}.'
>>> if (np.array_equal(huge_number_results, expected_output), f'Expected {expected_output}, but got {huge_number_results}.'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(huge_number_results))
```
failure_message: Failed: The function does not efficiently process very large numbers.
success_message: Success: The function efficiently handles very large numbers!


Tests for symbolic-heart-curve (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(16.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(161.0)
>>> log_variable('total-points', f'week6-practicemidterm, 1_practicemidterm_q', 161.0)
>>> question_id = 'symbolic-heart-curve-1'
>>> max_score = 2.0
>>> score = 0
>>> function_exists = callable(get_equation)
>>> assert function_exists, 'The function `get_equation` must be defined and callable.'
>>> if (function_exists, 'The function `get_equation` must be defined and callable.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(function_exists))
```
failure_message: Failed: The function `get_equation` is not defined correctly.
success_message: Success: The function `get_equation` exists and is callable!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> result = get_equation()
>>> heart_eq_returned = result[0]
>>> from sympy import Eq
>>> question_id = 'symbolic-heart-curve-2'
>>> max_score = 4.0
>>> score = 0
>>> exec(base64.b64decode('aXNfZXEgPSBpc2luc3RhbmNlKGhlYXJ0X2VxX3JldHVybmVkLCBFcSk=').decode())
>>> assert is_eq, 'Expected a SymPy `Eq` equation as the first return value.'
>>> if (is_eq, 'Expected a SymPy `Eq` equation as the first return value.'):
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(is_eq))
```
failure_message: Failed: The function does not return a valid SymPy equation.
success_message: Success: The function correctly returns a SymPy equation!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import sympy
>>> result = get_equation()
>>> heart_plot_returned = result[1]
>>> from sympy.plotting.plot import Plot
>>> question_id = 'symbolic-heart-curve-3'
>>> max_score = 4.0
>>> score = 0
>>> exec(base64.b64decode('aXNfcGxvdCA9IGlzaW5zdGFuY2UoaGVhcnRfcGxvdF9yZXR1cm5lZCwgUGxvdCk=').decode())
>>> assert is_plot, 'Expected a SymPy `Plot` object as the second return value.'
>>> if (is_plot, 'Expected a SymPy `Plot` object as the second return value.'):
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(is_plot))
```
failure_message: Failed: The function does not return a valid SymPy plot object.
success_message: Success: The function correctly returns a SymPy plot object!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from sympy import Eq
>>> question_id = 'symbolic-heart-curve-4'
>>> max_score = 6.0
>>> score = 0
>>> expected_equation = Eq((x ** 2 + y ** 2 - 1) ** 3 - x ** 2 * y ** 3, 0)
>>> result = get_equation()
>>> heart_eq_returned = result[0]
>>> matches_equation = heart_eq_returned == expected_equation
>>> assert matches_equation, f'Expected the equation to be {expected_equation}, but got {heart_eq_returned}.'
>>> if (matches_equation, f'Expected the equation to be {expected_equation}, but got {heart_eq_returned}.'):
...     score = 6.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicemidterm_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(matches_equation))
```
failure_message: Failed: The heart equation is not correctly defined.
success_message: Success: The heart equation matches the expected symbolic expression!
