```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_homework_q", "week_8", "homework", assignment_points = 75.0, assignment_tag = 'week8-homework')

# Initialize Otter
import otter
grader = otter.Notebook("1_homework_q.ipynb")
```

# 💻🧪 Lab 10: Combining Concepts to Build Interoperable Code for Plotting

This lab uses all of the concepts we have learned in python programming to make reusable building blocks for generating, plotting, and fitting data. 

As an engineer you will regularly be presented with noisy data and want to fit that data to a function. In this assignment, we will build some machinery and a coding schema to generate noisy data (since we do not have real data), plot data, and fit the results. We will build tools to simplify the visualization of the results. The code will be designed to be interoperable. If we were to have a new type of data or mathematical expression, we could reuse all the code we have written. 

## Question 1 (Points: 3.0): Importing Functions

Here we will import some packages that we need for this assignment

1. import `numpy` as `np`

2. import the submodule `matplotlib.pyplot` and assign it to `plt`

3. import from `scipy.optimize` the `curve_fit` function

```python
# 1. you need to import numpy as np
...
# 2. You need to import the submodule matplotlib.pyplot and assign it to plt
...
# 3. You need to import from scipy.optimize the curve_fit function
...
```

```python
grader.check("q1-Importing Functions")
```

## Question 2 (Points: 13.0): 2. Implementing a Class for MathExpressions

In Python and programming in general, it is common to have a base class that adds functionality to an object. For example, if you have a car that you are building, you might want to add blind spot detection. You can create the blind spot detection software and hardware and install it on multiple cars types. We will do the same with a math function. We will build a class that adds a methods to fit and evaluate a math expression. 

When we build this class we need it to be flexible to accept math functions with a flexible number of input parameters. We implement a fit function where we can input any parameters into the function (this is important for optimization required to fit data), and an evaluate method which uses the parameters set during initialization. 

Follow these steps:

1. Define a `class` `MathExpression`.

2. Build an initialization function. 
- The function should take a variable `func` which is a mathematical function
- The function should accept `**kwargs` 
    - The `kwargs` will be the parameters of the fitting function 
    - We want to save the `kwargs`, each key-value pair of `kwargs` where the key is the variable name, and `func` as an attribute of the object

3. Define a method `fit`.
- That accepts a required input `x`, and `**kwargs`
- Return the result when you call the objects method `func` that accepts a required input `x`, and `**kwargs`

4. Define an method for the `class` `evaluate`.
- Takes a required input `x`
- Returns the result when you call the objects method `func` that accepts a required input `x`, and the `**kwargs` values set at initialization.

```python
# 1. define a class MathExpression
...
    # 2. Use the built in initialization function
    # The function should take a variable func which is a mathematical function for generation and fitting (we will define this later)
    # The function should accept **kwargs
    # The kwargs will be the parameters of the fitting function 
    ...
        
        # use the .items() method to extract from kwargs a list of tuples with the key value pairs
        # Create a for loop that extracts the key value pairs from the kwargs
        #   make sure to unpack the tuple into the variables key and value
        # Inside the for loop, use the built in python method setattr to add an attribute to the MathExpression Object 
        #   for the key, use the variable key, 
        #   for the value, use the variable value
        ...
        
        # save the kwargs as an attribute of the object named kwargs
        ...
        
        # save the function as an attribute of the object named func
        ...
        
    # 3. for the class, define a method fit that accepts a required input x, and **kwargs
    #   x is the points to sample along the x axis
    #   **kwargs are new parameters to test
    #   Note: this is a requirement of the fitting function
    ...
        
        # return the result when you call the object's method func that accepts a required input x and **kwargs
        ...
        
    # 4. for the class, define a method evaluate that takes a required input x
    ...
       
        # return the result when you call the object's method func that accepts a required input x and the kwargs values set at initialization.
        ...
```

```python
grader.check("q2-Implement MathExpression")
```

## Question 3 (Points: 3.0): 3. Define a Linear Function

$$ Y = slope * x + intercept $$

We want to define a simple linear function. A linear model is one of the simplest fitting functions used in science and engineering. It says that an independent variable directly influences a dependent variable. 

1. Build a function `LinearFunction`
- `LinearFunction` should take three inputs: `x` for the independent variable, `m` for the slope, `b` for the intercept.
- Return the calculation of `y`, where y = slope * x + intercept. 
- Make sure to replace slope and intercept with the correct variables.

```python
# 1. build a function LinearFunction
# linear function should take two inputs m - for the slope, and b for the intercept
# return the calculation of y, where y = slope * x + intercept. 
# Make sure to replace slope and intercept with the correct variables.
...

```

```python
grader.check("q3 - Linear Function")
```

## Question 4 (Points: 6.0): 4. Plotting

One of the best ways to understand your code is doing what you expect is to visualize the results. We can do this by plotting the data. 

1. Instantiate an object `Linear_` using the `MathExpression` `Class`.
- The input `func` should be the `LinearFunction` you wrote
- The variables needed for kwargs must be provided as key value pairs. 
- In this case, those key value pairs are `m = 1`, and `b = -3`. 

2. Derive a linearly-spaced vector for the independent variable.
- Use the `np.linspace` function to make a linearly-spaced array from `0` to `10` with `100` steps

3. call the built-in `evaluate` method of the `LinearFunction` object over the range defined by `x`.
- assign the output to a variable `y`

4. Use the function `plt.plot` to plot `x` vs `y`.
- assign the plot to a variable called `linear_plot`

```python
# 1. Instantiate an object Linear_ using the MathExpression Class
# The variables needed for the Kwargs m = 1, and b = -3 need to be provided as key value pairs. 
# example my_function(x=1), this would provide a **kwargs key (x) value (1)

...

# 2. use the np.linspace function to make a linearly-spaced array from 0 to 10, with 100 steps
# save this to the variable x

...

# 3. call the built-in `evaluate` method of the LinearFunction object over the range defined by x
# assign the output to a variable y

...

# 4. use the function plt.plot to plot x vs y

...

# we do not need to add more labels to this plot. It was just for testing.
```

```python
grader.check("q4 - Plotting")
```

# 5. A NoisyFunction

You are responsible for making a 3-4 minute VoiceThread submission that briefly explains each function of this class. *You are not responsible for writing any code or comments in this Task.*

We have written a function with several useful methods that you might want to use as an engineer. 

1 It inherits information from an object of type `MathExpression`

2 It can generate noisy data, which is to simulate real experimental data.

3 It can fit the data to an known function. 

4 It can fit, plot and format the results for quick visualization.

```python
# build a class Noisy Function that Inherits MathExpression
# This is simply done writing class NoisyFunction(MathExpression)
...
    # This class does not require any specific initialization variables, but it will require that we pass the inherited variable through.
    # we have done this for you 
    # you should not touch this code. If you do, here it is
    # def __init__(self, *args, **kwargs):
    #    super().__init__(*args, **kwargs)
    ...
    # build a function generate noisy data with two required inputs
    #   x_range - a list of the minimum and max ranges for to generate the data
    #   noise_amplitude - a float from which sets the magnitude of the noise
    ...
        # use np.linspace to generate x data along the range
        # generate 100 points
        ...
        # call the evaluate method of the base function object to return the true values
        ...
        # generate the noise vector by multiplying the noise amplitude by an random noise vector of length x  
        # the random noise vector should be generated with np.random.randn
        ...
        # add the noise to the y_true value
        # save it in a new variable y
        ...
        # return x, y, and y_true
        ...
    # build a function fit_data that takes 2 inputs x and y
    ...
        # call the curvefit function
        # provide the inherited pointed to the function
        # For the initial guess provide the true values they are contained within the saved variable kwargs
        # You can get the values of a dictionary using the built-in method .values()
        # P0 takes a list. You need to convert the values to a list.
        # Curvefit should return two values assign them to popt and pcov. 
        # 
        # is the fit results
        ...
        # Return popt and pcov
        ...
    # build a function plot_fit_results that plots the raw and fit results
    # This function should take two input parameters:
    #   x_range - a tuple defining the x range to plot
    #   noise_amplitude - a float used to control the magnitude of the noise
    ...

            # use the subplots function in pyplot to create a multiplot figure with 2 columns and 1 row.
            # set the figure size = 10, 5 - figsize takes a tuple
            # assign the output to two variables fig and axs
            ...

            # make a list of strings called labels 
            # The 0th index should be "True Fit"
            # The 1st index should be "Noisy Fit"
            ...

            # call the generate_noisy_data function from this class
            # save the returned variables to x, y, and y_true
            ...

            # make a look that loops over the axs with enumerate.
            # set the iterator equal to the variable i
            # set the value equal to the variable ax
            ...
    
                # create a scatterplot on the axis object 
                # plot x vs y
                # make the marker size = 5 (with the `s` tag)
                # add a label 'Raw Data' (with the `label` tag)
                ...
    
                # write an if statement that is true if the label of the graph is "True Fit". This is ordered by the list labels you created.
                ...
        
                    # if the statement was true on the axis object plot
                    # x vs y true
                    # make the line red using the 'r' flag
                    # add a label as indexed from the list
                    ...
    
                # add an elif statement that is true if the label of the graph is "Noisy Fit" this is ordered by the list labels you created.
                ...
        
                    # call the method in the class fit_data
                    # This should fit the x and y data
                    # have it return the variables popt and pcov
                    ...
        
                    # use the predicted popt and the inherited function to generate the y values for the fit results
                    # make sure to unpack the popt list with the * operator
                    ...
        
                    # on the axis object plot
                    # x vs y_noisy_fit
                    # make the line blue using the 'b' flag
                    # add a label as indexed from the list
                    ...
    
                # else passes to the next loop
                ...
    
                # use the set_xlabel method to set the x-label of the graph to 'x'
                # use the set_ylabel method to set the y-label of the graph to 'y'
                ...
    
                # call the legend method of the axis object to show the legend
                ...
            # return the fig object
            ...
```

```python
grader.check("q5 - Noisy Function")
```

## Question 6 (Points: 11.0): 6. Testing our Class

In coding it is always a good idea to run all parts of your code to check it is performing as expected. 

1. Create a `NoisyFunction` instance for a linear function
- Assign this object to the variable `linear_func`
- The input `func` should be the `LinearFunction` you wrote
- The slope should be `3` and the intercept should be `4`

2. Call the proper method of `linear_func` to generate noisy data over the range from 0- $2\pi$
- the noise amplitude is `1`
- save the results to `x`, `y`, and `y_true`

3. Fit the `x` and `y` data using the proper method of `linear_func`
- save the fit results to `popt`, `pcov`
- (these names represent the optimal parameters (`popt`) and the parameters' covariance (`pcov`))

4. Use the built-in method to plot the fit result
- set the range equal to `0` to `16`
- set the noise amplitude equal to `0.3`
- Save this to the variable plot


```python
seed = np.random.seed(42)
# 1. Create a NoisyFunction instance for a linear function
# Assign this object to the variable linear_func
# The slope should be 3 and the intercept should be 4
...
# 2. call the generate_noisy_data function of the object the linear_func to produce noisy data over the range from 0-2pi
#   use a tuple to indicate the range from 0-2pi
#   the noise amplitude is 1
#   save the results to x, y, and y_true
...
# 3. Fit the data
# save the fit results to popt, pcov
...

# 4. Use the built-in method to plot the fit result
# set the range equal to 0 to 16
# set the noise amplitude equal to 0.3
# Save this to the variable plot
...
```

```python
grader.check("q6 - plotting and fitting")
```

## Question 7 (Points: 6.0): 7. Defining new math functions

Now we can view the flexibility (interoperability) of our code. 

Define two new functions. Follow the logic of the way you created `LinearFunction`.

1. `SineFunction`
   
$$g(x) = A (sin(2\pi f x+\phi))$$

where $A$ represents amplitude, $f$ represents the frequency, and $\phi$ represents the phase.

2. `ExponentialFunction`

$$h(x) = ae^{bx}$$
   

```python
# 1. Write the SineFunction here
# the inputs should be x, amplitude, frequency, phase
...

# 2. Write the ExponentialFunction here
# the inputs should be x, a, b
...
```

```python
grader.check("q7-new functions")
```

# 8 Using our class with the sine function

Now we can use nearly the same code from Task 7 to test our `SineFunction`.

1. Create a `NoisyFunction` instance for the sine function
- Assign this object to the variable `sine_func`
- The input `func` should be the `SineFunction` you wrote
- The amplitude should be 2, the frequency should be .2, and the phase should be .5

2. Apply the `generate_noisy_data` method of `sine_func` to produce noisy data over the range from 0-6 $\pi$
- use a tuple to indicate the range from 0-6 $\pi$
- the noise amplitude is 1.5
- save the results to `x`, `y`, and `y_true`

3. Fit the `x` and `y` data using the proper method of `sine_func`
- save the fit results to `popt`, `pcov`

4. Use the built-in method to plot the fit results
- set the range equal to 0 to 6 $\pi$
- set the noise amplitude equal to 1.5
- Save this to the variable `plot`

```python
seed = np.random.seed(42)
# 1. Create a NoisyFunction instance for a sine function
# Assign this object to the variable sine_func
# The amplitude should be 2, the frequency should be .2, and the phase should be .5
...
# 2. evaluate the sine_func to produce noisy data over the range from 0-6pi
# the noise amplitude is 1.5
# save the results to x, y, and y_true
...
# 3. Fit the data
# save the fit results to popt, pcov
...
# 4. Use the built-in method to plot the fit results
# set the range equal to 0 to 6 pi
# set the noise amplitude equal to 1.5
# Save this to the variable plot
...
```

```python
grader.check("q8-using the class with new Sine functions")
```

# 9 Using our class with the Exponential function

Now we can use nearly the exact same code to test our Exponential function. We wrote one code that we can use over and over again in different ways!

1. Create a `NoisyFunction` instance for a exponential function
- Assign this object to the variable `exp_func`
- The input `func` should be the `ExponentialFunction` you wrote
- `a` should be 2, and `b` should be .2

2. Call the proper method of `exp_func` to generate noisy data over the range from 0-6 $\pi$
- use a tuple to indicate the range from 0-6 $\pi$
- the noise amplitude is 1.5
- save the results to `x`, `y`, and `y_true`

3. Fit the `x` and `y` data using the proper method of `exp_func`
- save the fit results to `popt`, `pcov`

4. Use the built-in method to plot the fit results
- set the range equal to 0 to 6 $\pi$
- set the noise amplitude equal to 1.5
- save this to the variable `plot`

```python
seed = np.random.seed(42)
# 1. Create a NoisyFunction instance for a exponential function
# Assign this object to the variable exp_func
# The a should be 2, and b should be .2
...
# 2. evaluate the exp_func to produce noisy data over the range from 0-6pi
# the noise amplitude is 1.5
# save the results to x, y, and y_true
...
# 3. Fit the data
# save the fit results to popt, pcov
...
# 4. Use the built-in method to plot the fit results
# set the range equal to 0 to 6 pi
# set the noise amplitude equal to 1.5
# Save this to the variable plot
...

```

```python
grader.check("q9-using the class with new exponential functions")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week8-homework", "1_homework_q")
```

Tests for q1-Importing Functions (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(3.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(75.0)
>>> log_variable('total-points', f'week8-homework, 1_homework_q', 75.0)
>>> question_id = 'q1-Importing Functions-1'
>>> max_score = 1.0
>>> score = 0
>>> assert np.__version__ is not None
>>> if np.__version__ is not None:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: numpy random incorrectly implemented.
success_message: numpy correctly implemented.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q1-Importing Functions-2'
>>> max_score = 1.0
>>> score = 0
>>> assert plt.__package__ is not None
>>> if plt.__package__ is not None:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: plt random incorrectly implemented.
success_message: plt correctly implemented.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q1-Importing Functions-3'
>>> max_score = 1.0
>>> score = 0
>>> assert curve_fit.__name__ is not None
>>> if curve_fit.__name__ is not None:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: plt random incorrectly implemented.
success_message: plt correctly implemented.


Tests for q2-Implement MathExpression (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(13.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(75.0)
>>> log_variable('total-points', f'week8-homework, 1_homework_q', 75.0)
>>> question_id = 'q2-Implement MathExpression-1'
>>> max_score = 1.0
>>> score = 0
>>> assert isinstance(MathExpression, type)
>>> if isinstance(MathExpression, type):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: MathExpression is not a Class
success_message: MathExpression is a Class
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q2-Implement MathExpression-2'
>>> max_score = 1.0
>>> score = 0
>>> assert '__init__' in dir(MathExpression)
>>> if '__init__' in dir(MathExpression):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: __init__ not implemented.
success_message: __init__ implemented.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q2-Implement MathExpression-3'
>>> max_score = 2.0
>>> score = 0
>>> out = MathExpression('a', b=1, test='test')
>>> assert out.kwargs['b'] == 1
>>> assert out.kwargs['test'] == 'test'
>>> if out.kwargs['b'] == 1 and out.kwargs['test'] == 'test':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: MathExpression inputs kwargs not correctly implemented.
success_message: MathExpression inputs kwargs correctly implemented.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q2-Implement MathExpression-4'
>>> max_score = 1.0
>>> score = 0
>>> out = MathExpression(np.sum)
>>> assert out.func is np.sum
>>> if out.func is np.sum:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: MathExpression function input not correctly implemented.
success_message: MathExpression function input correctly implemented.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q2-Implement MathExpression-5'
>>> max_score = 1.0
>>> score = 0
>>> assert 'fit' in dir(MathExpression)
>>> if 'fit' in dir(MathExpression):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: fit is not implemented.
success_message: fit is implemented.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q2-Implement MathExpression-6'
>>> max_score = 3.0
>>> score = 0
>>> def test(x, **kwargs):
...     return (x, kwargs)
>>> out = MathExpression(test)
>>> a, b = out.fit(2, a=2, b=3)
>>> assert a == 2
>>> assert b == {'a': 2, 'b': 3}
>>> if a == 2 and b == {'a': 2, 'b': 3}:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: fit is not implemented correctly.
success_message: fit is implemented correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q2-Implement MathExpression-7'
>>> max_score = 1.0
>>> score = 0
>>> assert 'evaluate' in dir(MathExpression)
>>> if 'evaluate' in dir(MathExpression):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: evaluate is not implemented.
success_message: evaluate is implemented.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q2-Implement MathExpression-8'
>>> max_score = 3.0
>>> score = 0
>>> def test(x, **kwargs):
...     return (x, kwargs)
>>> out = MathExpression(test, a=5, b=3)
>>> a, b = out.evaluate(2)
>>> assert a == 2
>>> assert b == {'a': 5, 'b': 3}
>>> if a == 2 and b == {'a': 5, 'b': 3}:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: evaluate is not implemented correctly.
success_message: evaluate is implemented correctly.


Tests for q3 - Linear Function (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from inspect import isfunction
>>> max_question_points = str(3.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(75.0)
>>> log_variable('total-points', f'week8-homework, 1_homework_q', 75.0)
>>> question_id = 'q3 - Linear Function-1'
>>> max_score = 1.0
>>> score = 0
>>> assert isfunction(LinearFunction)
>>> if isfunction(LinearFunction):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: LinearFunction is not a function
success_message: LinearFunction is a function
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import inspect
>>> question_id = 'q3 - Linear Function-2'
>>> max_score = 1.0
>>> score = 0
>>> signature = inspect.signature(LinearFunction)
>>> num_parameters = len(signature.parameters)
>>> assert num_parameters == 3
>>> if num_parameters == 3:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: LinearFunction function has the incorrect inputs.
success_message: LinearFunction function has the correct inputs.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q3 - Linear Function-3'
>>> max_score = 1.0
>>> score = 0
>>> x = np.array([2, 3, 4, 5, 8, 10])
>>> out = LinearFunction(x, 7, 4)
>>> assert all(out == [18, 25, 32, 39, 60, 74])
>>> if all(out == [18, 25, 32, 39, 60, 74]):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: LinearFunction function implemented incorrect inputs.
success_message: LinearFunction function implemented correctly


Tests for q4 - Plotting (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(6.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(75.0)
>>> log_variable('total-points', f'week8-homework, 1_homework_q', 75.0)
>>> question_id = 'q4 - Plotting-1'
>>> max_score = 1.0
>>> score = 0
>>> assert isinstance(Linear_, MathExpression)
>>> assert Linear_.m == 1
>>> assert Linear_.b == -3
>>> if isinstance(Linear_, MathExpression) and Linear_.m == 1 and (Linear_.b == -3):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: MathExpression is not instantiated correctly
success_message: MathExpression is instantiated correctly
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q4 - Plotting-2'
>>> max_score = 1.0
>>> score = 0
>>> seed = np.random.seed(42)
>>> assert np.isclose(sum(x[::4]), 121.212121, rtol=0.01)
>>> assert x.__len__() == 100
>>> if np.isclose(sum(x[::4]), 121.212121, rtol=0.01) and x.__len__() == 100:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: X vector is incorrect
success_message: X vector is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q4 - Plotting-3'
>>> max_score = 1.0
>>> score = 0
>>> seed = np.random.seed(42)
>>> assert sum(y) == 200
>>> if sum(y) == 200:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: evaluate computed incorrectly
success_message: evaluate computed correctly
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q4 - Plotting-4'
>>> max_score = 3.0
>>> score = 0
>>> seed = np.random.seed(42)
>>> assert 700 == np.sum(linear_plot[0].get_data())
>>> if 700 == np.sum(linear_plot[0].get_data()):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The plot is incorrect
success_message: The plot is correct


Tests for q5 - Noisy Function (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(11.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(75.0)
>>> log_variable('total-points', f'week8-homework, 1_homework_q', 75.0)
>>> question_id = 'q5 - Noisy Function-1'
>>> max_score = 1.0
>>> score = 0
>>> condition = issubclass(NoisyFunction, MathExpression)
>>> assert condition
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: NoisyFunction does not correctly inherit from MathExpression.
success_message: NoisyFunction correctly inherits from MathExpression.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q5 - Noisy Function-2'
>>> max_score = 2.0
>>> score = 0
>>> def test_func(x, a, b):
...     return a * x + b
>>> nf = NoisyFunction(test_func, a=2, b=5)
>>> condition = nf.func is test_func
>>> condition1 = nf.kwargs == {'a': 2, 'b': 5}
>>> assert condition
>>> assert condition1
>>> if condition and condition1:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: NoisyFunction class does not initialize correctly.
success_message: NoisyFunction class initializes correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q5 - Noisy Function-3'
>>> max_score = 2.0
>>> score = 0
>>> def test_func(x, a, b):
...     return a * x + b
>>> nf = NoisyFunction(test_func, a=2, b=5)
>>> x_range = (0, 10)
>>> noise_amplitude = 1.0
>>> x, y, y_true = nf.generate_noisy_data(x_range, noise_amplitude)
>>> condition = len(x) == 100
>>> assert condition
>>> condition = len(y) == 100
>>> assert condition
>>> condition = len(y_true) == 100
>>> assert condition
>>> condition = not np.array_equal(y, y_true)
>>> assert condition
>>> if condition and condition and condition and condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: generate_noisy_data does not produce correct outputs.
success_message: generate_noisy_data produces correct outputs.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q5 - Noisy Function-4'
>>> max_score = 3.0
>>> score = 0
>>> def test_func(x, a, b):
...     return a * x + b
>>> nf = NoisyFunction(test_func, a=2, b=5)
>>> popt, pcov = nf.fit_data(x, y)
>>> condition = len(popt) == 2
>>> assert condition
>>> condition = pcov.shape == (2, 2)
>>> assert condition
>>> if condition and condition:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: fit_data function does not return expected parameters.
success_message: fit_data function returns expected parameters.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q5 - Noisy Function-5'
>>> max_score = 3.0
>>> score = 0
>>> def test_func(x, a, b):
...     return a * x + b
>>> nf = NoisyFunction(test_func, a=2, b=5)
>>> x_range = (0, 10)
>>> noise_amplitude = 1.0
>>> fig = nf.plot_fit_results(x_range, noise_amplitude)
>>> condition = fig is not None
>>> assert condition
>>> condition = len(fig.axes) == 2
>>> assert condition
>>> if condition and condition:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: plot_fit_results does not generate a valid figure.
success_message: plot_fit_results generates a valid figure.


Tests for q6 - plotting and fitting (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> max_question_points = str(11.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(75.0)
>>> log_variable('total-points', f'week8-homework, 1_homework_q', 75.0)
>>> question_id = 'q6 - plotting and fitting-1'
>>> max_score = 1.0
>>> score = 0
>>> seed = np.random.seed(42)
>>> assert isinstance(linear_func, NoisyFunction)
>>> assert linear_func.m == 3
>>> assert linear_func.b == 4
>>> assert linear_func.func == LinearFunction
>>> if isinstance(linear_func, NoisyFunction) and linear_func.m == 3 and (linear_func.b == 4) and (linear_func.func == LinearFunction):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: linear_func implemented incorrectly
success_message: linear_func implemented correctly
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> question_id = 'q6 - plotting and fitting-2'
>>> max_score = 0.0
>>> score = 0
>>> assert seed == np.random.seed(42)
>>> if seed == np.random.seed(42):
...     score = 0.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: You might have deleted the line that sets the random seed. Please add it at the top of the cell. You will fail the autograder if you do not place the seed in the top of the cell.
success_message: Your seed is still correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> question_id = 'q6 - plotting and fitting-3'
>>> max_score = 1.0
>>> score = 0
>>> seed = np.random.seed(42)
>>> assert np.isclose(sum([x, y, y_true])[4], 9.542909136398162, rtol=0.01)
>>> assert np.isclose(np.mean(sum([x, y, y_true])[4::3]), 30.365526866857397, rtol=0.01)
>>> if np.isclose(sum([x, y, y_true])[4], 9.542909136398162, rtol=0.01) and np.isclose(np.mean(sum([x, y, y_true])[4::3]), 30.365526866857397, rtol=0.01):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: x, y, y_true return incorrectly
success_message: x, y, y_true return correctly
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> question_id = 'q6 - plotting and fitting-4'
>>> max_score = 4.0
>>> score = 0
>>> seed = np.random.seed(42)
>>> assert np.isclose(popt[0], 3.0219526649250503, rtol=0.01)
>>> assert np.isclose(popt[1], 3.8271871517503655, rtol=0.01)
>>> if np.isclose(popt[0], 3.0219526649250503, rtol=0.01) and np.isclose(popt[1], 3.8271871517503655, rtol=0.01):
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Fitting function incorrectly implemented
success_message: Fitting function correctly implemented
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> question_id = 'q6 - plotting and fitting-5'
>>> max_score = 5.0
>>> score = 0
>>> seed = np.random.seed(42)
>>> assert np.sum(plot.axes[0].get_lines()[0].get_xydata()) == 3600
>>> assert np.isclose(np.sum(plot.axes[1].get_lines()[0].get_xydata()), 3600.6691376113495, rtol=0.01)
>>> if np.sum(plot.axes[0].get_lines()[0].get_xydata()) == 3600 and np.isclose(np.sum(plot.axes[1].get_lines()[0].get_xydata()), 3600.6691376113495, rtol=0.01):
...     score = 5.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: plot incorrectly implemented
success_message: plot implemented correctly


Tests for q7-new functions (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import inspect
>>> max_question_points = str(6.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(75.0)
>>> log_variable('total-points', f'week8-homework, 1_homework_q', 75.0)
>>> question_id = 'q7-new functions-1'
>>> max_score = 3.0
>>> score = 0
>>> signature = inspect.signature(SineFunction)
>>> num_parameters = len(signature.parameters)
>>> seed = np.random.seed(42)
>>> assert num_parameters == 4
>>> assert np.allclose([-1.51360499, -0.77557439, -0.456128], SineFunction(np.array([0.0, 0.1, 0.3]), 2, 3, 4), rtol=0.01)
>>> if num_parameters == 4 and np.allclose([-1.51360499, -0.77557439, -0.456128], SineFunction(np.array([0.0, 0.1, 0.3]), 2, 3, 4), rtol=0.01):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: SineFunction incorrectly implemented
success_message: SineFunction implemented correctly
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import inspect
>>> question_id = 'q7-new functions-2'
>>> max_score = 3.0
>>> score = 0
>>> signature = inspect.signature(ExponentialFunction)
>>> num_parameters = len(signature.parameters)
>>> seed = np.random.seed(42)
>>> assert num_parameters == 3
>>> assert np.allclose([2, 2.69971762, 4.91920622], ExponentialFunction(np.array([0.0, 0.1, 0.3]), 2, 3), rtol=0.01)
>>> if num_parameters == 3 and np.allclose([2, 2.69971762, 4.91920622], ExponentialFunction(np.array([0.0, 0.1, 0.3]), 2, 3), rtol=0.01):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: SineFunction incorrectly implemented
success_message: SineFunction implemented correctly


Tests for q8-using the class with new Sine functions (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(11.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(75.0)
>>> log_variable('total-points', f'week8-homework, 1_homework_q', 75.0)
>>> question_id = 'q8-using the class with new Sine functions-1'
>>> max_score = 1.0
>>> score = 0
>>> seed = np.random.seed(42)
>>> assert isinstance(sine_func, NoisyFunction)
>>> assert sine_func.amplitude == 2
>>> assert sine_func.frequency == 0.2
>>> assert sine_func.phase == 0.5
>>> assert sine_func.func == SineFunction
>>> if isinstance(sine_func, NoisyFunction) and sine_func.amplitude == 2 and (sine_func.frequency == 0.2) and (sine_func.phase == 0.5) and (sine_func.func == SineFunction):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: sine_func implemented incorrectly
success_message: sine_func implemented correctly
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q8-using the class with new Sine functions-2'
>>> max_score = 0.0
>>> score = 0
>>> seed = np.random.seed(42)
>>> assert seed == np.random.seed(42)
>>> if seed == np.random.seed(42):
...     score = 0.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: You might have deleted the line that sets the random seed. Please add it at the top of the cell. You will fail the autograder if you do not place the seed in the top of the cell.
success_message: Your seed is still correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q8-using the class with new Sine functions-3'
>>> max_score = 1.0
>>> score = 0
>>> seed = np.random.seed(42)
>>> assert np.isclose(sum([x, y, y_true])[4], 4.384520747622161, rtol=0.01)
>>> assert np.isclose(np.mean(sum([x, y, y_true])[4::3]), 9.50518884876772, rtol=0.01)
>>> if np.isclose(sum([x, y, y_true])[4], 4.384520747622161, rtol=0.01) and np.isclose(np.mean(sum([x, y, y_true])[4::3]), 9.50518884876772, rtol=0.01):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: x, y, y_true return incorrectly
success_message: x, y, y_true return correctly
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q8-using the class with new Sine functions-4'
>>> max_score = 4.0
>>> score = 0
>>> seed = np.random.seed(42)
>>> assert np.isclose(popt[0], 2.2550781375831015, rtol=0.01)
>>> assert np.isclose(popt[1], 0.19762745754501546, rtol=0.01)
>>> if np.isclose(popt[0], 2.2550781375831015, rtol=0.01) and np.isclose(popt[1], 0.19762745754501546, rtol=0.01):
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Fitting function incorrectly implemented
success_message: Fitting function correctly implemented
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q8-using the class with new Sine functions-5'
>>> max_score = 5.0
>>> score = 0
>>> seed = np.random.seed(42)
>>> assert np.isclose(np.sum(plot.axes[0].get_lines()[0].get_xydata()), 944.5788167212941, rtol=0.01)
>>> assert np.isclose(np.sum(plot.axes[1].get_lines()[0].get_xydata()), 943.8768017567118, rtol=0.01)
>>> if np.isclose(np.sum(plot.axes[0].get_lines()[0].get_xydata()), 944.5788167212941, rtol=0.01) and np.isclose(np.sum(plot.axes[1].get_lines()[0].get_xydata()), 943.8768017567118, rtol=0.01):
...     score = 5.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: plot incorrectly implemented
success_message: plot implemented correctly


Tests for q9-using the class with new exponential functions (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(11.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(75.0)
>>> log_variable('total-points', f'week8-homework, 1_homework_q', 75.0)
>>> question_id = 'q9-using the class with new exponential functions-1'
>>> max_score = 1.0
>>> score = 0
>>> seed = np.random.seed(42)
>>> assert isinstance(exp_func, NoisyFunction)
>>> assert exp_func.a == 2
>>> assert exp_func.b == 0.2
>>> assert exp_func.func == ExponentialFunction
>>> if isinstance(exp_func, NoisyFunction) and exp_func.a == 2 and (exp_func.b == 0.2) and (exp_func.func == ExponentialFunction):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: exp_func implemented incorrectly
success_message: exp_func implemented correctly
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q9-using the class with new exponential functions-2'
>>> max_score = 0.0
>>> score = 0
>>> assert seed == np.random.seed(42)
>>> if seed == np.random.seed(42):
...     score = 0.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: You might have deleted the line that sets the random seed. Please add it at the top of the cell. You will fail the autograder if you do not place the seed in the top of the cell.
success_message: Your seed is still correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q9-using the class with new exponential functions-3'
>>> max_score = 1.0
>>> score = 0
>>> seed = np.random.seed(42)
>>> assert np.isclose(sum([x, y, y_true])[4], 5.068497807066033, rtol=0.01)
>>> assert np.isclose(np.mean(sum([x, y, y_true])[4::3]), 54.848917146243046, rtol=0.01)
>>> if np.isclose(sum([x, y, y_true])[4], 5.068497807066033, rtol=0.01) and np.isclose(np.mean(sum([x, y, y_true])[4::3]), 54.848917146243046, rtol=0.01):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: x, y, y_true return incorrectly
success_message: x, y, y_true return correctly
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q9-using the class with new exponential functions-4'
>>> max_score = 4.0
>>> score = 0
>>> seed = np.random.seed(42)
>>> assert np.isclose(popt[0], 1.9635667563774726, rtol=0.01)
>>> assert np.isclose(popt[1], 0.20098217114795966, rtol=0.01)
>>> if np.isclose(popt[0], 1.9635667563774726, rtol=0.01) and np.isclose(popt[1], 0.20098217114795966, rtol=0.01):
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Fitting function incorrectly implemented
success_message: Fitting function correctly implemented
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q9-using the class with new exponential functions-5'
>>> max_score = 5.0
>>> score = 0
>>> seed = np.random.seed(42)
>>> assert np.isclose(np.sum(plot.axes[0].get_lines()[0].get_xydata()), 3212.76954915798, rtol=0.01)
>>> assert np.isclose(np.sum(plot.axes[1].get_lines()[0].get_xydata()), 3220.812931722522, rtol=0.01)
>>> if np.isclose(np.sum(plot.axes[0].get_lines()[0].get_xydata()), 3212.76954915798, rtol=0.01) and np.isclose(np.sum(plot.axes[1].get_lines()[0].get_xydata()), 3220.812931722522, rtol=0.01):
...     score = 5.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: plot incorrectly implemented
success_message: plot implemented correctly
