```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("10_dictionaries_q", "week_2", "readings", assignment_points = 11.0, assignment_tag = 'week2-readings')

# Initialize Otter
import otter
grader = otter.Notebook("10_dictionaries_q.ipynb")
```

# â“ Dictionaries and Search Engine Optimization (SEO) ðŸ‘€

![](./assets/figures/let-me-google-that.webp)

Search Engine Optimization (SEO) is the practice of improving a website to increase its visibility in search engines like Google. It involves tracking various metrics (measurements) about how people interact with your website:


```{admonition} What is SEO?
Search Engine Optimization (SEO) is the practice of improving a website to increase its visibility in search engines like Google. It involves tracking various metrics (measurements) about how people interact with your website:

- **Pageviews**: Number of times people view your pages ðŸ‘€
- **Bounce Rate**: Percentage of visitors who leave without interacting ðŸšª
- **Keywords**: Words people use to find your site ðŸ”
- **Rankings**: How high your site appears in search results ðŸ“Š
```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._10_dictionaries_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._10_dictionaries_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._10_dictionaries_q import Question3
Question3().show()

```

## Website Analytics Exercise ðŸ“Š

### Instructions:

Let's practice working with dictionaries by creating a simple analytics dashboard! 

```{admonition} Context
We'll create a function that manages website metrics. Don't worry about the SEO terms - focus on practicing dictionary operations!
- Pageviews = Number of times people view your pages
- Bounce rate = Percentage of visitors who leave quickly (shown as decimal)
- Session duration = How long visitors stay (in seconds)
```

1. Create a function that:
   - Creates a dictionary called `metrics` with these key-value pairs:
     - "pageviews": 1000
     - "bounce_rate": 0.35
     - "duration": 180
   - Adds a new key "date" with value "2024-03-15"
   - Calculates the bounce rate as a percentage (multiply by 100)
   - Returns the dictionary and the percentage

2. The function should print:
   ```
   Website Metrics for 2024-03-15:
   - Pageviews: 1000
   - Bounce Rate: 35%
   - Average Duration: 180 seconds
   ```

   We have provided this code for you.

```python
def website_analytics():
    # Create the `metrics` dictionary
    
    ...
    
    # Add the date
    ...
    
    
    # Calculate percentage bounce rate and save it to a variable called `bounce_rate_percent`
    ...
    
    return metrics, bounce_rate_percent

# Call the function
metrics, bounce_rate_percent = website_analytics()

# We have provided this code for you, you don't need to modify it
# Print the summary
print(f"Website Metrics for {metrics['date']}:")
print(f"- Pageviews: {metrics['pageviews']}")
print(f"- Bounce Rate: {int(bounce_rate_percent)}%")
print(f"- Average Duration: {metrics['duration']} seconds")
    
```

```python
grader.check("question-Website-Analytics")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week2-readings", "10_dictionaries_q")
```

Tests for question-Website-Analytics (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(6.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(6.0)
>>> log_variable('total-points', f'Reading-Week-X, 10_dictionaries_q', 6.0)
>>> question_id = 'question-Website-Analytics-1'
>>> max_score = 3.0
>>> score = 0
>>> metrics, _ = website_analytics()
>>> dictionary = {'pageviews': 1000, 'bounce_rate': 0.35, 'duration': 180, 'date': '2024-03-15'}
>>> assert metrics == dictionary, 'Metrics dictionary values or structure is incorrect.'
>>> if metrics == dictionary:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_dictionaries_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The metrics dictionary is incorrect.
success_message: Correctly created and modified the metrics dictionary.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Website-Analytics-2'
>>> max_score = 3.0
>>> score = 0
>>> _, bounce_rate_percent = website_analytics()
>>> assert bounce_rate_percent == 35.0, 'Bounce rate percentage calculation is incorrect.'
>>> if bounce_rate_percent == 35.0:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_dictionaries_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The bounce rate percentage calculation is incorrect.
success_message: Correctly calculated the bounce rate percentage.
