```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("2_strings_q", "week_2", "readings", assignment_points = 15.0, assignment_tag = 'week2-readings')

# Initialize Otter
import otter
grader = otter.Notebook("2_strings_q.ipynb")
```

# ❓Pootie Tang's Wild String Manipulation Adventures

![](./assets/figures/pootie.jpg)

<iframe width="1280" height="738" src="https://www.youtube.com/embed/yhBExhldRXQ" title="Pootie Tang - Trailer" frameborder="0" allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share" referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>"

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._2_strings_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._2_strings_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._2_strings_q import Question3
Question3().show()

```

## Pootie Tang’s Legendary Quotes

![](./assets/figures/pootie-dirty.gif)

### Instructions:
1. Write a Python program that:
    - Stores the following phrases in a list: `"Wa Da Tah!", "Sa Da Tay!", "Sine your pitty"`. List are defined using square brackets `[]` with each element separated by a comma.For example: `["element1", "element2", "element3"]`.
    - Joins the phrases into a single string, separated by a comma and a space. You should use the built in string method `join()`.
    - Prints the resulting string.


```python
phrases = ...
result = ...
print(result)
```

```python
grader.check("Pootie-Tangs-Famous-Quotes")
```

## Pootie Tang’s Famous Quotes

### Instructions:
Pootie Tang’s iconic phrase "Wa Da Tah!" is already attention-grabbing, but let’s make it LOUDER by converting it to uppercase.

1. Write Python code to convert "Wa Da Tah!" into all uppercase letters.
- `upper()` is a string method that converts all lowercase letters in a string to uppercase.
- This ensures Pootie’s catchphrase gets the volume it deserves—perfect for a high-energy scene in the movie.
- save the original string to the variable `phrase` and the uppercase version to the variable `phrase_upper`.

```python
def catchphrase():
        
    phrase = ...
    upper_phrase = ...
    print(upper_phrase)
    
catchphrase()
```

```python
grader.check("coding-uppercase-pootie-tang")
```

## Pootie Tang’s is Awesome

### Instructions:
PootieTang always knows how to tweak his phrases for maximum impact. Suppose the string "PootieTang is cool" needs an upgrade to "PootieTang is awesome".

1. Replace the word "cool" with "awesome" using Python.
- The `replace(old, new)` method in Python replaces all occurrences of `old` with `new` in a string.
- This allows Pootie Tang to adapt his style on the fly, just as he does in the film when he updates his approach to saving the day.


```python
def replace_name():
        
    pt_phrase = "PootieTang is cool"
    modfied_phrase = ...
    print(modfied_phrase)

replace_name()
```

```python
grader.check("coding-replace-pootie-tang")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week2-readings", "2_strings_q")
```

Tests for Pootie-Tangs-Famous-Quotes (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(2.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(10.0)
>>> log_variable('total-points', f'Reading-Week-X, 2_strings_q', 10.0)
>>> question_id = 'Pootie-Tangs-Famous-Quotes-1'
>>> max_score = 2.0
>>> score = 0
>>> assert result == 'Wa Da Tah!, Sa Da Tay!, Sine your pitty', 'The output string is incorrect.'
>>> if result == 'Wa Da Tah!, Sa Da Tay!, Sine your pitty':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_strings_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The `dining_center` variable is incorrect.
success_message: Correctly assigned the `dining_center` variable.


Tests for coding-replace-pootie-tang (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import io
>>> import sys
>>> max_question_points = str(4.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(10.0)
>>> log_variable('total-points', f'Reading-Week-X, 2_strings_q', 10.0)
>>> question_id = 'coding-replace-pootie-tang-1'
>>> max_score = 1.0
>>> score = 0
>>> captured_output = io.StringIO()
>>> sys.stdout = captured_output
>>> replace_name()
>>> sys.stdout = sys.__stdout__
>>> printed_output = captured_output.getvalue().strip()
>>> expected_output = 'PootieTang is awesome'
>>> assert printed_output == expected_output, f"Expected '{expected_output}', but got '{printed_output}'"
>>> if printed_output == expected_output:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_strings_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The `dining_center` variable is incorrect.
success_message: Correctly assigned the `dining_center` variable.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import inspect
>>> import ast
>>> question_id = 'coding-replace-pootie-tang-2'
>>> max_score = 3.0
>>> score = 0
>>> def check_method_usage_in_function(func, method_name):
...     source_code = inspect.getsource(func)
...     tree = ast.parse(source_code)
...     class MethodUsageChecker(ast.NodeVisitor):
...         def __init__(self, method_name):
...             self.method_name = method_name
...             self.is_used = False
...         def visit_Attribute(self, node):
...             if isinstance(node, ast.Attribute) and node.attr == self.method_name:
...                 self.is_used = True
...             self.generic_visit(node)
...     checker = MethodUsageChecker(method_name)
...     checker.visit(tree)
...     return checker.is_used
>>> method_name_to_check = 'replace'
>>> is_replace_used = check_method_usage_in_function(replace_name, method_name_to_check)
>>> assert is_replace_used, f"The method '{method_name_to_check}' is not used in the function."
>>> if is_replace_used:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_strings_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: You did not use the `replace` method in the function.
success_message: Correctly used the `replace` method in the function.


Tests for coding-uppercase-pootie-tang (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import io
>>> import sys
>>> max_question_points = str(4.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(10.0)
>>> log_variable('total-points', f'Reading-Week-X, 2_strings_q', 10.0)
>>> question_id = 'coding-uppercase-pootie-tang-1'
>>> max_score = 1.0
>>> score = 0
>>> captured_output = io.StringIO()
>>> sys.stdout = captured_output
>>> catchphrase()
>>> sys.stdout = sys.__stdout__
>>> printed_output = captured_output.getvalue().strip()
>>> expected_output = 'WA DA TAH!'
>>> assert printed_output == expected_output, f"Expected '{expected_output}', but got '{printed_output}'"
>>> if printed_output == expected_output:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_strings_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: print statement is incorrect.
success_message: print statement is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import inspect
>>> import ast
>>> question_id = 'coding-uppercase-pootie-tang-2'
>>> max_score = 3.0
>>> score = 0
>>> def check_method_usage_in_function(func, method_name):
...     source_code = inspect.getsource(func)
...     tree = ast.parse(source_code)
...     class MethodUsageChecker(ast.NodeVisitor):
...         def __init__(self, method_name):
...             self.method_name = method_name
...             self.is_used = False
...         def visit_Attribute(self, node):
...             if isinstance(node, ast.Attribute) and node.attr == self.method_name:
...                 self.is_used = True
...             self.generic_visit(node)
...     checker = MethodUsageChecker(method_name)
...     checker.visit(tree)
...     return checker.is_used
>>> method_name_to_check = 'upper'
>>> is_upper_used = check_method_usage_in_function(catchphrase, method_name_to_check)
>>> assert is_upper_used, f"The method '{method_name_to_check}' is not used in the function."
>>> if is_upper_used:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_strings_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: You did not use the `upper` method in the function.
success_message: Correctly used the `upper` method in the function.
