```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("4_datatypes_q", "week_2", "readings", assignment_points = 13.0, assignment_tag = 'week2-readings')

# Initialize Otter
import otter
grader = otter.Notebook("4_datatypes_q.ipynb")
```

# ❓ Python Datatypes: The Periodic Table Edition 🧪

<iframe class="pubchem-widget" src="https://pubchem.ncbi.nlm.nih.gov/periodic-table/#view=table&embed=true" style="width: 100%; height: 1000px;"></iframe>



```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._4_datatypes_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._4_datatypes_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._4_datatypes_q import Question3
Question3().show()

```

## Chemical Element Data Structure 🔬

### Instructions:

Create a Python program that organizes information about chemical elements efficiently using appropriate data types. Your solution should:

1. Create a dictionary for the element Sodium (Na) with the following properties:

   - `symbol`: String "Na"
   - `atomic_number`: Integer 11
   - `atomic_mass`: Float 22.990
   - `electron_config`: String "1s2 2s2 2p6 3s1"
   - `period`: Integer 3
   - `group`: Integer 1
   - `properties`: List containing strings ["alkali metal", "reactive", "soft"]

2. Create a dictionary called `electron_shells` that stores the electron configuration by shell:
   - Shell K: Tuple (1, "1s2")
   - Shell L: Tuple (2, "2s2 2p6")
   - Shell M: Tuple (3, "3s1")

   The keys for the dictionary should be the shell names (e.g., "K", "L", "M").

Note: You do not need to write the exponentiation for the electron configuration (e.g., "2s2" instead of "2s²").


```python
def create_sodium_data():
    # Define the element dictionary
    sodium = ...

    # Define the electron shells dictionary
    electron_shells = ...

    return sodium, electron_shells

# We have provided this code for you
# Please do not modify the following code beyond this line
# Get the data structures
sodium, electron_shells = create_sodium_data()

# Print the structures to verify
print("Sodium Element Data:")
for key, value in sodium.items():
    print(f"{key}: {value}")

print("\nElectron Shell Configuration:")
for shell, config in electron_shells.items():
    print(f"Shell {shell}: Level {config[0]} - Configuration {config[1]}")
```

```python
grader.check("question-Element-Structure")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week2-readings", "4_datatypes_q")
```

Tests for question-Element-Structure (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(5.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(5.0)
>>> log_variable('total-points', f'Reading-Week-X, 4_datatypes_q', 5.0)
>>> question_id = 'question-Element-Structure-1'
>>> max_score = 0.5
>>> score = 0
>>> sodium, _ = create_sodium_data()
>>> assert isinstance(sodium, dict), 'Sodium dictionary not created correctly.'
>>> if isinstance(sodium, dict):
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_datatypes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Sodium dictionary is incorrectly created.
success_message: Sodium dictionary is correctly created.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Element-Structure-2'
>>> max_score = 0.5
>>> score = 0
>>> assert sodium['symbol'] == 'Na', 'Incorrect symbol'
>>> if sodium['symbol'] == 'Na':
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_datatypes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Symbol is incorrect.
success_message: Symbol is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Element-Structure-3'
>>> max_score = 0.5
>>> score = 0
>>> assert sodium['atomic_number'] == 11, 'Incorrect atomic number'
>>> if sodium['atomic_number'] == 11:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_datatypes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Atomic number is incorrect.
success_message: Atomic number is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Element-Structure-4'
>>> max_score = 0.5
>>> score = 0
>>> assert abs(sodium['atomic_mass'] - 22.99) < 0.001, 'Incorrect atomic mass'
>>> if abs(sodium['atomic_mass'] - 22.99) < 0.001:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_datatypes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Atomic mass is incorrect.
success_message: Atomic mass is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Element-Structure-5'
>>> max_score = 0.5
>>> score = 0
>>> assert sodium['electron_config'] == '1s2 2s2 2p6 3s1', 'Incorrect electron configuration'
>>> if sodium['electron_config'] == '1s2 2s2 2p6 3s1':
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_datatypes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Electron configuration is incorrect.
success_message: Electron configuration is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Element-Structure-6'
>>> max_score = 0.5
>>> score = 0
>>> assert sodium['period'] == 3, 'Incorrect period'
>>> if sodium['period'] == 3:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_datatypes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Period is incorrect.
success_message: Period is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Element-Structure-7'
>>> max_score = 0.5
>>> score = 0
>>> assert sodium['group'] == 1, 'Incorrect group'
>>> if sodium['group'] == 1:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_datatypes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Group is incorrect.
success_message: Group is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Element-Structure-8'
>>> max_score = 0.5
>>> score = 0
>>> description = set(['alkali metal', 'reactive', 'soft'])
>>> assert set(sodium['properties']) == description, 'Incorrect properties'
>>> if set(sodium['properties']) == description:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_datatypes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Properties are incorrect.
success_message: Properties are correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Element-Structure-9'
>>> max_score = 1.0
>>> score = 0
>>> _, electron_shells = create_sodium_data()
>>> k_out = (1, '1s2')
>>> l_out = (2, '2s2 2p6')
>>> m_out = (3, '3s1')
>>> assert electron_shells['K'] == k_out, 'Incorrect K shell configuration'
>>> assert electron_shells['L'] == l_out, 'Incorrect L shell configuration'
>>> assert electron_shells['M'] == m_out, 'Incorrect M shell configuration'
>>> if electron_shells['K'] == k_out and electron_shells['L'] == l_out and (electron_shells['M'] == m_out):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_datatypes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The electron_shells dictionary is missing shells or has incorrect configurations.
success_message: Correctly created the electron_shells dictionary with all required shell configurations.
