```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("6_comments_q", "week_2", "readings", assignment_points = 12.0, assignment_tag = 'week2-readings')

# Initialize Otter
import otter
grader = otter.Notebook("6_comments_q.ipynb")
```

# â“ Commenting

![](./assets/figures/simon.jpg)

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._6_comments_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._6_comments_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._6_comments_q import Question3
Question3().show()

```

## Commenting Simon Says

We have provided you with a function that: 
1. Takes a single argument command, which is a string.
2. Executes one of four predefined actions based on the value of the command string:
    - If the command is "jump", return "Simon says: Jump high!".
    - If the command is "run", return "Simon says: Run fast!".
    - If the command is "stop", return "Simon says: Stop immediately!".
    - If the command is "dance", return "Simon says: Show your moves!".
    - If the command is anything else, return "Simon didn't say that!".

### Instructions:

Include comments within the function to describe each action explicitly. The comments should be structured like this:

```python
if command == "walk":
        # If the command is "walk", Simon says to walk slow.
        return "Simon says: Walk slow!"
```



```python
def simon_says_actions(command):

    # Check the command and respond accordingly
    if command == "jump":
        ...
        return "Simon says: Jump high!"
    elif command == "run":
        ...
        return "Simon says: Run fast!"
    elif command == "stop":
        ...
        return "Simon says: Stop immediately!"
    elif command == "dance":
        ...
        return "Simon says: Show your moves!"
    else:
        # If the command is unrecognized, Simon didn't say that.
        return "Simon didn't say that!"
    
print(simon_says_actions("jump")) # Should return "Simon says
```

```python
grader.check("question-Simon-Says")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week2-readings", "6_comments_q")
```

Tests for question-Simon-Says (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> import inspect
>>> max_question_points = str(4.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(4.0)
>>> log_variable('total-points', f'Reading-Week-X, 6_comments_q', 4.0)
>>> question_id = 'question-Simon-Says-1'
>>> max_score = 1.0
>>> score = 0
>>> def test_jump_comment():
...     source_code = inspect.getsource(simon_says_actions)
...     assert re.search('#.*jump.*simon.*jump.*high', source_code, re.IGNORECASE), "The 'jump' action comment is missing or does not describe the action properly."
...     return source_code
>>> source_code = test_jump_comment()
>>> if re.search('#.*jump.*simon.*jump.*high', source_code, re.IGNORECASE):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_comments_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The 'jump' action comment is missing or incorrect.
success_message: Correctly commented the 'jump' action.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> import inspect
>>> question_id = 'question-Simon-Says-2'
>>> max_score = 1.0
>>> score = 0
>>> def test_run_comment():
...     source_code = inspect.getsource(simon_says_actions)
...     assert re.search('#.*run.*simon.*run.*fast', source_code, re.IGNORECASE), "The 'run' action comment is missing or does not describe the action properly."
...     return source_code
>>> source_code = test_run_comment()
>>> if re.search('#.*run.*simon.*run.*fast', source_code, re.IGNORECASE):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_comments_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The 'run' action comment is missing or incorrect.
success_message: Correctly commented the 'run' action.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> import inspect
>>> question_id = 'question-Simon-Says-3'
>>> max_score = 1.0
>>> score = 0
>>> def test_stop_comment():
...     source_code = inspect.getsource(simon_says_actions)
...     assert re.search('#.*stop.*simon.*stop.*immediately', source_code, re.IGNORECASE), "The 'stop' action comment is missing or does not describe the action properly."
...     return source_code
>>> source_code = test_stop_comment()
>>> if re.search('#.*stop.*simon.*stop.*immediately', source_code, re.IGNORECASE):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_comments_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The 'stop' action comment is missing or incorrect.
success_message: Correctly commented the 'stop' action.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> import inspect
>>> question_id = 'question-Simon-Says-4'
>>> max_score = 1.0
>>> score = 0
>>> def test_dance_comment():
...     source_code = inspect.getsource(simon_says_actions)
...     assert re.search('#.*dance.*simon.*show.*moves', source_code, re.IGNORECASE), "The 'dance' action comment is missing or does not describe the action properly."
...     return source_code
>>> source_code = test_dance_comment()
>>> if re.search('#.*dance.*simon.*show.*moves', source_code, re.IGNORECASE):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_comments_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The 'dance' action comment is missing or incorrect.
success_message: Correctly commented the 'dance' action.
