```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("8_lists_and_tuples_q", "week_2", "readings", assignment_points = 11.0, assignment_tag = 'week2-readings')

# Initialize Otter
import otter
grader = otter.Notebook("8_lists_and_tuples_q.ipynb")
```

# â“ ðŸ”¬  Materials Design: Lists and Tuples

![](./assets/figures/nature-automated.jpg)

{cite:p}`Unknown2023-me`


```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._8_lists_and_tuples_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._8_lists_and_tuples_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._8_lists_and_tuples_q import Question3
Question3().show()

```

## Materials Processing Sequence ðŸ­

### Instructions:

1. Create the following variables:

   - `element_properties`: A tuple storing properties of titanium ("Ti", 22, 47.87)
   - `processing_steps`: A list storing heat treatment steps ["Solution Treatment", "Quenching", "Aging"]
   - `temperatures`: A list of processing temperatures [955, 25, 540]
   - `times`: A list of processing times in hours [1.0, 0.1, 24.0]
   - `total_time`: Use sum() to calculate total processing time
   - `max_temp`: Use max() to find the highest temperature
   - `num_steps`: Calculate the number of processing steps using len()
   - `process_complete`: A boolean indicating if all steps are completed (True)

2. Print a descriptive message that incorporates all the above variables, formatted like this:

   ```
   Processing Ti (Atomic number: 22, Mass: 47.87) using 3 steps:
   Solution Treatment at 955Â°C for 1.0 hours
   Quenching at 25Â°C for 0.1 hours
   Aging at 540Â°C for 24.0 hours
   Total processing time: 25.1 hours
   Maximum temperature: 955Â°C
   Process completion status: True
   ```


```python
def materials_processing():
    # Define the variables
    element_properties = ...
    processing_steps = ...
    temperatures = ...
    times = ...
    total_time = ...
    max_temp = ...
    num_steps = ...
    process_complete = ...

    return (
        element_properties,
        processing_steps,
        temperatures,
        times,
        total_time,
        max_temp,
        num_steps,
        process_complete,
    )

# This code is provided for you to test the function to ensure it is correct
# please do not modify the provided code
# Get all variables
(
    element_properties,
    processing_steps,
    temperatures,
    times,
    total_time,
    max_temp,
    num_steps,
    process_complete,
) = materials_processing()

# Print the descriptive message
print(f"Processing {element_properties[0]} (Atomic number: {element_properties[1]}, Mass: {element_properties[2]}) using {num_steps} steps:")
for step, temp, time in zip(processing_steps, temperatures, times):
    print(f"{step} at {temp}Â°C for {time} hours")
print(f"Total processing time: {total_time} hours")
print(f"Maximum temperature: {max_temp}Â°C")
print(f"Process completion status: {process_complete}")

```

```python
grader.check("question-Materials-Processing-Sequence")
```

## References

```{bibliography}

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week2-readings", "8_lists_and_tuples_q")
```

Tests for question-Materials-Processing-Sequence (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(4.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(4.0)
>>> log_variable('total-points', f'Reading-Week-X, 8_lists_and_tuples_q', 4.0)
>>> question_id = 'question-Materials-Processing-Sequence-1'
>>> max_score = 1.0
>>> score = 0
>>> element_properties, *_ = materials_processing()
>>> assert element_properties == ('Ti', 22, 47.87), 'Variable element_properties is incorrect.'
>>> if element_properties == ('Ti', 22, 47.87):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('8_lists_and_tuples_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The element_properties tuple is incorrect.
success_message: Correctly created the element_properties tuple.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Materials-Processing-Sequence-2'
>>> max_score = 1.0
>>> score = 0
>>> _, _, _, times, total_time, *_ = materials_processing()
>>> assert total_time == sum(times), 'Variable total_time is incorrect.'
>>> if total_time == sum(times):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('8_lists_and_tuples_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The total_time calculation is incorrect.
success_message: Correctly calculated the total processing time.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Materials-Processing-Sequence-3'
>>> max_score = 1.0
>>> score = 0
>>> _, _, temperatures, _, _, max_temp, *_ = materials_processing()
>>> assert max_temp == max(temperatures), 'Variable max_temp is incorrect.'
>>> if max_temp == max(temperatures):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('8_lists_and_tuples_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The max_temp calculation is incorrect.
success_message: Correctly determined the maximum temperature.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Materials-Processing-Sequence-4'
>>> max_score = 1.0
>>> score = 0
>>> _, processing_steps, _, _, _, _, num_steps, _ = materials_processing()
>>> assert num_steps == len(processing_steps), 'Variable num_steps is incorrect.'
>>> if num_steps == len(processing_steps):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('8_lists_and_tuples_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The num_steps calculation is incorrect.
success_message: Correctly counted the number of processing steps.
