```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_week_7_lecture_attendance", "week_7", "lecture", assignment_points = 1.0, assignment_tag = 'week7-lecture')

# Initialize Otter
import otter
grader = otter.Notebook("1_week_7_lecture_attendance.ipynb")
```

# ðŸ—“ï¸ðŸ“š Week 7: Lecture Attendance âœ…

To get credit for attendance make sure you run all blocks of code. 

```python
I_attended_lecture = True
```

```python
grader.check("week-3-lab-attendance")
```

## ðŸŽ‰ Well Done!

Keep up the good work.

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week7-lecture", "1_week_7_lecture_attendance")
```

Tests for week-3-lab-attendance (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(1.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(1.0)
>>> log_variable('total-points', f'week7-lecture, 1_week_7_lecture_attendance', 1.0)
>>> question_id = 'week-3-lab-attendance-1'
>>> max_score = 1.0
>>> score = 0
>>> assert I_attended_lecture, 'You must attend the lab to get credit for this assignment.'
>>> if I_attended_lecture:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_week_7_lecture_attendance', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Output does not include 'Each friend gets'.
success_message: Correctly included the phrase 'Each friend gets'.
