```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_practicequiz_q", "week_7", "practicequiz", assignment_points = 28.0, assignment_tag = 'week7-practicequiz')

# Initialize Otter
import otter
grader = otter.Notebook("1_practicequiz_q.ipynb")
```

# ❓A Manufacturing System using Object-Oriented Programming

![](./assets/figures/subway-assembly.gif)




```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_practicequiz_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_practicequiz_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_practicequiz_q import Question3
Question3().show()

```

## Question 1 (Points: 12.0): Modify an Inventory Class

Given the `Inventory` class below, modify it to:
- Add an `__init__()` method that initializes the `name` and `stock` attributes. The `stock` attribute should default to 0.
- Add an `add_stock(quantity)` method that increases the `stock` attribute.
- Add a `remove_stock(quantity)` method that decreases `stock` but ensures it doesn’t go below zero.
- Implement a `__str__()` method that returns `"Product: <name>, Stock: <stock>"`.

### Example:
```python
item = Inventory("Gears", 50)
item.add_stock(20)
print(item)  # Output: Product: Gears, Stock: 70
item.remove_stock(80)
print(item)  # Output: Product: Gears, Stock: 0
```

```python
...

# Example usage
item = Inventory("Gears", 50)
item.add_stock(20)
print(item)
item.remove_stock(80)
print(item)
```

```python
grader.check("Free-Response-Modify-an-Inventory-Class")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week7-practicequiz", "1_practicequiz_q")
```

Tests for Free-Response-Modify-an-Inventory-Class (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(12.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(12.0)
>>> log_variable('total-points', f'week7-practicequiz, 1_practicequiz_q', 12.0)
>>> question_id = 'Free-Response-Modify-an-Inventory-Class-1'
>>> max_score = 2.0
>>> score = 0
>>> item = Inventory('Gears', 50)
>>> inventory_initialization = item.name == 'Gears' and item.stock == 50
>>> assert inventory_initialization, f"Expected Inventory('Gears', 50), but got name={item.name}, stock={item.stock}."
>>> if inventory_initialization:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicequiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(inventory_initialization))
```
failure_message: Failed: Inventory object does not initialize correctly.
success_message: Success: Inventory object initializes correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Free-Response-Modify-an-Inventory-Class-2'
>>> max_score = 2.0
>>> score = 0
>>> item = Inventory('Bolts')
>>> default_stock_check = item.stock == 0
>>> assert default_stock_check, f'Expected default stock to be 0, but got {item.stock}.'
>>> if default_stock_check:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicequiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(default_stock_check))
```
failure_message: Failed: Default stock value is incorrect.
success_message: Success: Default stock value is correctly set to 0!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Free-Response-Modify-an-Inventory-Class-3'
>>> max_score = 2.0
>>> score = 0
>>> item = Inventory('Nuts', 10)
>>> item.add_stock(20)
>>> stock_after_addition = item.stock
>>> assert stock_after_addition == 30, f'Expected stock to be 30 after addition, but got {item.stock}.'
>>> if stock_after_addition == 30:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicequiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(stock_after_addition))
```
failure_message: Failed: `add_stock()` does not increase stock correctly.
success_message: Success: `add_stock()` correctly increases stock!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Free-Response-Modify-an-Inventory-Class-4'
>>> max_score = 2.0
>>> score = 0
>>> item = Inventory('Bolts', 50)
>>> item.remove_stock(30)
>>> stock_after_removal = item.stock
>>> assert stock_after_removal == 20, f'Expected stock to be 20 after removal, but got {item.stock}.'
>>> if stock_after_removal == 20:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicequiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(stock_after_removal))
```
failure_message: Failed: `remove_stock()` does not decrease stock correctly.
success_message: Success: `remove_stock()` correctly decreases stock!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Free-Response-Modify-an-Inventory-Class-5'
>>> max_score = 2.0
>>> score = 0
>>> item = Inventory('Washers', 40)
>>> item.remove_stock(100)
>>> stock_after_excessive_removal = item.stock
>>> assert stock_after_excessive_removal == 0, f'Expected stock to be 0 (not negative), but got {item.stock}.'
>>> if stock_after_excessive_removal == 0:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicequiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(stock_after_excessive_removal))
```
failure_message: Failed: Stock goes below 0.
success_message: Success: Stock does not drop below 0!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Free-Response-Modify-an-Inventory-Class-6'
>>> max_score = 2.0
>>> score = 0
>>> item = Inventory('Bearings', 15)
>>> string_representation = str(item)
>>> expected_string = 'Product: Bearings, Stock: 15'
>>> assert string_representation == expected_string, f"Expected '{expected_string}', but got '{string_representation}'."
>>> if string_representation == expected_string:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicequiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(string_representation))
```
failure_message: Failed: `__str__()` does not return the expected format.
success_message: Success: `__str__()` returns the correct format!
