```python
from pykubegrader.tokens.validate_token import validate_token
validate_token('type the key provided by your instructor here')

```

```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("3_if_else_q", "week_4", "lecture", assignment_points = 8.0, assignment_tag = 'week4-lecture')

# Initialize Otter
import otter
grader = otter.Notebook("3_if_else_q.ipynb")
```

# 💻 Activity 5.1: Nuclear Plant Temperature Alert

![Nuclear Reactor](./assets/figures/nuclear-reactor-inside.jpg)

We have just covered the relationship of a process value to a desired set point and discussed the importance of redundancy, especially in systems that have the potential to harm people.

## Problem Statement

Consider a nuclear reactor. A meltdown would be sufficiently catastrophic that the temperature of the cooling water around the nuclear fuel rods is monitored by four sensors.

An alarm should go off if:

1. **Condition 1**: Any two sensor readings differ by strictly more than 8 degrees.
2. **Condition 2**: Any of the two sensors report a temperature greater than 100 degrees Celsius.

Write a function `coolingWaterAlarm(t1, t2, t3, t4)` where `t1`, `t2`, `t3`, and `t4` are the temperature readings for sensor 1, sensor 2, sensor 3, and sensor 4, respectively.

The function should:
    - Print `Alarm! Temperatures differ.` if the first condition is met.
    - Print `Alarm! Temperatures exceed 100 deg Celsius.` if the second condition is met.
    - Print `Normal operation.` if neither condition is met.

## Implementation

1. We have provided the function for you to use. 
2. Convert the temperature readings to a list, `tlist`.
3. Check if the difference between the maximum and minimum values of `tlist` is strictly greater than 8. If so, print the first alarm message. You should use the `max()` and `min()` functions.
4. Check if the maximum value of `tlist` is greater than 100. If so, print the second alarm message.
5. If neither condition is met, print the normal operation message.


```python
def coolingWaterAlarm(t1, t2, t3, t4):
    ...


# Test the function
coolingWaterAlarm(93, 95, 96, 101)
coolingWaterAlarm(95, 103, 96, 101)
coolingWaterAlarm(90, 95, 96, 101)
coolingWaterAlarm(90, 95, 96, 92)
```

```python
grader.check("cooling-water-problems")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.tokens.validate_token import validate_token
validate_token()


from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week4-lecture", "3_if_else_q")
```

Tests for cooling-water-problems (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> from io import StringIO
>>> import sys
>>> max_question_points = str(8.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(8.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token()
>>> log_variable('total-points', f'Reading-Week-X, 3_if_else_q', 8.0)
>>> question_id = 'cooling-water-problems-1'
>>> max_score = 2.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> coolingWaterAlarm(93, 103, 96, 101)
>>> sys.stdout = sys.stdout
>>> output_value = output.getvalue().strip()
>>> condition = re.search('Alarm! Temperatures differ\\.', output_value, re.IGNORECASE)
>>> assert condition, 'The program did not correctly identify large temperature differences.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_if_else_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(output))
```
failure_message: Failed: The program did not correctly identify large temperature differences.
success_message: Success: The program correctly identifies large temperature differences.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> from io import StringIO
>>> import sys
>>> question_id = 'cooling-water-problems-2'
>>> max_score = 2.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> coolingWaterAlarm(95, 103, 96, 101)
>>> sys.stdout = sys.stdout
>>> output_value = output.getvalue().strip()
>>> condition = re.search('Alarm! Temperatures exceed 100 deg Celsius\\.', output_value, re.IGNORECASE)
>>> assert condition, 'The program did not correctly identify temperatures exceeding 100 degrees.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_if_else_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(output))
```
failure_message: Failed: The program did not correctly identify temperatures exceeding 100 degrees.
success_message: Success: The program correctly identifies temperatures exceeding 100 degrees.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> from io import StringIO
>>> import sys
>>> question_id = 'cooling-water-problems-3'
>>> max_score = 2.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> coolingWaterAlarm(90, 95, 96, 92)
>>> sys.stdout = sys.stdout
>>> output_value = output.getvalue().strip()
>>> print(output_value)
>>> condition = re.search('Normal operation\\.', output_value, re.IGNORECASE)
>>> assert condition, 'The program did not correctly identify normal operation.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_if_else_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(output))
```
failure_message: Failed: The program did not correctly identify normal operation.
success_message: Success: The program correctly handles normal operation.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> from io import StringIO
>>> import sys
>>> question_id = 'cooling-water-problems-4'
>>> max_score = 2.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> coolingWaterAlarm(95, 105, 96, 91)
>>> sys.stdout = sys.stdout
>>> output_value = output.getvalue().strip()
>>> condition_diff = re.search('Alarm! Temperatures differ\\.', output_value, re.IGNORECASE)
>>> assert condition_diff, 'The program did not correctly prioritize temperature difference alarms.'
>>> if condition_diff:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_if_else_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(output))
```
failure_message: Failed: The program did not correctly handle mixed conditions.
success_message: Success: The program correctly handles mixed conditions.
