```python
from pykubegrader.tokens.validate_token import validate_token
validate_token('type the key provided by your instructor here')

```

```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("9_loops_activity_q", "week_4", "lecture", assignment_points = 8.0, assignment_tag = 'week4-lecture')

# Initialize Otter
import otter
grader = otter.Notebook("9_loops_activity_q.ipynb")
```

# 💻 Activity: Taking Integrals with Loops

![Riemann Sums](https://math.libretexts.org/@api/deki/files/6076/5.3.2.png?revision=1&size=bestfit&width=441&height=378)

## Problem Statement

Let's explore how two important choices impact the accuracy of this method for approximating integrals:

1. The width of the rectangle $h$, which is determined by the number of rectangles between $a$ and $b$.
2. The value of $x$ at which the function is evaluated (left, midpoint, or right).

## Implementation

Write a function `riemann(numRectangles, evalPoint)` that calculates the approximate integral of a given function using Riemann sums with the following requirements:

1. Define the boundaries of the integral: $a = 0$ and $b = 5$.
2. Use the curve $f(x) = \sin(x/2) + 1$.
   - Create an array of $x$ values between $a$ and $b$, using `np.linspace()`. The number of steps is determined by `numRectangles`.
3. Calculate the width of each rectangle, $h$, as:
   $$h = \frac{b - a}{\text{numRectangles} - 1}$$
4. Initialize the variable `integral_R` to zero.
5. Use a loop to calculate the integral using the value of $f$ at the right side of each rectangle. This is if i < (numRectangles - 1).
6. Add and assign the area of each rectangle to `integral_R` as:
   $$\text{integral_R} += h \times f[i + 1]$$

We have provided the necessary print statement and test statements to tests your code. 

## Example Usage

Call the function using different values for `numRectangles` (e.g., 2, 5, and 11) and `"right"` as the evaluation point to compare the approximations.

```python
import numpy as np


def riemann(numRectangles, evalPoint):
    # The function riemann takes two inputs: numRectangles and evalPoint
    # numRectangles is the number of rectangles to use in the Riemann sum

    # Define the boundaries of the integral
    ...

    # Define the curve
    x = ...
    f = ...

    # Determine the width of each rectangle
    h = ...

    # Initialize the integral
    integral_R = ...

    # Integrate using the value of f at the right
    ...

    # Print the integral, number of rectangles, and the chosen point for evaluation
    print(
        f"The integral is {integral_R:.6f} for {numRectangles} using the {evalPoint} point for evaluating the function."
    )


# Call the function with different inputs
riemann(2, "right")
riemann(5, "right")
riemann(11, "right")
riemann(51, "right")
```

```python
grader.check("integrals-riemann-sums")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.tokens.validate_token import validate_token
validate_token()


from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week4-lecture", "9_loops_activity_q")
```

Tests for integrals-riemann-sums (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> from io import StringIO
>>> import sys
>>> max_question_points = str(8.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(8.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token()
>>> log_variable('total-points', f'Reading-Week-X, 9_loops_activity_q', 8.0)
>>> question_id = 'integrals-riemann-sums-1'
>>> max_score = 2.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> riemann(2, 'right')
>>> sys.stdout = sys.__stdout__
>>> output_value = output.getvalue().strip()
>>> condition = re.search('The integral is\\s+7\\.\\d{6,9}\\s+for\\s+2\\s+using\\s+the\\s+right', output_value)
>>> assert condition, 'The program did not calculate the integral correctly for 2 rectangles.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('9_loops_activity_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(output))
```
failure_message: Failed: The program did not calculate the integral correctly for 2 rectangles.
success_message: Success: The program correctly calculates the integral for 2 rectangles.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> from io import StringIO
>>> import sys
>>> question_id = 'integrals-riemann-sums-2'
>>> max_score = 2.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> riemann(5, 'right')
>>> sys.stdout = sys.__stdout__
>>> output_value = output.getvalue().strip()
>>> condition = re.search('The integral is\\s+8\\.85\\d{4,6}\\s+for\\s+5\\s+using\\s+the\\s+right', output_value)
>>> assert condition, 'The program did not calculate the integral correctly for 5 rectangles.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('9_loops_activity_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(output))
```
failure_message: Failed: The program did not calculate the integral correctly for 5 rectangles.
success_message: Success: The program correctly calculates the integral for 5 rectangles.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> from io import StringIO
>>> import sys
>>> question_id = 'integrals-riemann-sums-3'
>>> max_score = 2.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> riemann(11, 'right')
>>> sys.stdout = sys.__stdout__
>>> output_value = output.getvalue().strip()
>>> condition = re.search('The integral is\\s+8\\.73\\d{2,6}\\s+for\\s+11\\s+using\\s+the\\s+right', output_value)
>>> assert condition, 'The program did not calculate the integral correctly for 11 rectangles.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('9_loops_activity_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(output))
```
failure_message: Failed: The program did not calculate the integral correctly for 11 rectangles.
success_message: Success: The program correctly calculates the integral for 11 rectangles.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> from io import StringIO
>>> import sys
>>> question_id = 'integrals-riemann-sums-4'
>>> max_score = 2.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> riemann(51, 'right')
>>> sys.stdout = sys.__stdout__
>>> output_value = output.getvalue().strip()
>>> condition = re.search('The integral is\\s+8\\.631\\d{2,6}\\s+for\\s+51\\s+using\\s+the\\s+right', output_value)
>>> assert condition, 'The program did not improve precision with more rectangles.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('9_loops_activity_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(output))
```
failure_message: Failed: The program did not improve precision with more rectangles.
success_message: Success: The program correctly calculates the integral with increased precision.
