```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("10_args_kwargs_q", "week_5", "readings", assignment_points = 30.0, assignment_tag = 'week5-readings')

# Initialize Otter
import otter
grader = otter.Notebook("10_args_kwargs_q.ipynb")
```

# ❓ `*args` and `**kwargs` in Brawl Stars

![](./assets/figures/brawl-stars-characters.webp)

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._10_args_kwargs_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._10_args_kwargs_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._10_args_kwargs_q import Question3
Question3().show()

```

## Write a Function to Display Brawler Stats

Write a Python function named `display_brawler_stats` that:
- Accepts a brawler's name as the first parameter.
- Accepts additional statistics as keyword arguments (`**kwargs`).
- Prints the brawler's name and their statistics in the format:
    ```
    Name: <name>
    Health: <value>
    Power: <value>
    ```
- Use the string method `capitalize()` to capitalize the the stat which you are describing

```python
...
    
# Example usage
display_brawler_stats("Shelly", health=3600, power=10)
```

```python
grader.check("Brawler-Stats-question")
```

## **Brawler's Survival Challenge**

![](./assets/figures/josh-brawl.jpeg)

In **Brawl Stars**, characters take damage from enemy attacks and receive healing from allies or abilities. Your task is to implement a system that calculates a **brawler's remaining hitpoints** after taking damage and receiving healing.

### **Instructions**
You need to implement the following **three functions**:

1. **`apply_damage(hitpoints, damage)`**  
   - **Parameters**: 
     - `hitpoints` (int): The brawler’s current HP.
     - `damage` (int): The amount of damage taken.
   - **Returns**: The new HP after taking damage. If HP falls below **zero**, it should return **zero**. This can be done several ways. One way is to use the `max()` function to return the maximum of 0 and the result of the subtraction.

2. **`apply_healing(hitpoints, healing, max_hitpoints)`**  
   - **Parameters**:  
     - `hitpoints` (int): The brawler’s current HP.
     - `healing` (int): The amount of healing received.
     - `max_hitpoints` (int): The brawler’s maximum HP.

    Make sure healing does not happen it the Brawler has been eliminated (HP = 0), or if the healing would exceed the maximum HP.
   - **Returns**: The new HP after healing, but it **cannot exceed** `max_hitpoints`. You can use the `min()` function to return the minimum of `max_hitpoints` and the result of the addition.

3. **`calculate_remaining_hp(starting_hp, max_hp, damage_taken, healing_received)`**  
   - Calls the first two functions to compute the **final** HP of the brawler.
   - **Parameters**:  
     - `starting_hp` (int): The brawler’s initial HP.
     - `max_hp` (int): The brawler’s maximum HP.
     - `damage_taken` (int): The total damage received.
     - `healing_received` (int): The total healing received.
   - **Returns**: The final HP after applying damage and healing.


### **Example Usage**
```python
# Shelly starts with 3600 HP and takes 1200 damage, but receives 500 healing.
remaining_hp = calculate_remaining_hp(3600, 3600, 1200, 500)
print(remaining_hp)  # Expected Output: 2900
```

```python
# Function to apply damage to a brawler's HP
...


# Function to apply healing to a brawler's HP
...


# Function to calculate remaining HP after applying damage and healing
...
```

```python
grader.check("Brawler-Damage-Stats")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week5-readings", "10_args_kwargs_q")
```

Tests for Brawler-Damage-Stats (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(14.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(21.0)
>>> log_variable('total-points', f'Reading-Week-X, 10_args_kwargs_q', 21.0)
>>> question_id = 'Brawler-Damage-Stats-1'
>>> max_score = 2.0
>>> score = 0
>>> hitpoints = 3600
>>> damage = 1200
>>> expected_output = 2400
>>> actual_output = apply_damage(hitpoints, damage)
>>> assert actual_output == expected_output, f'Expected {expected_output}, but got {actual_output}.'
>>> if (actual_output == expected_output, f'Expected {expected_output}, but got {actual_output}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_args_kwargs_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(hitpoints))
>>> responses = update_responses(question_id, str(damage))
>>> responses = update_responses(question_id, str(expected_output))
>>> responses = update_responses(question_id, str(actual_output))
```
failure_message: Failed: Damage calculation is incorrect.
success_message: Success: Damage is applied correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Brawler-Damage-Stats-2'
>>> max_score = 2.0
>>> score = 0
>>> hitpoints = 1000
>>> damage = 1500
>>> expected_output = 0
>>> actual_output = apply_damage(hitpoints, damage)
>>> assert actual_output == expected_output, f'Expected {expected_output}, but got {actual_output}.'
>>> if (actual_output == expected_output, f'Expected {expected_output}, but got {actual_output}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_args_kwargs_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(hitpoints))
>>> responses = update_responses(question_id, str(damage))
>>> responses = update_responses(question_id, str(expected_output))
>>> responses = update_responses(question_id, str(actual_output))
```
failure_message: Failed: HP can go below zero.
success_message: Success: Damage does not drop below zero!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Brawler-Damage-Stats-3'
>>> max_score = 2.0
>>> score = 0
>>> hitpoints = 2000
>>> healing = 500
>>> max_hitpoints = 3600
>>> expected_output = 2500
>>> actual_output = apply_healing(hitpoints, healing, max_hitpoints)
>>> assert actual_output == expected_output, f'Expected {expected_output}, but got {actual_output}.'
>>> if (actual_output == expected_output, f'Expected {expected_output}, but got {actual_output}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_args_kwargs_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(hitpoints))
>>> responses = update_responses(question_id, str(healing))
>>> responses = update_responses(question_id, str(max_hitpoints))
>>> responses = update_responses(question_id, str(expected_output))
>>> responses = update_responses(question_id, str(actual_output))
```
failure_message: Failed: Healing calculation is incorrect.
success_message: Success: Healing is applied correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Brawler-Damage-Stats-4'
>>> max_score = 2.0
>>> score = 0
>>> hitpoints = 3500
>>> healing = 500
>>> max_hitpoints = 3600
>>> expected_output = 3600
>>> actual_output = apply_healing(hitpoints, healing, max_hitpoints)
>>> assert actual_output == expected_output, f'Expected {expected_output}, but got {actual_output}.'
>>> if (actual_output == expected_output, f'Expected {expected_output}, but got {actual_output}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_args_kwargs_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(hitpoints))
>>> responses = update_responses(question_id, str(healing))
>>> responses = update_responses(question_id, str(max_hitpoints))
>>> responses = update_responses(question_id, str(expected_output))
>>> responses = update_responses(question_id, str(actual_output))
```
failure_message: Failed: Healing allows HP to exceed the max.
success_message: Success: Healing does not exceed max HP!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Brawler-Damage-Stats-5'
>>> max_score = 3.0
>>> score = 0
>>> starting_hp = 3600
>>> max_hp = 3600
>>> damage_taken = 1200
>>> healing_received = 500
>>> expected_output = 2900
>>> actual_output = calculate_remaining_hp(starting_hp, max_hp, damage_taken, healing_received)
>>> assert actual_output == expected_output, f'Expected {expected_output}, but got {actual_output}.'
>>> if (actual_output == expected_output, f'Expected {expected_output}, but got {actual_output}.'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_args_kwargs_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(starting_hp))
>>> responses = update_responses(question_id, str(max_hp))
>>> responses = update_responses(question_id, str(damage_taken))
>>> responses = update_responses(question_id, str(healing_received))
>>> responses = update_responses(question_id, str(expected_output))
>>> responses = update_responses(question_id, str(actual_output))
```
failure_message: Failed: The function does not correctly compute remaining HP.
success_message: Success: The function correctly calculates remaining HP!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Brawler-Damage-Stats-6'
>>> max_score = 3.0
>>> score = 0
>>> starting_hp = 1000
>>> max_hp = 1000
>>> damage_taken = 2000
>>> healing_received = 500
>>> expected_output = 0
>>> actual_output = calculate_remaining_hp(starting_hp, max_hp, damage_taken, healing_received)
>>> assert actual_output == expected_output, f'Expected {expected_output}, but got {actual_output}.'
>>> if (actual_output == expected_output, f'Expected {expected_output}, but got {actual_output}.'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_args_kwargs_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(starting_hp))
>>> responses = update_responses(question_id, str(max_hp))
>>> responses = update_responses(question_id, str(damage_taken))
>>> responses = update_responses(question_id, str(healing_received))
>>> responses = update_responses(question_id, str(expected_output))
>>> responses = update_responses(question_id, str(actual_output))
```
failure_message: Failed: HP drops below zero.
success_message: Success: HP does not go negative!


Tests for Brawler-Stats-question (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import io
>>> import sys
>>> max_question_points = str(7.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(21.0)
>>> log_variable('total-points', f'Reading-Week-X, 10_args_kwargs_q', 21.0)
>>> question_id = 'Brawler-Stats-question-1'
>>> max_score = 1.0
>>> score = 0
>>> name = 'Shelly'
>>> kwargs = {'health': 3600, 'power': 10}
>>> condition = 'display_brawler_stats' in globals()
>>> if condition:
...     output_capture = io.StringIO()
...     sys.stdout = output_capture
...     display_brawler_stats(name, **kwargs)
...     sys.stdout = sys.__stdout__
>>> assert condition, 'The function raised an error when called.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_args_kwargs_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function encounters an error when called.
success_message: Success: The function executes without errors!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import io
>>> import sys
>>> question_id = 'Brawler-Stats-question-2'
>>> max_score = 1.0
>>> score = 0
>>> name = 'Shelly'
>>> kwargs = {'health': 3600, 'power': 10}
>>> expected_output = 'Name: Shelly'
>>> output_capture = io.StringIO()
>>> sys.stdout = output_capture
>>> display_brawler_stats(name, **kwargs)
>>> sys.stdout = sys.__stdout__
>>> actual_output = output_capture.getvalue().strip().split('\n')[0]
>>> assert actual_output == expected_output, f'Expected output:\n{expected_output}\nBut got:\n{actual_output}'
>>> if (actual_output == expected_output, f'Expected output:\n{expected_output}\nBut got:\n{actual_output}'):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_args_kwargs_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not correctly print the name.
success_message: Success: The function correctly prints the name!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import io
>>> import sys
>>> question_id = 'Brawler-Stats-question-3'
>>> max_score = 1.0
>>> score = 0
>>> name = 'Shelly'
>>> kwargs = {'health': 3600, 'power': 10}
>>> expected_output = 'Name: Shelly'
>>> expected_stat_output = 'Health: 3600'
>>> output_capture = io.StringIO()
>>> sys.stdout = output_capture
>>> display_brawler_stats(name, **kwargs)
>>> sys.stdout = sys.__stdout__
>>> actual_output_lines = output_capture.getvalue().strip().split('\n')
>>> condition = any((line == expected_stat_output for line in actual_output_lines))
>>> assert condition, f"Expected to find '{expected_stat_output}' in the output, but it was missing."
>>> if (condition, f"Expected to find '{expected_stat_output}' in the output, but it was missing."):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_args_kwargs_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not correctly print a stat.
success_message: Success: The function correctly prints at least one stat!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import io
>>> import sys
>>> question_id = 'Brawler-Stats-question-4'
>>> max_score = 1.0
>>> score = 0
>>> name = 'Shelly'
>>> kwargs = {'health': 3600, 'power': 10}
>>> expected_output = 'Name: Shelly'
>>> expected_stat_output = 'Health: 3600'
>>> output_capture = io.StringIO()
>>> sys.stdout = output_capture
>>> display_brawler_stats(name, **kwargs)
>>> sys.stdout = sys.__stdout__
>>> expected_capitalization = ['Health: 3600', 'Power: 10']
>>> actual_output_lines = output_capture.getvalue().strip().split('\n')
>>> condition = all((stat in actual_output_lines for stat in expected_capitalization))
>>> assert condition, f'Expected stats with correct capitalization:\n{expected_capitalization}\nBut got:\n{actual_output_lines}'
>>> if (condition, f'Expected stats with correct capitalization:\n{expected_capitalization}\nBut got:\n{actual_output_lines}'):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_args_kwargs_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not capitalize stats correctly.
success_message: Success: The function correctly capitalizes stats!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import io
>>> import sys
>>> question_id = 'Brawler-Stats-question-5'
>>> max_score = 2.0
>>> score = 0
>>> name = 'Shelly'
>>> kwargs = {'health': 3600, 'power': 10}
>>> expected_output = 'Name: Shelly'
>>> expected_stat_output = 'Health: 3600'
>>> output_capture = io.StringIO()
>>> sys.stdout = output_capture
>>> display_brawler_stats(name, **kwargs)
>>> sys.stdout = sys.__stdout__
>>> expected_output = 'Name: Shelly\nHealth: 3600\nPower: 10'
>>> actual_output = output_capture.getvalue().strip()
>>> assert actual_output == expected_output, f'Expected output:\n{expected_output}\nBut got:\n{actual_output}'
>>> if (actual_output == expected_output, f'Expected output:\n{expected_output}\nBut got:\n{actual_output}'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_args_kwargs_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not print multiple stats correctly.
success_message: Success: The function correctly prints multiple stats!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import io
>>> import sys
>>> question_id = 'Brawler-Stats-question-6'
>>> max_score = 1.0
>>> score = 0
>>> output_capture = io.StringIO()
>>> sys.stdout = output_capture
>>> display_brawler_stats('Colt')
>>> sys.stdout = sys.__stdout__
>>> expected_output = 'Name: Colt'
>>> actual_output = output_capture.getvalue().strip()
>>> assert actual_output == expected_output, f'Expected output:\n{expected_output}\nBut got:\n{actual_output}'
>>> if (actual_output == expected_output, f'Expected output:\n{expected_output}\nBut got:\n{actual_output}'):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_args_kwargs_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not handle empty stats correctly.
success_message: Success: The function correctly handles empty stats!
