```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("3_flow_control_q", "week_5", "readings", assignment_points = 16.0, assignment_tag = 'week5-readings')

# Initialize Otter
import otter
grader = otter.Notebook("3_flow_control_q.ipynb")
```

# ❓ Functions and Flow Control in the Mall

![](./assets/figures/Edith-Macefields-House-1.jpg)

Edith Macefield became a symbol of defiance and resilience when she refused a $1 million offer to sell her small, humble home in Seattle’s Ballard neighborhood to make way for a commercial development. Developers planned to construct a sprawling shopping center but faced an immovable obstacle in Macefield, who chose to stay in the house where she had lived for decades. Her steadfastness turned her into a local legend, and her story inspired comparisons to the movie Up. Rather than forcing her out, the developers built the five-story complex around her home, leaving it dwarfed by the surrounding structure. Macefield’s determination to remain true to her roots captured the public's imagination, making her a beloved figure and a poignant reminder of the value of standing up for what matters most to you.

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._3_flow_control_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._3_flow_control_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._3_flow_control_q import Question3
Question3().show()

```

## Write a Function to Calculate Total Revenue
The mall is analyzing its revenue from parking fees. Write a Python function `calculate_total_revenue` that:

- Accepts a list of hours parked by multiple cars.

Calculates the total revenue, using the following parking fee rules:

- \$5 for up to 2 hours.
- \$10 for 3-5 hours.
- \$15 for more than 5 hours.
- Returns the total revenue.


```python
...

# Example usage
hours_parked = [1, 3, 6, 2, 5]
print(calculate_total_revenue(hours_parked)) # Should print 45
```

```python
grader.check("Mall-calculation-total-revenue")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week5-readings", "3_flow_control_q")
```

Tests for Mall-calculation-total-revenue (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(7.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(7.0)
>>> log_variable('total-points', f'Reading-Week-X, 3_flow_control_q', 7.0)
>>> question_id = 'Mall-calculation-total-revenue-1'
>>> max_score = 1.0
>>> score = 0
>>> input_hours = [1, 2, 2, 1]
>>> expected_revenue = len(input_hours) * 5
>>> calculated_revenue = calculate_total_revenue(input_hours)
>>> condition = calculated_revenue == expected_revenue
>>> assert condition, f'Expected revenue: {expected_revenue}, but got: {calculated_revenue} for input {input_hours}.'
>>> if (condition, f'Expected revenue: {expected_revenue}, but got: {calculated_revenue} for input {input_hours}.'):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_flow_control_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(input_hours))
>>> responses = update_responses(question_id, str(expected_revenue))
>>> responses = update_responses(question_id, str(calculated_revenue))
```
failure_message: Failed: The function does not correctly calculate revenue for short parking durations.
success_message: Success: The function correctly calculates revenue for short parking durations!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Mall-calculation-total-revenue-2'
>>> max_score = 1.0
>>> score = 0
>>> input_hours = [3, 4, 5, 4]
>>> expected_revenue = len(input_hours) * 10
>>> calculated_revenue = calculate_total_revenue(input_hours)
>>> condition = calculated_revenue == expected_revenue
>>> assert condition, f'Expected revenue: {expected_revenue}, but got: {calculated_revenue} for input {input_hours}.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_flow_control_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(input_hours))
>>> responses = update_responses(question_id, str(expected_revenue))
>>> responses = update_responses(question_id, str(calculated_revenue))
```
failure_message: Failed: The function does not correctly calculate revenue for medium parking durations.
success_message: Success: The function correctly calculates revenue for medium parking durations!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Mall-calculation-total-revenue-3'
>>> max_score = 1.0
>>> score = 0
>>> input_hours = [3, 4, 5, 4]
>>> expected_revenue = len(input_hours) * 10
>>> calculated_revenue = calculate_total_revenue(input_hours)
>>> condition = calculated_revenue == expected_revenue
>>> assert condition, f'Expected revenue: {expected_revenue}, but got: {calculated_revenue} for input {input_hours}.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_flow_control_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(input_hours))
>>> responses = update_responses(question_id, str(expected_revenue))
>>> responses = update_responses(question_id, str(calculated_revenue))
```
failure_message: Failed: The function does not correctly calculate revenue for medium parking durations.
success_message: Success: The function correctly calculates revenue for medium parking durations!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Mall-calculation-total-revenue-4'
>>> max_score = 1.0
>>> score = 0
>>> input_hours = [6, 7, 8, 10]
>>> expected_revenue = len(input_hours) * 15
>>> calculated_revenue = calculate_total_revenue(input_hours)
>>> condition = calculated_revenue == expected_revenue
>>> assert condition, f'Expected revenue: {expected_revenue}, but got: {calculated_revenue} for input {input_hours}.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_flow_control_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(input_hours))
>>> responses = update_responses(question_id, str(expected_revenue))
>>> responses = update_responses(question_id, str(calculated_revenue))
```
failure_message: Failed: The function does not correctly calculate revenue for long parking durations.
success_message: Success: The function correctly calculates revenue for long parking durations!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Mall-calculation-total-revenue-5'
>>> max_score = 2.0
>>> score = 0
>>> input_hours = [1, 3, 6, 2, 5]
>>> expected_revenue = 5 + 10 + 15 + 5 + 10
>>> calculated_revenue = calculate_total_revenue(input_hours)
>>> condition = calculated_revenue == expected_revenue
>>> assert condition, f'Expected revenue: {expected_revenue}, but got: {calculated_revenue} for input {input_hours}.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_flow_control_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(input_hours))
>>> responses = update_responses(question_id, str(expected_revenue))
>>> responses = update_responses(question_id, str(calculated_revenue))
```
failure_message: Failed: The function does not correctly calculate revenue for mixed parking durations.
success_message: Success: The function correctly calculates revenue for mixed parking durations!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Mall-calculation-total-revenue-6'
>>> max_score = 1.0
>>> score = 0
>>> input_hours = []
>>> expected_revenue = 0
>>> calculated_revenue = calculate_total_revenue(input_hours)
>>> condition = calculated_revenue == expected_revenue
>>> assert condition, f'Expected revenue: {expected_revenue}, but got: {calculated_revenue} for input {input_hours}.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('3_flow_control_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(input_hours))
>>> responses = update_responses(question_id, str(expected_revenue))
>>> responses = update_responses(question_id, str(calculated_revenue))
```
failure_message: Failed: The function does not correctly calculate revenue for an empty list.
success_message: Success: The function correctly calculates revenue for an empty list!
