```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("5_shrinkflation_q", "week_5", "readings", assignment_points = 15.0, assignment_tag = 'week5-readings')

# Initialize Otter
import otter
grader = otter.Notebook("5_shrinkflation_q.ipynb")
```

# â“ Shrinkflation Function Scope 

![](./assets/figures/shrinkflation1.png)

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._5_shrinkflation_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._5_shrinkflation_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._5_shrinkflation_q import Question3
Question3().show()

```

## Question 1 (Points: 6.0): Calculate the Shrinkflation Price

![](./assets/figures/shrinkflation2.jpg)

Write a Python function named `adjust_price` that:
- Accepts two arguments: `original_price` (float) and `shrink_factor` (float).
- Calculates the adjusted price by multiplying `original_price` by `shrink_factor`.
- Returns the adjusted price.
Ensure that the function uses local variables only and does not rely on any global variables.

```python
...

print(adjust_price(10.0, 0.8)) # Should print 8.0
```

```python
grader.check("Shrinkflation-Price")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week5-readings", "5_shrinkflation_q")
```

Tests for Shrinkflation-Price (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(6.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(6.0)
>>> log_variable('total-points', f'Reading-Week-X, 5_shrinkflation_q', 6.0)
>>> question_id = 'Shrinkflation-Price-1'
>>> max_score = 1.0
>>> score = 0
>>> input_original_price = 10.0
>>> input_shrink_factor = 0.8
>>> expected_price = 8.0
>>> calculated_price = adjust_price(input_original_price, input_shrink_factor)
>>> condition = calculated_price == expected_price
>>> assert condition, f'Expected price: {expected_price}, but got: {calculated_price} for input ({input_original_price}, {input_shrink_factor}).'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_shrinkflation_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not correctly adjust the price with a basic shrink factor.
success_message: Success: The function correctly adjusts the price with a basic shrink factor!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Shrinkflation-Price-2'
>>> max_score = 1.0
>>> score = 0
>>> input_original_price = 50.0
>>> input_shrink_factor = 1.0
>>> expected_price = 50.0
>>> calculated_price = adjust_price(input_original_price, input_shrink_factor)
>>> condition = calculated_price == expected_price
>>> assert condition, f'Expected price: {expected_price}, but got: {calculated_price} for input ({input_original_price}, {input_shrink_factor}).'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_shrinkflation_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not correctly handle a shrink factor of 1.
success_message: Success: The function correctly handles a shrink factor of 1 (no adjustment)!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Shrinkflation-Price-3'
>>> max_score = 2.0
>>> score = 0
>>> input_original_price = 20.0
>>> input_shrink_factor = 0.0
>>> expected_price = 0.0
>>> calculated_price = adjust_price(input_original_price, input_shrink_factor)
>>> condition = calculated_price == expected_price
>>> assert condition, f'Expected price: {expected_price}, but got: {calculated_price} for input ({input_original_price}, {input_shrink_factor}).'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_shrinkflation_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not correctly handle a shrink factor of 0.
success_message: Success: The function correctly handles a shrink factor of 0 (price adjusted to 0)!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Shrinkflation-Price-4'
>>> max_score = 2.0
>>> score = 0
>>> input_original_price = 15.0
>>> input_shrink_factor = 1.5
>>> expected_price = 22.5
>>> calculated_price = adjust_price(input_original_price, input_shrink_factor)
>>> condition = calculated_price == expected_price
>>> assert condition, f'Expected price: {expected_price}, but got: {calculated_price} for input ({input_original_price}, {input_shrink_factor}).'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_shrinkflation_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not correctly handle a shrink factor greater than 1.
success_message: Success: The function correctly handles a shrink factor greater than 1!
