```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("8_arguments_and_return_q", "week_5", "readings", assignment_points = 20.0, assignment_tag = 'week5-readings')

# Initialize Otter
import otter
grader = otter.Notebook("8_arguments_and_return_q.ipynb")
```

# â“ Norman Doors - How, "Not" to Design Arguments, Parameters, and Return Statements 

![](./assets/figures/norman_door_2.png)

A **Norman door** is a term for a poorly designed door that confuses users about whether to push or pull, often due to unclear visual or tactile cues. Coined by Don Norman in *The Design of Everyday Things*, it highlights the failure of the design to communicate its function intuitively, often requiring labels like "PUSH" or "PULL" to compensate. Good design avoids such confusion by aligning the door's appearance and affordances (e.g., handles or flat plates) with its intended operation, ensuring usability without explanation.

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._8_arguments_and_return_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._8_arguments_and_return_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._8_arguments_and_return_q import Question3
Question3().show()

```

## Create a Function to Evaluate Door Usability

![](./assets/figures/norman-doors.jpg)

Write a Python function named `evaluate_door` that:

- Accepts two arguments: `handle_type` (a string) and `door_action` (a string).
- Returns `"Confusing"` if:
    - The handle type is "pull" but the action is "push".
    - The handle type is "push" but the action is "pull".
- Returns `"Clear"` if the handle type matches the action.


```python
...

print(evaluate_door("pull", "push")) # Should print Confusing
print(evaluate_door("push", "push")) # Should print Clear
```

```python
grader.check("Door-Assessment")
```

## Create a Function to Recommend Door Fixes

![](./assets/figures/norman-3.avif)

Write a Python function named `recommend_fix` that:
- Accepts two arguments: `handle_type` (a string) and `door_action` (a string).
- Returns a string recommendation:
    - `"Change the handle"` if the handle type does not match the action.
    - `"No changes needed"` if the handle type matches the action.


```python
...

# Example usage
print(recommend_fix("pull", "push")) # Should print Change the handle
print(recommend_fix("push", "push")) # Should print No changes needed
```

```python
grader.check("Door-Recommendations")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week5-readings", "8_arguments_and_return_q")
```

Tests for Door-Assessment (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(6.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(11.0)
>>> log_variable('total-points', f'Reading-Week-X, 8_arguments_and_return_q', 11.0)
>>> question_id = 'Door-Assessment-1'
>>> max_score = 1.0
>>> score = 0
>>> handle_type = 'pull'
>>> door_action = 'push'
>>> expected_output = 'Confusing'
>>> actual_output = evaluate_door(handle_type, door_action)
>>> condition = actual_output == expected_output
>>> assert condition, f"Expected '{expected_output}' for input ({handle_type}, {door_action}), but got '{actual_output}'."
>>> if (condition, f"Expected '{expected_output}' for input ({handle_type}, {door_action}), but got '{actual_output}'."):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('8_arguments_and_return_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not correctly identify a pull handle with a push action as 'Confusing'.
success_message: Success: The function correctly identifies a pull handle with a push action as 'Confusing'!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Door-Assessment-2'
>>> max_score = 1.0
>>> score = 0
>>> handle_type = 'push'
>>> door_action = 'pull'
>>> expected_output = 'Confusing'
>>> actual_output = evaluate_door(handle_type, door_action)
>>> condition = actual_output == expected_output
>>> assert condition, f"Expected '{expected_output}' for input ({handle_type}, {door_action}), but got '{actual_output}'."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('8_arguments_and_return_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not correctly identify a push handle with a pull action as 'Confusing'.
success_message: Success: The function correctly identifies a push handle with a pull action as 'Confusing'!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Door-Assessment-3'
>>> max_score = 1.0
>>> score = 0
>>> handle_type = 'pull'
>>> door_action = 'pull'
>>> expected_output = 'Clear'
>>> actual_output = evaluate_door(handle_type, door_action)
>>> condition = actual_output == expected_output
>>> assert condition, f"Expected '{expected_output}' for input ({handle_type}, {door_action}), but got '{actual_output}'."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('8_arguments_and_return_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not correctly identify a pull handle with a pull action as 'Clear'.
success_message: Success: The function correctly identifies a pull handle with a pull action as 'Clear'!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Door-Assessment-4'
>>> max_score = 1.0
>>> score = 0
>>> handle_type = 'push'
>>> door_action = 'push'
>>> expected_output = 'Clear'
>>> actual_output = evaluate_door(handle_type, door_action)
>>> condition = actual_output == expected_output
>>> assert condition, f"Expected '{expected_output}' for input ({handle_type}, {door_action}), but got '{actual_output}'."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('8_arguments_and_return_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not correctly identify a push handle with a push action as 'Clear'.
success_message: Success: The function correctly identifies a push handle with a push action as 'Clear'!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Door-Assessment-5'
>>> max_score = 1.0
>>> score = 0
>>> handle_type = 'rotate'
>>> door_action = 'turn'
>>> expected_output = 'Clear'
>>> actual_output = evaluate_door(handle_type, door_action)
>>> condition = actual_output == expected_output
>>> assert condition, f"Expected '{expected_output}' for input ({handle_type}, {door_action}), but got '{actual_output}'."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('8_arguments_and_return_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not correctly handle invalid inputs.
success_message: Success: The function correctly handles invalid inputs!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Door-Assessment-6'
>>> max_score = 1.0
>>> score = 0
>>> handle_type = 'rotate'
>>> door_action = 'turn'
>>> expected_output = 'Clear'
>>> actual_output = evaluate_door(handle_type, door_action)
>>> condition = actual_output == expected_output
>>> assert condition, f"Expected '{expected_output}' for input ({handle_type}, {door_action}), but got '{actual_output}'."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('8_arguments_and_return_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not correctly handle invalid inputs.
success_message: Success: The function correctly handles invalid inputs!


Tests for Door-Recommendations (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(5.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(11.0)
>>> log_variable('total-points', f'Reading-Week-X, 8_arguments_and_return_q', 11.0)
>>> question_id = 'Door-Recommendations-1'
>>> max_score = 1.0
>>> score = 0
>>> handle_type = 'pull'
>>> door_action = 'push'
>>> expected_output = 'Change the handle'
>>> actual_output = recommend_fix(handle_type, door_action)
>>> condition = actual_output == expected_output
>>> assert condition, f"Expected '{expected_output}' for input ({handle_type}, {door_action}), but got '{actual_output}'."
>>> if (condition, f"Expected '{expected_output}' for input ({handle_type}, {door_action}), but got '{actual_output}'."):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('8_arguments_and_return_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(handle_type))
>>> responses = update_responses(question_id, str(door_action))
>>> responses = update_responses(question_id, str(expected_output))
>>> responses = update_responses(question_id, str(actual_output))
```
failure_message: Failed: The function does not correctly recommend 'Change the handle' for mismatched input (pull, push).
success_message: Success: The function correctly recommends 'Change the handle' for mismatched input (pull, push)!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Door-Recommendations-2'
>>> max_score = 1.0
>>> score = 0
>>> handle_type = 'push'
>>> door_action = 'pull'
>>> expected_output = 'Change the handle'
>>> actual_output = recommend_fix(handle_type, door_action)
>>> condition = actual_output == expected_output
>>> assert condition, f"Expected '{expected_output}' for input ({handle_type}, {door_action}), but got '{actual_output}'."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('8_arguments_and_return_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(handle_type))
>>> responses = update_responses(question_id, str(door_action))
>>> responses = update_responses(question_id, str(expected_output))
>>> responses = update_responses(question_id, str(actual_output))
```
failure_message: Failed: The function does not correctly recommend 'Change the handle' for mismatched input (push, pull).
success_message: Success: The function correctly recommends 'Change the handle' for mismatched input (push, pull)!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Door-Recommendations-3'
>>> max_score = 1.0
>>> score = 0
>>> handle_type = 'pull'
>>> door_action = 'pull'
>>> expected_output = 'No changes needed'
>>> actual_output = recommend_fix(handle_type, door_action)
>>> condition = actual_output == expected_output
>>> assert condition, f"Expected '{expected_output}' for input ({handle_type}, {door_action}), but got '{actual_output}'."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('8_arguments_and_return_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(handle_type))
>>> responses = update_responses(question_id, str(door_action))
>>> responses = update_responses(question_id, str(expected_output))
>>> responses = update_responses(question_id, str(actual_output))
```
failure_message: Failed: The function does not correctly recommend 'No changes needed' for matched input (pull, pull).
success_message: Success: The function correctly recommends 'No changes needed' for matched input (pull, pull)!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Door-Recommendations-4'
>>> max_score = 1.0
>>> score = 0
>>> handle_type = 'push'
>>> door_action = 'push'
>>> expected_output = 'No changes needed'
>>> actual_output = recommend_fix(handle_type, door_action)
>>> condition = actual_output == expected_output
>>> assert condition, f"Expected '{expected_output}' for input ({handle_type}, {door_action}), but got '{actual_output}'."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('8_arguments_and_return_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(handle_type))
>>> responses = update_responses(question_id, str(door_action))
>>> responses = update_responses(question_id, str(expected_output))
>>> responses = update_responses(question_id, str(actual_output))
```
failure_message: Failed: The function does not correctly recommend 'No changes needed' for matched input (push, push).
success_message: Success: The function correctly recommends 'No changes needed' for matched input (push, push)!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Door-Recommendations-5'
>>> max_score = 1.0
>>> score = 0
>>> handle_type = 'rotate'
>>> door_action = 'turn'
>>> expected_output = 'Change the handle'
>>> actual_output = recommend_fix(handle_type, door_action)
>>> condition = actual_output == expected_output
>>> assert condition, f"Expected '{expected_output}' for input ({handle_type}, {door_action}), but got '{actual_output}'."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('8_arguments_and_return_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(handle_type))
>>> responses = update_responses(question_id, str(door_action))
>>> responses = update_responses(question_id, str(expected_output))
>>> responses = update_responses(question_id, str(actual_output))
```
failure_message: Failed: The function does not correctly handle invalid inputs.
success_message: Success: The function correctly handles invalid inputs and recommends 'Change the handle'!
