```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("Lab_4_Battleship", "week_4", "lab", assignment_points = 22.0, assignment_tag = 'week4-lab')

# Initialize Otter
import otter
grader = otter.Notebook("Lab_4_Battleship.ipynb")
```

# 🎮 LAB 4 : BATTLESHIP 🎮

Complete the following 4 functions and corresponding function calls which facilitate playing battleship in the Jupyter console. 

Note: bullet points represent tasks you must do to complete the lab.

## 🎲 Set up the game board 🎲


- Add the following doc string to `initialize_board()`. It goes between the pair of `"""`  at the top of the function (replacing `...`). 
```python
    """
    Initialize a game board for tracking ships or shots.

    :param size: The size of the game board (size x size).
    :return: A grid filled with '🌊' to represent water.
    """
```

None of the next functions are called in the cells in which they are defined by default. You may add calls to run the functions inside their cells if you like, or use the ones in **5: Play the game** below. 

`initialize_board()` is given to you in its entirety (minus the doc string above). The operations it performs use a for loop to execute a program designed to generate data for a game board. 

The for loop in `initialize_board()` uses an iterator `_` generated by `range(size)`, an iterable that produces size values. The `_` acts as a placeholder. It will automatically be replaced with `🌊` as the following two lines are executed. 

Each time the loop runs, a row of `🌊` symbols is created and appended to the board object. The function returns the complete `size` × `size` grid, called `board`. 

```python
def initialize_board(size):
    """
    ...
    """
    board = []
    for _ in range(size):
        row = ["🌊"] * size
        board.append(row)
    return board
```

```python
grader.check("task1-initialize_board")
```

## 🚢 Place ships on the board 🚢


```python
    """
    Place ships on the board.

    :param board: The game board.
    :param coordinates: List of (row, col) tuples for ship positions (1-based indexing).
    """
```
### Instructions for Placing Ships on the Board:

Now that we have a game board filled with `🌊` to represent water, we want to replace some of those `🌊` with `"🚢"` to represent ships placed by a player.

#### Step 1: Extract Row and Column Indices from Coordinates
To iterate through the list of ship coordinates (`coordinates`), you will assign each tuple's values to two variables, `row` and `col`. These variables will represent the row and column indices for each part of the ship's placement.

- write the for loop around the coordinates, the  `for ... in coordinates` loop with `row, col` so that each tuple in `coordinates` is unpacked into `row` (for the row index) and `col` (for the column index).

**Example:**  
If `coordinates = [(1, 2), (2, 2)]`, on the first iteration:
- `row` will be assigned `1`
- `col` will be assigned `2`

#### Step 2: Adjust Row and Column for Python's Indexing

Since this game requires human input and we want the board to count rows and columns starting from `(1, 1)` (upper leftmost element), we need to adjust the indices from the 1-based input to Python's 0-based indexing.

- Use a compound assignment operator (`-=`) to subtract `1` from `row` and update its value. Replace the `...` in `row -= ...` with `1`.
- Similarly, do the same for `col` by replacing `...` in `col -= ...` with `1`.

**Why?**  
This adjustment ensures that when you use `board[row][col]`, it correctly points to the intended element on the board.

#### Step 3: Place the Ship on the Board
Now that the `row` and `col` variables have been adjusted for Python's indexing:
- Use `board[row][col] = "🚢"` to mark the corresponding cell in the `board` with `"🚢"`.

When this line runs, it will:
- Look into the 2D list `board` at the index specified by `[row][col]`
- Replace the existing `"🌊"` with `"🚢"` to represent part of a ship.

```python
def place_ship(board, coordinates):
    """
    ...
    """
    ...
        ...
        ...
        board[row][col] = "🚢"
```

```python
grader.check("task2-place_ship")
```

## ⬜⬜⬜Display the board ⬜⬜⬜



```python
    """
    Display the board with colors.

    """
```
### Simplified Instructions:

Now that we have functions to generate and manage the game board, we want to display the board in the console. The board will eventually show four types of data: water, ships, hits, and misses.

- Inside each function, the board will use single string characters (`"🌊"`, `"🚢"`, `"🎯"`, `"💥"`) to represent water, ships, hits, and misses, respectively.

The `display_board()` function:
1. Uses **two nested loops** to iterate through each element of the board.
2. Performs **four logical tests** to check the type of each board element.
3. Appends formatted strings to a temporary list (`color_dat`) for display.

### Steps to Implement:

1. **Define `color_dat`:**
   - Inside the outer loop `for row in board:`, initialize a blank list called `color_dat` to store formatted data for the row.
   - Replace the `...` with `color_dat = []`.

2. **Perform Logical Tests:**
   - Inside the inner loop `for element in row:`, add string characters to each logical test based on the board's data.
   - Use the following mappings:
     - `"🌊"`: Water
     - `"🚢"`: Ship
     - `"🎯"`: Hit
     - `"💥"`: Miss

3. **Add Formatting Characters:**
   - Use the escape sequences below in the respective `color_dat.append(...)` calls nested in the `if` and `elif` logical tests:
     ```python
     color_dat.append("🌊")  # Blue for water
     color_dat.append("🚢")  # 🚢 for ships
     color_dat.append("🎯")  # 💥 for hits
     color_dat.append("💥")  # 💥 for misses
     ```

4. **Print the Board:**
   - Once the `color_dat` list is complete for a row, print it to display the formatted board. The function should return color_dat.

```python
def display_board(board):
    """
    ...
    """
    for row in board:
        ...
        for element in row:
            if element == "🌊":
                ...
            elif element == "🚢":
                ...
            elif element == "🎯":
                ...
            elif element == "💥":
                ...
        print(" ".join(color_dat))
    return color_dat
```

```python
grader.check("task3-display_board")
```

## 💥💥💥💥 Fire away 💥💥💥💥


```python
    """
    Fire a shot at the board and return the result.

    :param board: The game board.
    :param target: A tuple (row, col) where the shot is fired (1-based indexing).
    :return: The result of the shot: "Hit", "Miss", or "Already Fired".
    """

```

Here's the revised version of the instructions for clarity and conciseness:

---

### Instructions for Completing `fire_shot()`

We want to manipulate the `board` using indices provided in the tuple `target`. Here's how to break it down step-by-step:

1. **Extract Row and Column from Target:**
   - Define two variables, `row` and `col`, to store the values contained in `target`. Replace the first `...` in the code with this step.
   
2. **Adjust for Human-Friendly Input:**
   - Since `target` uses 1-based indexing (human-readable), but Python uses 0-based indexing, decrement `row` and `col` by 1. Replace the `...` in the next two lines with the appropriate adjustments.

3. **Determine the Action Based on the Board State:**
   - Analyze the current value of `board[row][col]`:
     - If it is `"🎯"`, this means the space has already been fired at. Replace the `...` with a return statement: `"Already Fired"`.
     - If it contains `"🚢"`, it represents a ship. Update the space to `"🎯"` (to mark a hit) and return `"Hit"`.
     - If it contains `"🌊"`, it represents water. Update the space to `"💥"` (to mark a miss) and return `"Miss"`.
     - For any other value, it is an invalid input. Replace the `...` with a return statement: `"Invalid Target"`.

4. **Fill in Logical Expressions:**
   - Complete the 4 logical checks and actions as described above by replacing the 9 `...` placeholders in the provided code.

```python
def fire_shot(board, target):
    """
    ...
    """
    # Unpack the tuple
    ...
    
    ...

    ...
    else:
        return "Invalid Target"   
```

```python
grader.check("task4-fire_shot")
```

## 🎮 Play the game 🎮

The variables `player_board` and `tracking_board` are populated with data using the initialize board function to produce 10 x 10 game boards filled with water.

To play battleship, you should place 5 ships on your board. These ships are: carrier (occupies 5 spaces), battleship (occupies 4), submarine (occupies 3), destroyer (occupies 3 spaces) and patrol boat (occupies 2). The destroyer and patrol boat have already been placed for you with `place_ship`. 

Keep your screen hidden from your opponent so they don't know where your ships are!

- Edit the placement of the destroyer and patrol boat so your opponent doesn't know where they are
     - Or try to be sneaky and keep them in their default position as the one place your opponent won't check?
- Add 3 new calls to `place_ship` to place the reamining 3 ships
- Run the cell to save your ship placement and view the board

You can run the cell as many times as you like before the game begins to help make sure you are placing ships where you think you're placing them, but **don't re-initialize the board once you've started to play or you'll lose all progress**.

Ships should be continuous and not overlap. They can run on diagonals. 



```python
# Initialize boards
player_board = initialize_board(10)
tracking_board = initialize_board(10)

# Place remaining 3 ships on the board
place_ship(player_board, [(2, 2), (2, 3), (2, 4)]) #destroyer
place_ship(player_board, [(7, 7), (8, 8)]) #patrol boat

# Display board
out = display_board(player_board)
```

Once both you and your oponent have boards set up you can use the following cell to target spaces on their board (and vice versa). You can also use the tracking function below to record your shots on your opponent's board (unless you can keep track of all that in your head?!) 

- Player 1 says: "5,5"
- Player 2 runs: result = `fire_shot(player_board, (5, 5))` and `display_board(player_board)`
- Player 2 says: "Miss!"
- Optional: Player 1 runs: `record_opponent_feedback(tracking_board)` and provides input to the box, then runs `display_board(tracking_board)`

```python
# Fire shots 
result = fire_shot(player_board, (2, 3))
print(result)  # Should print "Hit"

# Display the updated board
out = display_board(player_board)
```

## 💥💥💥💥💥💥 | Track your shots on your opponent's board |  💥💥💥💥💥💥 |

You want to make sure you're not aiming at the same place twice.

Run the following cell to edit the tracking board.

```python
def record_opponent_feedback(tracking_board):
    """
    Prompt the user for a target and result, then update the tracking board.

    :param tracking_board: The board tracking your shots on the opponent's board.
    """
    try:
        # Get input for target coordinates as a tuple
        target_input = input("row,col: ")
        target = tuple(map(int, target_input.split(",")))

        # Get input for result ("Hit" or "Miss")
        result = input("hit or miss?: ").capitalize()
        
        # Update the tracking board
        row, col = target
        row -= 1  # Convert to 0-based index
        col -= 1  # Convert to 0-based index

        if result == "Hit":
            tracking_board[row][col] = "🎯"  # Mark hit
        elif result == "Miss":
            tracking_board[row][col] = "💥"  # Mark miss
        else:
            print("Invalid result. Please enter 'hit' or 'miss'.")
    except ValueError:
        print("Invalid input. Please ensure coordinates are in the format row,col.")

tracking_board = initialize_board(size=10)
```

```python
#Running this following cell will prompt for input. 

#First, tell your opponent your shot and input coordinates as a comma separated pair of integers and hit enter. 

#Then, input "hit" or "miss" (no quotes) depending on what your opponent says.

record_opponent_feedback(tracking_board)
display_board(tracking_board)
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week4-lab", "Lab_4_Battleship")
```

Tests for task1-initialize_board (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(2.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(22.0)
>>> log_variable('total-points', f'Reading-Week-X, Lab_4_Battleship', 22.0)
>>> question_id = 'task1-initialize_board-1'
>>> max_score = 2.0
>>> score = 0
>>> assert 'Initialize a game board for tracking ships or shots.' in initialize_board.__doc__
>>> if 'Initialize a game board for tracking ships or shots.' in initialize_board.__doc__:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: doc string is incorrect (copy/paste the whole docstring provided).
success_message: doc string is correct.


Tests for task2-place_ship (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(6.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(22.0)
>>> log_variable('total-points', f'Reading-Week-X, Lab_4_Battleship', 22.0)
>>> question_id = 'task2-place_ship-1'
>>> max_score = 2.0
>>> score = 0
>>> board1 = initialize_board(5)
>>> place_ship(board1, [(2, 3), (3, 2)])
>>> assert board1[1][2] == '🚢', 'row, col are not correctly defined'
>>> assert board1[2][1] == '🚢', 'row, col are not correctly defined'
>>> if board1[1][2] == '🚢' and board1[2][1] == '🚢':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: row, col are not correctly defined
success_message: row, col are correctly defined.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task2-place_ship-2'
>>> max_score = 2.0
>>> score = 0
>>> board1 = initialize_board(5)
>>> place_ship(board1, [(1, 1), (2, 2)])
>>> assert board1[0][0] == '🚢', 'row space is not 1-based'
>>> assert board1[1][1] == '🚢', 'row space is not 1-based'
>>> if board1[0][0] == '🚢' and board1[1][1] == '🚢':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: row space is not 1-based
success_message: row space is 1-based
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task2-place_ship-3'
>>> max_score = 2.0
>>> score = 0
>>> board1 = initialize_board(5)
>>> place_ship(board1, [(3, 3), (4, 4)])
>>> assert board1[2][2] == '🚢', 'row space is not 1-based'
>>> assert board1[3][3] == '🚢', 'row space is not 1-based'
>>> if board1[2][2] == '🚢' and board1[3][3] == '🚢':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: col space is not 1-based
success_message: col space is 1-based


Tests for task3-display_board (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> board3 = ['💥', '🎯', '🚢', '🌊']
>>> from io import StringIO
>>> import sys
>>> max_question_points = str(10.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(22.0)
>>> log_variable('total-points', f'Reading-Week-X, Lab_4_Battleship', 22.0)
>>> question_id = 'task3-display_board-1'
>>> max_score = 2.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> test3 = display_board(board3)
>>> sys.stdout = sys.__stdout__
>>> assert type(test3) == list
>>> if type(test3) == list:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: color_dat is not a list
success_message: color_dat is a list
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> board3 = (['💥', '🎯', '🚢', '🌊'], ['💥', '🎯', '🚢', '🌊'])
>>> from io import StringIO
>>> import sys
>>> question_id = 'task3-display_board-2'
>>> max_score = 2.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> test3 = display_board(board3)
>>> sys.stdout = sys.__stdout__
>>> assert test3[0] == '💥'
>>> if test3[0] == '💥':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: blue 🌊 are not formatted
success_message: blue 🌊 are formatted
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> board3 = (['💥', '🎯', '🚢', '🌊'], ['💥', '🎯', '🚢', '🌊'])
>>> from io import StringIO
>>> import sys
>>> question_id = 'task3-display_board-3'
>>> max_score = 2.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> test3 = display_board(board3)
>>> sys.stdout = sys.__stdout__
>>> assert test3[1] == '🎯'
>>> if test3[1] == '🎯':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: blue 🌊 are not formatted with 🌊
success_message: blue 🌊 are formatted with 🌊
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> board3 = (['💥', '🎯', '🚢', '🌊'], ['💥', '🎯', '🚢', '🌊'])
>>> from io import StringIO
>>> import sys
>>> question_id = 'task3-display_board-4'
>>> max_score = 2.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> test3 = display_board(board3)
>>> sys.stdout = sys.__stdout__
>>> assert test3[2] == '🚢'
>>> if test3[2] == '🚢':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: blue 🌊 are not formatted with 🌊
success_message: blue 🌊 are formatted with 🌊
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> board3 = (['💥', '🎯', '🚢', '🌊'], ['💥', '🎯', '🚢', '🌊'])
>>> from io import StringIO
>>> import sys
>>> question_id = 'task3-display_board-5'
>>> max_score = 2.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> test3 = display_board(board3)
>>> sys.stdout = sys.__stdout__
>>> assert test3[3] == '🌊'
>>> if test3[3] == '🌊':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: blue 🌊 are not formatted with 🌊
success_message: blue 🌊 are formatted with 🌊


Tests for task4-fire_shot (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(4.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(22.0)
>>> log_variable('total-points', f'Reading-Week-X, Lab_4_Battleship', 22.0)
>>> question_id = 'task4-fire_shot-1'
>>> max_score = 1.0
>>> score = 0
>>> board4 = (['💥', '🎯', '🚢', '🌊'], ['💥', '🎯', '🚢', '🌊'])
>>> test4 = fire_shot(board4, (1, 2))
>>> assert test4 == 'Already Fired'
>>> if test4 == 'Already Fired':
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Hitting X does not yield 'Already Fired'.
success_message: Hitting X yields 'Already Fired'.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task4-fire_shot-2'
>>> max_score = 1.0
>>> score = 0
>>> board4 = (['💥', '🎯', '🚢', '🌊'], ['💥', '🎯', '🚢', '🌊'])
>>> test4 = fire_shot(board4, (1, 2))
>>> assert test4 == 'Already Fired'
>>> if test4 == 'Already Fired':
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Hitting X does not yield 'Already Fired'.
success_message: Hitting X yields 'Already Fired'.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task4-fire_shot-3'
>>> max_score = 1.0
>>> score = 0
>>> board4 = (['💥', '🎯', '🚢', '🌊'], ['💥', '🎯', '🚢', '🌊'])
>>> test4 = fire_shot(board4, (1, 3))
>>> assert test4 == 'Hit'
>>> if test4 == 'Hit':
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Hitting X does not yield 'Hit'.
success_message: Hitting X yields 'Hit'.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task4-fire_shot-4'
>>> max_score = 1.0
>>> score = 0
>>> board4 = (['💥', '🎯', '🚢', '🌊'], ['💥', '🎯', '🚢', '🌊'])
>>> test4 = fire_shot(board4, (1, 4))
>>> assert test4 == 'Miss'
>>> if test4 == 'Miss':
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Hitting X does not yield 'Miss'.
success_message: Hitting X yields 'Miss'.
