```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("Lab_4_Battleship_q__", "week_4", "lab", assignment_points = 22.0, assignment_tag = 'week4-lab')

# Initialize Otter
import otter
grader = otter.Notebook("Lab_4_Battleship_q__.ipynb")
```

| 🎮 | LAB 4 : BATTLESHIP | 🎮 |
|---|----------|---|





Complete the following 4 functions and corresponding function calls which facilitate playing battleship in the Jupyter console. 

Note: bullet points represent tasks you must do to complete the lab.

| 💥 | 1: Set up the game board |  💥 |
|---|----------|---|

- Add the following doc string to `initialize_board()`. It goes between the pair of `"""`  at the top of the function (replacing `...`). 
```python
    """
    Initialize a game board for tracking ships or shots.

    :param size: The size of the game board (size x size).
    :return: A grid filled with '~' to represent water.
    """
```

None of the next functions are called in the cells in which they are defined by default. You may add calls to run the functions inside their cells if you like, or use the ones in **5: Play the game** below. 

`initialize_board()` is given to you in its entirety (minus the doc string above). The operations it performs use a for loop to execute a program designed to generate data for a game board. 

The for loop in `initialize_board()` uses an iterator `_` generated by `range(size)`, an iterable that produces size values. The `_` acts as a placeholder. It will automatically be replaced with `~` as the following two lines are executed. 

Each time the loop runs, a row of `~` symbols is created and appended to the board object. The function returns the complete `size` × `size` grid, called `board`. 

```python
def initialize_board(size):
    """
    ...
    """
    board = []
    for _ in range(size):
        row = ["~"] * size
        board.append(row)
    return board
```

```python
grader.check("task1-initialize_board")
```

| 💥💥 | 2: Place ships on the board |  💥💥 |
|---|----------|---|

```python
    """
    Place ships on the board.

    :param board: The game board.
    :param coordinates: List of (row, col) tuples for ship positions (1-based indexing).
    """
```
Now that we have a game board filled with `~` to represent water, we want to replace some of them with `"S"` to represent ships placed by a player. 

- First, place `coordinates`, the second input to the function `place_ship()`, in the function definition below, which is initially populated by `...`  

Next, you will have to name variables which will be individually assigned values from the list of tupples `coordinates` in the first line of the for-loop. These variables will represent the row and column indices of (part of) the ship's coordinates on the board. 

- Define two variables to be updated simultaneously by replacing the Ellipisis in `for ... in coordinates` with `row,col`.

Since this game requires human input to play, we probably don't want to index in the matrix like Python normally would — starting from 0. 

Let's format the `row` and `col` variables to count such that (1,1) is the uupper leftmost element — instead of (0,0). 

- Convert to 1-based indexing in row space by using a compound assignment operator to subtract 1 from `row` and reassign that value to `row`, replacing the `...` after `row`
- Do the same for `col` on the next line


Now, when `board[row][col] = "S"` runs on the following line, it will look into the matrix `board` at the index indicated by `[row][col]` and populate that space with `"S"` to represent (part of) a ship.

```python
def place_ship(board, coordinates):
    """
    ...
    """
    ...
        ...
        ...
        board[row][col] = "S"
```

```python
grader.check("task2-place_ship")
```

| 💥💥💥 | 3: Display the board |  💥💥💥 |
|---|----------|---|

```python
    """
    Display the board with colors.

    """
```
Now that we have functions to quickly generate data about the board and ships, we want to see the whole thing put together in the console. We also know that we will eventually want to incorporate two other types of data on the board: visual indicators of hits and misses. So, we will have 4 total things to display: water, ships, hits, and missing. 

To make it easier to visually parse, `Board`, the object we use to pass data between functions, will be displayed using a unicode character (`■` in different colors for ships, hits, and misses). However, inside each of the functions `board` will be defined using the single string characters, `"~"`, `"S"`, `"X"`, and `"M"`. This is becuase python cannot perform operations using unicode characters as objects. However, it can display unicode characters — even with formatting options like color — to the console, which is how we will display the board as the last step once operations related to ship placement or targeting are performed. 

`display_board()` uses two for loops, 4 logical tests, and a print() call to display a modified version of `board` directly in the console. It generates a blank list which will be used to store data representing the modified version that actually gets printed to the console.

The outer loop iterates through each `row` of `board`, making a new blank list, which we will call `color_dat`. It has exactly  enough entries to store data describing the desired form and color of each individual `element` in that `row`. The inner loop performs a logical test on each `element` in the `row` of `board`, and appends a string of formatting characters to `color_dat` depending on what was found at that `element`. This process continues running until every entry in the matrix has been checked and `color_dat` is a matrix of strings containing exactly the same number of elements as `board`, but in a form that looks much better when output the console. We don't need to worry about retaining this data to perform operations on, so we just print directly to the console (`display_board` has no return line).

A major benefit of using a new blank list populated with formatting data for the printed output is that `board` still exists as an object retaining all of its original data (along with any updates from `place_ship()` or `fire_shot()`). Since `board` is easy to update and maintain, that let's us continually add new data to it without having to reinitialize every time. 

- Define a blank list saved as the variable `color_dat` by replacing the `...` in the for loop `for row in board:` with `color_dat =` an empty list.

Inside the for loop `for element in row:`, a series of logical tests will be performed against the presence of a string character which we are using to track data about the current status of any individual `element` of the matrix (well, really an element of `row`, but each individual `row` will be presented sequentially which makes a matrix, so...).



- Add the respective single string characters to the execution statements of the four logical tests inside `for element in row:`.
    - If you don't know which string character to place with each logical test, consult the other functions in this lab.
- Next, paste the following escape characters in the respective `row.append("...")` method calls nested in the logical statements `if element == ...` and `elif element == ...` nested in `for element in row:` nested in `for row in board:` which make up `display_board()` below.
```python
                color_dat.append("\033[96m~\033[0m")  # Blue for water
                color_dat.append("\033[30m■\033[0m")  # Black for ships
                color_dat.append("\033[31m■\033[0m")  # Red for hits
                color_dat.append("\033[37m■\033[0m")  # Grey/White for misses

```


```python
def display_board(board):
    """
    ...
    """
    for row in board:
        ...
        for element in row:
            if element == "~":
                ...
            elif element == "S":
                ...
            elif element == "X":
                ...
            elif element == "M":
                ...
        print(" ".join(color_dat))
        return color_dat
```

```python
grader.check("task3-display_board")
```

| 💥💥💥💥 | 4: Fire away |  💥💥💥💥 |
|---|----------|---|

```python
    """
    Fire a shot at the board and return the result.

    :param board: The game board.
    :param target: A tuple (row, col) where the shot is fired (1-based indexing).
    :return: The result of the shot: "Hit", "Miss", or "Already Fired".
    """

```

We want to manipulate `board` using indices based on the input tuple `target`. Let's break `target` into individual integers representing the row and column indices. 
- Add a line to define two variables, `row` and `col` which are to store values contained in `target`, replacing the first `...` after the doc string
- We want to allow for human input data, so we will need to convert `row` and `col` to a 1-based indexing system. Replace the `...` on the next two lines
    
Next, we need to determine what is currently occupying the index corresponding to our targeting information and what will happen to that index as a result of a shot being fired there. 

If the target is an `X`, that means we've already shot at that space, so we want to return a message stating `"Already Fired"`. 

If the target contains an `S` that means it contains part of a ship, so we want to update the space with an `X` and return `"Hit"`. 

If the target contains `~` that means it contains water, so we want to update the space with an `M` and return `"Miss"`. 

Since we built the board, we know those are the only possible options, so everything else should be an invalid input, returning `"Invalid Target"`.

- Fill in the 4 logical expressions below to complete `fire_shot()`, replacing 9 sets of `...`

```python
def fire_shot(board, target):
    """
    ...
    """
    ...

    if board[row][col] == "X":
        ...
    elif board[row][col] == "S":
        ...
    elif board[row][col] == "~":
        ...
    else:
        return "Invalid Target"   
```

```python
grader.check("task4-fire_shot")
```

## | 💥💥💥💥💥 | 5: Play the game |  💥💥💥💥💥 |
|---|----------|---|

The variables `player_board` and `tracking_board` are populated with data using the initialize board function to produce 10 x 10 game boards filled with water.

To play battleship, you should place 5 ships on your board. These ships are: carrier (occupies 5 spaces), battleship (occupies 4), submarine (occupies 3), destroyer (occupies 3 spaces) and patrol boat (occupies 2). The destroyer and patrol boat have already been placed for you with `place_ship`. 

Keep your screen hidden from your opponent so they don't know where your ships are!

- Edit the placement of the destroyer and patrol boat so your opponent doesn't know where they are
     - Or try to be sneaky and keep them in their default position as the one place your opponent won't check?
- Add 3 new calls to `place_ship` to place the reamining 3 ships
- Run the cell to save your ship placement and view the board

You can run the cell as many times as you like before the game begins to help make sure you are placing ships where you think you're placing them, but **don't re-initialize the board once you've started to play or you'll lose all progress**.

Ships should be continuous and not overlap. They can run on diagonals. 



```python
# Initialize boards
player_board = initialize_board(10)
tracking_board = initialize_board(10)

# Place remaining 3 ships on the board
place_ship(player_board, [(2, 2), (2, 3), (2, 4)]) #destroyer
place_ship(player_board, [(7, 7), (8, 8)]) #patrol boat

# Display board
display_board(player_board)
```

Once both you and your oponent have boards set up you can use the following cell to target spaces on their board (and vice versa). You can also use the tracking function below to record your shots on your opponent's board (unless you can keep track of all that in your head?!) 

- Player 1 says: "5,5"
- Player 2 runs: result = `fire_shot(player_board, (5, 5))` and `display_board(player_board)`
- Player 2 says: "Miss!"
- Optional: Player 1 runs: `record_opponent_feedback(tracking_board)` and provides input to the box, then runs `display_board(tracking_board)`

```python
# Fire shots 
result = fire_shot(player_board, (2, 3))

# Display the updated board
display_board(player_board)
```

| 💥💥💥💥💥💥 | Track your shots on your opponent's board |  💥💥💥💥💥💥 |
|---|----------|---|

You want to make sure you're not aiming at the same place twice.

Run the following cell to edit the tracking board.

```python
def record_opponent_feedback(tracking_board):
    """
    Prompt the user for a target and result, then update the tracking board.

    :param tracking_board: The board tracking your shots on the opponent's board.
    """
    try:
        # Get input for target coordinates as a tuple
        target_input = input("row,col: ")
        target = tuple(map(int, target_input.split(",")))

        # Get input for result ("Hit" or "Miss")
        result = input("hit or miss?: ").capitalize()
        
        # Update the tracking board
        row, col = target
        row -= 1  # Convert to 0-based index
        col -= 1  # Convert to 0-based index

        if result == "Hit":
            tracking_board[row][col] = "X"  # Mark hit
        elif result == "Miss":
            tracking_board[row][col] = "M"  # Mark miss
        else:
            print("Invalid result. Please enter 'hit' or 'miss'.")
    except ValueError:
        print("Invalid input. Please ensure coordinates are in the format row,col.")

tracking_board = initialize_board(size=10)
```

```python
#Running this following cell will prompt for input. 

#First, tell your opponent your shot and input coordinates as a comma separated pair of integers and hit enter. 

#Then, input "hit" or "miss" (no quotes) depending on what your opponent says.

record_opponent_feedback(tracking_board)
display_board(tracking_board)
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week4-lab", "Lab_4_Battleship_q__")
```

Tests for task1-initialize_board (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(2.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(22.0)
>>> log_variable('total-points', f'Reading-Week-X, Lab_4_Battleship_q__', 22.0)
>>> question_id = 'task1-initialize_board-1'
>>> max_score = 2.0
>>> score = 0
>>> assert 'Initialize a game board for tracking ships or shots.' in initialize_board.__doc__
>>> if 'Initialize a game board for tracking ships or shots.' in initialize_board.__doc__:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship_q__', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: doc string is incorrect (copy/paste the whole docstring provided).
success_message: doc string is correct.


Tests for task2-place_ship (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(6.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(22.0)
>>> log_variable('total-points', f'Reading-Week-X, Lab_4_Battleship_q__', 22.0)
>>> question_id = 'task2-place_ship-1'
>>> max_score = 2.0
>>> score = 0
>>> board1 = initialize_board(5)
>>> place_ship(board1, [(2, 3), (3, 2)])
>>> assert board1[1][2] == 'S', 'row, col are not correctly defined'
>>> assert board1[2][1] == 'S', 'row, col are not correctly defined'
>>> if board1[1][2] == 'S' and board1[2][1] == 'S':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship_q__', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: row, col are not correctly defined
success_message: row, col are correctly defined.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task2-place_ship-2'
>>> max_score = 2.0
>>> score = 0
>>> board1 = initialize_board(5)
>>> place_ship(board1, [(1, 1), (2, 2)])
>>> assert board1[0][0] == 'S', 'row space is not 1-based'
>>> assert board1[1][1] == 'S', 'row space is not 1-based'
>>> if board1[0][0] == 'S' and board1[1][1] == 'S':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship_q__', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: row space is not 1-based
success_message: row space is 1-based
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task2-place_ship-3'
>>> max_score = 2.0
>>> score = 0
>>> board1 = initialize_board(5)
>>> place_ship(board1, [(3, 3), (4, 4)])
>>> assert board1[2][2] == 'S', 'row space is not 1-based'
>>> assert board1[3][3] == 'S', 'row space is not 1-based'
>>> if board1[2][2] == 'S' and board1[3][3] == 'S':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship_q__', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: col space is not 1-based
success_message: col space is 1-based


Tests for task3-display_board (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> board3 = ['M', 'X', 'S', '~']
>>> from io import StringIO
>>> import sys
>>> max_question_points = str(10.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(22.0)
>>> log_variable('total-points', f'Reading-Week-X, Lab_4_Battleship_q__', 22.0)
>>> question_id = 'task3-display_board-1'
>>> max_score = 2.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> test3 = display_board(board3)
>>> sys.stdout = sys.__stdout__
>>> assert type(test3) == list
>>> if type(test3) == list:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship_q__', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: color_dat is not a list
success_message: color_dat is a list
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> board3 = (['M', 'X', 'S', '~'], ['M', 'X', 'S', '~'])
>>> from io import StringIO
>>> import sys
>>> question_id = 'task3-display_board-2'
>>> max_score = 2.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> test3 = display_board(board3)
>>> sys.stdout = sys.__stdout__
>>> assert test3[0] == '\x1b[37m■\x1b[0m'
>>> if test3[0] == '\x1b[37m■\x1b[0m':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship_q__', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: blue ~ are not formatted with \033[94m~\033[0m
success_message: blue ~ are formatted with \033[94m~\033[0m
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> board3 = (['M', 'X', 'S', '~'], ['M', 'X', 'S', '~'])
>>> from io import StringIO
>>> import sys
>>> question_id = 'task3-display_board-3'
>>> max_score = 2.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> test3 = display_board(board3)
>>> sys.stdout = sys.__stdout__
>>> assert test3[1] == '\x1b[31m■\x1b[0m'
>>> if test3[1] == '\x1b[31m■\x1b[0m':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship_q__', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: blue ~ are not formatted with \033[94m~\033[0m
success_message: blue ~ are formatted with \033[94m~\033[0m
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> board3 = (['M', 'X', 'S', '~'], ['M', 'X', 'S', '~'])
>>> from io import StringIO
>>> import sys
>>> question_id = 'task3-display_board-4'
>>> max_score = 2.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> test3 = display_board(board3)
>>> sys.stdout = sys.__stdout__
>>> assert test3[2] == '\x1b[30m■\x1b[0m'
>>> if test3[2] == '\x1b[30m■\x1b[0m':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship_q__', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: blue ~ are not formatted with \033[94m~\033[0m
success_message: blue ~ are formatted with \033[94m~\033[0m
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> board3 = (['M', 'X', 'S', '~'], ['M', 'X', 'S', '~'])
>>> from io import StringIO
>>> import sys
>>> question_id = 'task3-display_board-5'
>>> max_score = 2.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> test3 = display_board(board3)
>>> sys.stdout = sys.__stdout__
>>> assert test3[3] == '\x1b[94m~\x1b[0m'
>>> if test3[3] == '\x1b[94m~\x1b[0m':
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship_q__', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: blue ~ are not formatted with \033[94m~\033[0m
success_message: blue ~ are formatted with \033[94m~\033[0m


Tests for task4-fire_shot (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(4.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(22.0)
>>> log_variable('total-points', f'Reading-Week-X, Lab_4_Battleship_q__', 22.0)
>>> question_id = 'task4-fire_shot-1'
>>> max_score = 1.0
>>> score = 0
>>> board4 = (['M', 'X', 'S', '~'], ['M', 'X', 'S', '~'])
>>> test4 = fire_shot(board4, (1, 2))
>>> assert test4 == 'Already Fired'
>>> if test4 == 'Already Fired':
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship_q__', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Hitting X does not yield 'Already Fired'.
success_message: Hitting X yields 'Already Fired'.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task4-fire_shot-2'
>>> max_score = 1.0
>>> score = 0
>>> board4 = (['M', 'X', 'S', '~'], ['M', 'X', 'S', '~'])
>>> test4 = fire_shot(board4, (1, 2))
>>> assert test4 == 'Already Fired'
>>> if test4 == 'Already Fired':
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship_q__', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Hitting X does not yield 'Already Fired'.
success_message: Hitting X yields 'Already Fired'.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task4-fire_shot-3'
>>> max_score = 1.0
>>> score = 0
>>> board4 = (['M', 'X', 'S', '~'], ['M', 'X', 'S', '~'])
>>> test4 = fire_shot(board4, (1, 3))
>>> assert test4 == 'Hit'
>>> if test4 == 'Hit':
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship_q__', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Hitting X does not yield 'Hit'.
success_message: Hitting X yields 'Hit'.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'task4-fire_shot-4'
>>> max_score = 1.0
>>> score = 0
>>> board4 = (['M', 'X', 'S', '~'], ['M', 'X', 'S', '~'])
>>> test4 = fire_shot(board4, (1, 4))
>>> assert test4 == 'Miss'
>>> if test4 == 'Miss':
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('Lab_4_Battleship_q__', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Hitting X does not yield 'Miss'.
success_message: Hitting X yields 'Miss'.
