```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_practicequiz_q", "week_3", "practicequiz", assignment_points = 17.0, assignment_tag = 'week3-practicequiz')

# Initialize Otter
import otter
grader = otter.Notebook("1_practicequiz_q.ipynb")
```

# â“ ðŸŽï¸ NumPy Drift: Speeding Through Arrays at Lightning Speed ðŸ

![](./assets/figures/1972-ford-mustang-mach-i.jpg)

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_practicequiz_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_practicequiz_q import Question2
Question2().show()

```

 ## NumPy: Array Creation in the World of Sport Cars
 ### Instructions:
 You are designing the internal array processing system for a sport car's telemetry system! Create a NumPy array of integers from 0 to 120, representing speed data in km/h, using NumPy's built-in `np.arange()` method. Write a function `generate_speed_array()` that:
 1. imports numpy using the common alias `np`.
 2. Uses `np.arange()` to generate the array, starting from 0 and ending at 120 (inclusive). Save it to the variable `array`.

```python
...
def generate_speed_array():
    # Create an array from 0 to 120
    ...
    return array
# Example usage
result = generate_speed_array()
print("Generated speed array:", result)
```

```python
grader.check("sport-car-dynamics")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week3-practicequiz", "1_practicequiz_q")
```

Tests for sport-car-dynamics (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(10.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(10.0)
>>> log_variable('total-points', f'Reading-Week-X, 1_practicequiz_q', 10.0)
>>> question_id = 'sport-car-dynamics-1'
>>> max_score = 2.0
>>> score = 0
>>> condition = 'np' in globals()
>>> assert condition, 'Failed: `np` is not defined globally.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicequiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: numpy was not imported as np.
success_message: Success: numpy was imported as np.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import numpy as np
>>> question_id = 'sport-car-dynamics-2'
>>> max_score = 5.0
>>> score = 0
>>> expected_array = np.arange(0, 121)
>>> output_array = generate_speed_array()
>>> is_array_correct = np.array_equal(output_array, expected_array)
>>> output_type = isinstance(output_array, np.ndarray)
>>> assert is_array_correct, 'Test failed: The generated array is incorrect.'
>>> assert output_type, 'Test failed: Output should be a NumPy array.'
>>> if is_array_correct and output_type:
...     score = 5.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicequiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(is_array_correct))
>>> responses = update_responses(question_id, str(output_type))
```
failure_message: Failed: The function did not generate the correct array.
success_message: Success: The function did generate the correct array.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import inspect
>>> import re
>>> question_id = 'sport-car-dynamics-3'
>>> max_score = 3.0
>>> score = 0
>>> code = inspect.getsource(generate_speed_array)
>>> arange_used = bool(re.search('np\\.arange\\(', code))
>>> assert arange_used, 'Test failed: np.arange() must be used in the function.'
>>> if arange_used:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicequiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function must generate an array using np.arange().
success_message: Success: The function generates a correct NumPy array using np.arange().
