```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("7_numpy_q", "week_3", "lecture", assignment_points = 5.0, assignment_tag = 'week3-lecture')

# Initialize Otter
import otter
grader = otter.Notebook("7_numpy_q.ipynb")
```

# â“ Sorting M&Ms with NumPy

![](./assets/figures/m-m-fun-questions.avif)



## Instructions:
Imagine you work at an M&M sorting facility. You receive a batch of M&Ms in random order, and your task is to sort them by size using NumPy's built-in sorting functionality.

Write a function `sort_mm_sizes()` that:
    1. Imports NumPy as `np`.
    2. Takes a NumPy array of integers representing the sizes of M&Ms.
    3. Sorts the array in ascending order using NumPy's `np.sort()` method.
    4. Returns the sorted array.

```python
...


def sort_mm_sizes(mm_sizes):

    # Sort the array using NumPy's built-in method
    ...
    return sorted_sizes


# Example usage
mm_sizes = np.array([3, 1, 4, 1, 5, 9])
sorted_sizes = sort_mm_sizes(mm_sizes)
print("Sorted M&M sizes:", sorted_sizes)
```

```python
grader.check("numpy-mm-sorting")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week3-lecture", "7_numpy_q")
```

Tests for numpy-mm-sorting (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(5.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(5.0)
>>> log_variable('total-points', f'Reading-Week-X, 7_numpy_q', 5.0)
>>> question_id = 'numpy-mm-sorting-1'
>>> max_score = 1.0
>>> score = 0
>>> condition = 'np' in globals()
>>> assert condition, "You must import numpy using the alias 'np'."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('7_numpy_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(condition))
```
failure_message: Failed: You must import numpy using the alias 'np'.
success_message: Success: 'numpy' is imported as 'np'!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import numpy as np
>>> question_id = 'numpy-mm-sorting-2'
>>> max_score = 2.0
>>> score = 0
>>> input_array = np.array([3, 1, 4, 1, 5, 9])
>>> expected_output = np.array([1, 1, 3, 4, 5, 9])
>>> actual_output = sort_mm_sizes(input_array)
>>> assert np.array_equal(actual_output, expected_output), f'Expected sorted array {expected_output}, but got {actual_output}.'
>>> if np.array_equal(actual_output, expected_output):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('7_numpy_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(input_array))
>>> responses = update_responses(question_id, str(expected_output))
>>> responses = update_responses(question_id, str(actual_output))
```
failure_message: Failed: sort_mm_sizes does not correctly sort a typical M&M sizes array.
success_message: Success: sort_mm_sizes correctly sorts a typical M&M sizes array!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import numpy as np
>>> question_id = 'numpy-mm-sorting-3'
>>> max_score = 2.0
>>> score = 0
>>> input_array = np.array([2, 2, 1, 3, 3, 1])
>>> expected_output = np.array([1, 1, 2, 2, 3, 3])
>>> actual_output = sort_mm_sizes(input_array)
>>> assert np.array_equal(actual_output, expected_output), f'Expected sorted array {expected_output}, but got {actual_output}.'
>>> if np.array_equal(actual_output, expected_output):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('7_numpy_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(input_array))
>>> responses = update_responses(question_id, str(expected_output))
>>> responses = update_responses(question_id, str(actual_output))
```
failure_message: Failed: sort_mm_sizes does not correctly handle arrays with duplicates or already sorted elements.
success_message: Success: sort_mm_sizes correctly handles arrays with duplicate and already sorted elements!
