```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("9_dodgeball_sympy_q", "week_3", "lecture", assignment_points = 5.0, assignment_tag = 'week3-lecture')

# Initialize Otter
import otter
grader = otter.Notebook("9_dodgeball_sympy_q.ipynb")
```

# Dodgeball: Simple Symbolic Projectile Motion

## Instructions:
In the dodgeball arena, a player throws the ball with an initial velocity. Your task is to:

1. Use SymPy to solve the height equation $h(t) = v_0t - \frac{1}{2}gt^2$ for the time when the ball hits the ground.
2. Assume $v_0 = 15$ m/s and $g = 9.8$ m/s².

Your solution must include:
A function `solve_projectile_time()` that (we have provided this function for you):

1. Define the symbols `t`, `v0`, and `g`.
2. Define the height equation `h` symbolically.
3. Solve the height equation for time when height is zero. To do this, use the `solve()` function, which takes two arguments: the equation to solve and the variable to solve for, in this case `t`.
4. Build the dictionaries with the values to substitute. We have provided the code for you to filter positive times - it would not make sense to have a negative time.
5. Compute the time by substituting the values of `v0` and `g` into the equation. We did this using the `subs()` function, which takes a dictionary of symbols and values to substitute. We didn't want you to do this step because it requires a loop which you have not learned yet. 


```python
from sympy import symbols, solve, N  # Import N for numeric evaluation

def solve_projectile_time():
    # 1. Define symbols
    ...
    
    # 2. Height equation
    ...
    
    # 3. Solve for time when h = 0
    ...
    
    # 4. Define the values dictionaries
    ...

    # Substitute numerical values for v0 and g to filter positive times
    substituted_solutions = [sol.subs(values_dictionary) for sol in time_solutions]

    # Filter positive numerical solutions
    positive_time = [sol for sol in substituted_solutions if N(sol) > 0][0]

    # Return the positive time as a numeric value
    return N(positive_time)


# Example usage
result = solve_projectile_time()
print(f"Time when the ball hits the ground: {result:.2f} seconds")
```

```python
grader.check("Dodgeball-w-sympy")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week3-lecture", "9_dodgeball_sympy_q")
```

Tests for Dodgeball-w-sympy (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import inspect
>>> import re
>>> from sympy import symbols, solve, N
>>> max_question_points = str(5.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(5.0)
>>> log_variable('total-points', f'Reading-Week-X, 9_dodgeball_sympy_q', 5.0)
>>> question_id = 'Dodgeball-w-sympy-1'
>>> max_score = 1.0
>>> score = 0
>>> def get_solve_projectile_time_code(func):
...     """
...     Returns the source code of solve_projectile_time() as a string.
...     """
...     return inspect.getsource(func)
>>> code = get_solve_projectile_time_code(solve_projectile_time)
>>> symbols_pattern = 'v0\\s*,\\s*g\\s*,\\s*t\\s*=\\s*symbols\\s*\\(\\s*[\\"\']v0\\s+g\\s+t[\\"\']\\s*\\)'
>>> symbols_pattern_found = bool(re.search(symbols_pattern, code))
>>> assert symbols_pattern_found, "You must define v0, g, t = symbols('v0 g t')."
>>> if symbols_pattern_found:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('9_dodgeball_sympy_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(symbols_pattern_found))
```
failure_message: Failed: You must define v0, g, t = symbols('v0 g t').
success_message: Success: v0, g, t are defined as symbols in the function.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import inspect
>>> import re
>>> from sympy import symbols, solve, N
>>> question_id = 'Dodgeball-w-sympy-2'
>>> max_score = 1.0
>>> score = 0
>>> def get_solve_projectile_time_code(func):
...     """
...     Returns the source code of solve_projectile_time() as a string.
...     """
...     return inspect.getsource(func)
>>> code = get_solve_projectile_time_code(solve_projectile_time)
>>> height_pattern = 'h\\s*=\\s*v0\\s*\\*\\s*t\\s*-\\s*\\(\\s*1\\s*\\/\\s*2\\s*\\)\\s*\\*\\s*g\\s*\\*\\s*t\\*\\*2'
>>> height_pattern_found = bool(re.search(height_pattern, code))
>>> assert height_pattern_found, 'You must define h = v0*t - (1/2)*g*t**2 exactly.'
>>> if height_pattern_found:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('9_dodgeball_sympy_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: You must define h = v0*t - (1/2)*g*t**2 in the code.
success_message: Success: The height equation is defined as h = v0*t - (1/2)*g*t^2.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import inspect
>>> import re
>>> from sympy import symbols, solve, N
>>> question_id = 'Dodgeball-w-sympy-3'
>>> max_score = 1.0
>>> score = 0
>>> def get_solve_projectile_time_code(func):
...     """
...     Returns the source code of solve_projectile_time() as a string.
...     """
...     return inspect.getsource(func)
>>> code = get_solve_projectile_time_code(solve_projectile_time)
>>> solve_pattern = 'solve\\s*\\(\\s*h\\s*,\\s*t\\s*\\)'
>>> solve_pattern_found = bool(re.search(solve_pattern, code))
>>> assert solve_pattern_found, 'You must call solve(h, t) to find the times when h=0.'
>>> if solve_pattern_found:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('9_dodgeball_sympy_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Must call solve(h, t).
success_message: Success: The code uses solve(h, t) from Sympy to find time solutions.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Dodgeball-w-sympy-4'
>>> max_score = 2.0
>>> score = 0
>>> expected_time = 2 * 15 / 9.8
>>> lower_bound = expected_time * 0.98
>>> upper_bound = expected_time * 1.02
>>> positive_time = solve_projectile_time()
>>> condition = lower_bound <= positive_time <= upper_bound
>>> assert condition, f'Time was {positive_time:.3f}, but expected ~{expected_time:.3f} within ±2%.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('9_dodgeball_sympy_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The time is incorrect or negative for v0=15, g=9.8.
success_message: Success: The function returns a physically correct positive time for v0=15, g=9.8.
