```python
from pykubegrader.tokens.validate_token import validate_token
validate_token('type the key provided by your instructor here')

# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("10_rube_goldberg", "week_5", "lecture", assignment_points = 21.0, assignment_tag = 'week5-lecture')

# Initialize Otter
import otter
grader = otter.Notebook("10_rube_goldberg.ipynb")
```

# 💻 Activity Rube Goldberg Machine: Modular Functions with *args and **kwargs

## Question 1 (Points: 21.0): Rube Goldberg Machine

![Rube Goldberg Machine](./assets/figures/200.gif)

### Instructions:
You are designing a **Rube Goldberg machine simulator** where multiple functions interact in a sequence to achieve a simple goal in an overly complex way.

Your task is to define multiple modular functions that take flexible inputs using `*args` and `**kwargs`.

### Requirements:
1. **Define a Function `start_machine(*args, **kwargs)`**  
   - Calls other functions based on `*args`.
   - Passes metadata (like `speed=`, `material=`) using `**kwargs`.
   
2. **Define Helper Functions for Components:**  
   - Each function should accept flexible arguments and pass data to the next function.
   - Example:  
     ```python
     def lever(action, *args, **kwargs):
         return f"Lever {action} → {next_step(*args, **kwargs)}"
     ```

3. **Ensure Function Interaction:**  
   - Functions should call one another, creating a step-by-step simulation.
   - The final function should return `"Goal Achieved!"`.
   
4. **Example Execution:**
```python
start_machine("lever", "pulley", "dominoes", speed="fast", material="wood")
```
**Expected Output:**
```
Lever triggered → Pulley activated → Dominoes toppled → Goal Achieved!
```

### Helpful Information:
- Use `*args` to dynamically chain components.
- Use `**kwargs` to pass machine properties like speed or material.
- The machine should always end with `"Goal Achieved!"`.

```python
# Function Definitions
# A Rube Goldberg machine is a contraption that is deliberately over-engineered to perform a simple task.


# Define a function lever that takes an argument and returns a string that says "Lever triggered → [next_step]".
...


# Define a function pulley that takes an action as an argument and returns a string that says "Pulley [action] → [next_step]".
...


# Define a function dominoes that takes an action as an argument and returns a string that says "Dominoes [action] → [next_step]".
...


# Define a function next_step that takes a variable number of arguments and returns a string that says "Goal Achieved!" if no arguments are passed.
# If arguments are passed, the function should determine the next step in the machine based on the first argument and call the corresponding function.
# The function should return "Unknown step!" if the first argument does not match any of the defined functions.
def next_step(*args, **kwargs):
    """Determines the next step in the machine."""

    # if no arguments are passed, return "Goal Achieved!"
    ...

    # From the first argument, determine the next component, save to the `next_component`` variable, and save the remaining arguments to the `remaining_steps` variable
    ...

    # Create a dictionary that maps the next component to the corresponding function
    # Save this dictionary to the `component_map` variable
    ...

    # If the next component is in the component map, call the corresponding function with the remaining steps and return the result
    # If the next component is not in the component map, return "Unknown step!"
    # Recall, you can use the in operator to check if a key is in a dictionary
    ...


# Define a function start_machine that takes a variable number of arguments and keyword arguments and initiates the Rube Goldberg machine.
# The function should call the next_step function with the provided arguments and return the result.
# This is the main function call that initiates the machine.
...

# Example execution
print(start_machine("lever", "pulley", "dominoes", speed="fast", material="wood"))
```

```python
grader.check("rube-goldberg-modularity")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.tokens.validate_token import validate_token
validate_token()


from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week5-lecture", "10_rube_goldberg")
```

Tests for rube-goldberg-modularity (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(21.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(21.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token()
>>> log_variable('total-points', f'Reading-Week-X, 10_rube_goldberg', 21.0)
>>> question_id = 'rube-goldberg-modularity-1'
>>> max_score = 1.0
>>> score = 0
>>> condition = 'lever' in globals() and callable(lever)
>>> assert condition, "The function 'lever' is not defined or not callable."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_rube_goldberg', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function 'lever' is not defined or not callable.
success_message: Success: The function 'lever' is defined and callable!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'rube-goldberg-modularity-2'
>>> max_score = 1.0
>>> score = 0
>>> condition = 'pulley' in globals() and callable(pulley)
>>> assert condition, "The function 'pulley' is not defined or not callable."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_rube_goldberg', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function 'pulley' is not defined or not callable.
success_message: Success: The function 'pulley' is defined and callable!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'rube-goldberg-modularity-3'
>>> max_score = 1.0
>>> score = 0
>>> condition = 'dominoes' in globals() and callable(dominoes)
>>> assert condition, "The function 'dominoes' is not defined or not callable."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_rube_goldberg', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function 'dominoes' is not defined or not callable.
success_message: Success: The function 'dominoes' is defined and callable!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'rube-goldberg-modularity-4'
>>> max_score = 1.0
>>> score = 0
>>> condition = 'next_step' in globals() and callable(next_step)
>>> assert condition, "The function 'next_step' is not defined or not callable."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_rube_goldberg', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function 'next_step' is not defined or not callable.
success_message: Success: The function 'next_step' is defined and callable!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'rube-goldberg-modularity-5'
>>> max_score = 1.0
>>> score = 0
>>> condition = 'start_machine' in globals() and callable(start_machine)
>>> assert condition, "The function 'start_machine' is not defined or not callable."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_rube_goldberg', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function 'start_machine' is not defined or not callable.
success_message: Success: The function 'start_machine' is defined and callable!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'rube-goldberg-modularity-6'
>>> max_score = 2.0
>>> score = 0
>>> expected_output = 'Lever triggered → Goal Achieved!'
>>> actual_output = lever()
>>> assert actual_output == expected_output, f'Expected:\n{expected_output}\nBut got:\n{actual_output}'
>>> if (actual_output == expected_output, f'Expected:\n{expected_output}\nBut got:\n{actual_output}'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_rube_goldberg', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function 'lever' does not return the expected format.
success_message: Success: The function 'lever' correctly returns the expected format!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'rube-goldberg-modularity-7'
>>> max_score = 2.0
>>> score = 0
>>> expected_output = 'Pulley activated → Goal Achieved!'
>>> actual_output = pulley()
>>> assert actual_output == expected_output, f'Expected:\n{expected_output}\nBut got:\n{actual_output}'
>>> if (actual_output == expected_output, f'Expected:\n{expected_output}\nBut got:\n{actual_output}'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_rube_goldberg', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function 'pulley' does not return the expected format.
success_message: Success: The function 'pulley' correctly returns the expected format!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'rube-goldberg-modularity-8'
>>> max_score = 2.0
>>> score = 0
>>> expected_output = 'Dominoes toppled → Goal Achieved!'
>>> actual_output = dominoes()
>>> assert actual_output == expected_output, f'Expected:\n{expected_output}\nBut got:\n{actual_output}'
>>> if (actual_output == expected_output, f'Expected:\n{expected_output}\nBut got:\n{actual_output}'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_rube_goldberg', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function 'dominoes' does not return the expected format.
success_message: Success: The function 'dominoes' correctly returns the expected format!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'rube-goldberg-modularity-9'
>>> max_score = 3.0
>>> score = 0
>>> expected_output = 'Lever triggered → Pulley activated → Dominoes toppled → Goal Achieved!'
>>> actual_output = next_step('lever', 'pulley', 'dominoes')
>>> assert actual_output == expected_output, f'Expected:\n{expected_output}\nBut got:\n{actual_output}'
>>> if (actual_output == expected_output, f'Expected:\n{expected_output}\nBut got:\n{actual_output}'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_rube_goldberg', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function 'next_step' does not resolve steps correctly.
success_message: Success: The function 'next_step' correctly resolves a sequence of steps!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'rube-goldberg-modularity-10'
>>> max_score = 2.0
>>> score = 0
>>> expected_output = 'Unknown step!'
>>> actual_output = next_step('catapult')
>>> assert actual_output == expected_output, f'Expected:\n{expected_output}\nBut got:\n{actual_output}'
>>> if (actual_output == expected_output, f'Expected:\n{expected_output}\nBut got:\n{actual_output}'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_rube_goldberg', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function 'next_step' does not handle unknown steps correctly.
success_message: Success: The function 'next_step' correctly identifies an unknown step!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'rube-goldberg-modularity-11'
>>> max_score = 3.0
>>> score = 0
>>> expected_output = 'Lever triggered → Pulley activated → Dominoes toppled → Goal Achieved!'
>>> actual_output = start_machine('lever', 'pulley', 'dominoes')
>>> assert actual_output == expected_output, f'Expected:\n{expected_output}\nBut got:\n{actual_output}'
>>> if (actual_output == expected_output, f'Expected:\n{expected_output}\nBut got:\n{actual_output}'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_rube_goldberg', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function 'start_machine' does not work correctly.
success_message: Success: The function 'start_machine' correctly initiates the machine!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'rube-goldberg-modularity-12'
>>> max_score = 2.0
>>> score = 0
>>> expected_output = 'Goal Achieved!'
>>> actual_output = start_machine()
>>> assert actual_output == expected_output, f'Expected:\n{expected_output}\nBut got:\n{actual_output}'
>>> if (actual_output == expected_output, f'Expected:\n{expected_output}\nBut got:\n{actual_output}'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('10_rube_goldberg', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function 'start_machine' does not return 'Goal Achieved!' for empty calls.
success_message: Success: The function 'start_machine' correctly handles an empty call!
