```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("4_toybee_tiles", "week_5", "lecture-not-graded", assignment_points = 10.0, assignment_tag = 'week5-lecture-not-graded')

# Initialize Otter
import otter
grader = otter.Notebook("4_toybee_tiles.ipynb")
```

# 💻 Activity: Toynbee Tiles: Error Handling, Function Definitions, and Returns

## Question 1 (Points: 10.0): Unraveling the Mystery of Toynbee Tiles with Python! 🕵️‍♂️🔍

![](./assets/figures/toynbee-tiles-map2.png)

![](./assets/figures/toynbee-2.jpg)

### Instructions:
The **Toynbee Tiles** are mysterious plaques embedded in urban streets, often containing cryptic messages. Your task is to write a function, `analyze_tile()`, that processes a tile's message while handling various errors.
### Requirements:
1. **Function Definition:** Define a function `analyze_tile(message, location="Unknown")`.
2. **Documentation:** Include a docstring with the function's name, parameters, and a brief description.

    ```python
    """
    Processes a Toynbee Tile message while handling errors.

    Args:
        message (str): The message on the tile.
        location (str, optional): The tile's location. Defaults to "Unknown".

    Returns:
        str: A formatted string describing the tile.
    """
    ```

2. **Input Handling:**
- `message` (required): A string containing the tile’s message.
- `location` (optional, default `"Unknown"`): The place where the tile was found.
3. **Error Handling:**
- Raise a `ValueError` if `message` is an empty string.
- Raise a `TypeError` if `message` is not a string.
4. **Return Statement:**
- If valid, return a formatted string:
```
Tile Found: <message>
Location: <location>
```
5. **Location Handling:** If `location` is not a string, return an error message `"Invalid location format."`.
### Example Usage:
```python
analyze_tile("Resurrect dead on planet Jupiter", "Philadelphia")
```
**Expected Output:**
```
Tile Found: Resurrect dead on planet Jupiter
Location: Philadelphia
```
If an error occurs, return a meaningful error message instead of crashing.
### Helpful Information:
- Use `if not message:` to check for an empty string.
- Use `isinstance(variable, str)` to check if a value is a string.
- Use `raise ValueError("message")` or `raise TypeError("message")` for errors.
- Ensure the function gracefully handles incorrect inputs.

![](./assets/figures/toynbee-3.jpeg)

```python
# Function Definition
...
    """
    ...
    """
    # Error Handling
    # When message is not a string, raise TypeError, "Message must be a string."
    ...

    # When message is empty, raise ValueError, "Message cannot be empty."
    ...

    # When location is not a string, return "Invalid location format."
    ...

    # Return formatted output
    ...
    
# Example usage
print(analyze_tile("Resurrect dead on planet Jupiter", "Philadelphia"))
```

```python
grader.check("toybee-tiles-function")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week5-lecture-not-graded", "4_toybee_tiles")
```

Tests for toybee-tiles-function (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import io
>>> import sys
>>> max_question_points = str(10.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(10.0)
>>> log_variable('total-points', f'Reading-Week-X, 4_toybee_tiles', 10.0)
>>> question_id = 'toybee-tiles-function-1'
>>> max_score = 1.0
>>> score = 0
>>> condition = 'analyze_tile' in globals()
>>> sys.stdout = io.StringIO()
>>> analyze_tile('Test message', 'Test location')
>>> sys.stdout = sys.__stdout__
>>> assert condition, 'The function raised an error when called.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_toybee_tiles', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function encounters an error when called.
success_message: Success: The function executes without errors!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'toybee-tiles-function-2'
>>> max_score = 1.0
>>> score = 0
>>> try:
...     analyze_tile(12345)
... except TypeError as e:
...     condition = str(e) == 'Message must be a string.'
... else:
...     condition = False
>>> assert condition, 'Expected a TypeError with the correct message, but it was not raised.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_toybee_tiles', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not handle non-string messages correctly.
success_message: Success: The function correctly raises an error for non-string messages!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'toybee-tiles-function-3'
>>> max_score = 1.0
>>> score = 0
>>> try:
...     analyze_tile('')
... except ValueError as e:
...     condition = str(e) == 'Message cannot be empty.'
... else:
...     condition = False
>>> assert condition, 'Expected a ValueError with the correct message, but it was not raised.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_toybee_tiles', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not handle empty messages correctly.
success_message: Success: The function correctly raises an error for empty messages!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'toybee-tiles-function-4'
>>> max_score = 2.0
>>> score = 0
>>> output = analyze_tile('Resurrect dead on planet Jupiter', 'Philadelphia')
>>> expected_output = 'Tile Found: Resurrect dead on planet Jupiter\nLocation: Philadelphia'
>>> assert output == expected_output
>>> if output == expected_output:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_toybee_tiles', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not format output correctly.
success_message: Success: The function correctly formats output for valid input!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'toybee-tiles-function-5'
>>> max_score = 2.0
>>> score = 0
>>> output = analyze_tile('Resurrect dead on planet Jupiter')
>>> expected_output = 'Tile Found: Resurrect dead on planet Jupiter\nLocation: Unknown'
>>> assert output == expected_output
>>> if output == expected_output:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_toybee_tiles', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not handle missing 'location' correctly.
success_message: Success: The function correctly handles missing 'location' argument!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'toybee-tiles-function-6'
>>> max_score = 2.0
>>> score = 0
>>> output = analyze_tile('Resurrect dead on planet Jupiter', 12345)
>>> expected_output = 'Invalid location format.'
>>> assert output == expected_output
>>> if output == expected_output:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_toybee_tiles', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not handle non-string location correctly.
success_message: Success: The function correctly handles non-string location input!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'toybee-tiles-function-7'
>>> max_score = 1.0
>>> score = 0
>>> docstring = analyze_tile.__doc__
>>> condition = docstring is not None and 'Args:' in docstring and ('Returns:' in docstring) and ('message (str): The message on the tile.' in docstring) and ("location (str, optional): The tile's location." in docstring) and ('A formatted string describing the tile.' in docstring)
>>> assert condition, "The function's docstring is incomplete."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_toybee_tiles', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not have a complete docstring.
success_message: Success: The function has a complete docstring!
