```python
# Initialize Otter
import otter
grader = otter.Notebook("1_Homework_1_Python_with_The_Office_with_Tests.ipynb")
```

# Homework 1 - Python with "The Office"

![the_office](./assets/figures/the-office.jpeg)

Welcome to Scranton Branch's official Python training! In this homework, you will step into the shoes of your favorite characters from _The Office_. Get ready to solve some "Dunder Mifflin" problems using Python, and maybe, just maybe, impress Michael Scott with your coding prowess.

```{note}
Remember, just like at Dunder Mifflin, effort counts! While tools like ChatGPT can assist you, it’s crucial to understand these concepts. After all, you never know when Dwight might quiz you during lunch.
```


## Question 1: Dwight's Beet Farm Calculations

Dwight needs to calculate the square footage of his beet farm to plan the next season's harvest. Help him by writing Python code to calculate the square root of 789 using the `math` module. Print the result so Dwight can note it down in his "Farm Log."

### Example:

Here’s an example to guide you. Let’s calculate the natural logarithm of 20 for Michael's "World’s Best Boss" mug production:

```python
import math

def example():
    # Calculate and print the natural logarithm of 20
    log_20 = math.log(20)
    print(log_20)

example()
```

Now, write code to calculate and print the square root of 789. Wrap it in a function named `question_1()` as shown in the example.


```python
# Import required packages here

...

# This line of code defines a function
# DO NOT CHANGE THIS LINE OF CODE
def question_1(printing=True):
    # Calculate the square root of 789
    sqrt_789 = ...
    
    if printing:
        # Print the output of the result
        output = ...

    return sqrt_789

# Call the function
sqrt_789 = question_1()
```

```python
grader.check("HW1-Q1-Dwitghts-Beet-Farm")
```

## Question 2: Pam's Art Supplies Budget

Pam is budgeting for her art supplies. She needs to calculate the total cost of 5 canvases priced at $12.99 each. Write a Python program to calculate the total cost and round it to 2 decimal places. Use the `math` module if needed.

1. Define the cost of a single canvas using a variable (`price_per_canvas`).
2. Define the number of canvases (`num_canvases`).
3. Use a formula to calculate the total cost (`price_per_canvas * num_canvases`). Use the `round()` function to round the result to 2 decimal places.
4. Use `print()` to display the total cost.


```python
def question_2():
    # 1. Define the cost of a single canvas using a variable (`price_per_canvas`).
    ...

    # 2. Define the number of canvases (`num_canvases`).
    ...

    # 3. Use a formula to calculate the total cost (`price_per_canvas * num_canvases`). Use the `round()` function to round the result to 2 decimal places.
    ...

    # 4. Use `print()` to display the total cost.
    ...

    # DO NOT CHANGE ANYTHING BELOW THIS LINE
    return locals()

out = question_2()
```

```python
grader.check("HW1-Q2-Pams-Art-Supplies-Budget")
```

## Question 3: Michael's Paper Inventory

Michael Scott needs help keeping track of paper inventory at Dunder Mifflin Scranton. Write a Python program that calculates and prints the following:

1. The total number of paper boxes (`total_boxes`) when combining Premium (`premium_boxes` = 15) and Standard (`standard_boxes` = 23) stock
2. The number of reams (`total_reams`) when each box of Deluxe paper (`deluxe_boxes` = 12) contains 4 reams per box (`reams_per_box` = 4)
3. How many complete client orders (`possible_orders`) can be fulfilled when there are 48 boxes in stock (`total_inventory` = 48) and each order requires 6 boxes (`boxes_per_order` = 6)
   4 The number of leftover boxes (`leftover_boxes`) when trying to evenly distribute 17 boxes (`specialty_boxes` = 17) among 5 salespeople (`num_salespeople` = 5)

The variables must use the names provided in the parentheses, with the values provided in block quotes.


```python
def question_3():
    # Define the number of Premium paper boxes
    ...

    # Define the number of Standard paper boxes
    ...

    # Calculate the total boxes by adding Premium and Standard boxes
    ...

    # Calculate total reams of Deluxe paper (12 boxes with 4 reams per box)
    ...

    # Calculate complete orders possible with 48 boxes when each order needs 6 boxes
    ...

    # Calculate leftover boxes when distributing 17 boxes among 5 salespeople
    ...

    # DO NOT CHANGE ANYTHING BELOW THIS LINE
    
    print(f"Total boxes: {total_boxes}")
    print(f"Total reams: {total_reams}")
    print(f"Possible orders: {possible_orders}")
    print(f"Leftover boxes: {leftover_boxes}")
    
    return locals()

out = question_3()
```

```python
grader.check("HW1-Q3-Michaels-Paper-Inventory")
```

## Question 3: Office Party Games

Michael Scott is planning a casino-themed office party at Dunder Mifflin and needs help creating some Python programs to make it more exciting.

## 1. `roll_dice()` function that:

- Uses Python's `random` module to simulate rolling two six-sided dice. [Python random module](https://docs.python.org/3/library/random.html)
- Prints each roll in this exact format: "Die 1: X, Die 2: Y, Sum: Z" (where X, Y are the individual rolls and Z is their sum). You should use f-string formatting for this. We have not covered f-strings formally yet, but this is a bonus question. [f-string guide](https://www.geeksforgeeks.org/formatted-string-literals-f-strings-python/)
- Returns a tuple of the two dice values in order die1, die2.

## 2. `print_date()` function that:

- Uses Python's `datetime` module to get today's date. [Python datetime module](https://docs.python.org/3/library/datetime.html)
- Prints the date in this exact format: "Today's date is: MM-DD-YYYY" (e.g., "Today's date is: 12-08-2024")
- Returns the datetime.date object

This will help Michael show off his Python skills while planning the office party!


```python
...

def roll_dice():
    die1 = ...
    die2 = ...
    ...
    return die1, die2

def print_date():
    today = ...
    ...
    return today
```

```python
roll_dice()
```

```python
print_date()
```

```python
grader.check("HW1-Bonus-Modules-Madness-with-Michael")
```

Tests for HW1-Bonus-Modules-Madness-with-Michael (points: None)
```
>>> import random
>>> random.seed(162)
>>> out = roll_dice()
>>> assert out == (1, 6)
```
failure_message: Dice rolls are incorrect.
success_message: Dice rolls are correct with seed 162.
```
>>> import io
>>> import sys
>>> import random
>>> random.seed(162)
>>> captured_output = io.StringIO()
>>> sys_stdout_original = sys.stdout
>>> sys.stdout = captured_output
>>> roll_dice()
>>> printed_value = captured_output.getvalue().strip()
>>> sys.stdout = sys_stdout_original
>>> expected_print_value = 'Die 1: 1, Die 2: 6, Sum: 7'
>>> assert printed_value == expected_print_value
```
failure_message: Dice roll output is incorrect.
success_message: Dice roll output is formatted correctly.
```
>>> import io
>>> import sys
>>> captured_output = io.StringIO()
>>> sys_stdout_original = sys.stdout
>>> sys.stdout = captured_output
>>> today = print_date()
>>> sys.stdout = sys_stdout_original
>>> printed_value = captured_output.getvalue().strip()
>>> assert isinstance(today, datetime.date)
>>> assert printed_value == f'Today the date is: {today.strftime('%m-%d-%Y')}'
```
failure_message: Date format is incorrect.
success_message: Date is printed in correct format.


Tests for HW1-Q1-Dwitghts-Beet-Farm (points: None)
```
>>> import numpy as np
>>> a = question_1(printing=False)
>>> assert np.isclose(a, 28.08914, atol=0.001)
```
failure_message: Variables incorrectly defined.
success_message: Variables correctly defined.
```
>>> from unittest.mock import patch
>>> with patch('builtins.print') as mock_print:
...     a = question_1()
...     mock_print.assert_called_once_with(a)
```
failure_message: Your code does not prints the result of sqrt_256.
success_message: Your code prints the result of sqrt_256.


Tests for HW1-Q2-Pams-Art-Supplies-Budget (points: None)
```
>>> import os
>>> original_stdout = sys.stdout
>>> sys.stdout = open(os.devnull, 'w')
>>> out = question_2()
>>> sys.stdout.close()
>>> sys.stdout = original_stdout
>>> price_per_canvas_test = 12.99
>>> assert out['price_per_canvas'] == price_per_canvas_test, f"'price_per_canvas' should be {price_per_canvas_test}, but got {out['price_per_canvas']}"
```
failure_message: The variable `price_per_canvas` is not correctly defined.
success_message: The variable `price_per_canvas` is correctly defined.
```
>>> import os
>>> original_stdout = sys.stdout
>>> sys.stdout = open(os.devnull, 'w')
>>> out = question_2()
>>> sys.stdout.close()
>>> sys.stdout = original_stdout
>>> num_canvases_test = 5
>>> assert 'num_canvases' in out or 'num_canvases' in globals(), "Variable 'num_canvases' is not defined."
>>> assert out['num_canvases'] == num_canvases_test, f"'num_canvases' should be {num_canvases_test}, but got {out['num_canvases']}"
64.95
```
failure_message: The variable `num_canvases` is not correctly defined.
success_message: The variable `num_canvases` is correctly defined.
```
>>> import os
>>> original_stdout = sys.stdout
>>> sys.stdout = open(os.devnull, 'w')
>>> out = question_2()
>>> sys.stdout.close()
>>> sys.stdout = original_stdout
>>> total_cost_test = round(out['price_per_canvas'] * out['num_canvases'], 2)
>>> assert 'total_cost' in out, "Variable 'total_cost' is not defined."
>>> assert out['total_cost'] == total_cost_test, f"'total_cost' should be {total_cost_test}, but got {total_cost}"
64.95
```
failure_message: The total cost is not calculated correctly.
success_message: The total cost is calculated correctly.
```
>>> original_stdout = sys.stdout
>>> sys.stdout = open(os.devnull, 'w')
>>> out = question_2()
>>> sys.stdout.close()
>>> sys.stdout = original_stdout
>>> assert out['total_cost'] == float(f'{out['total_cost']:.2f}'), 'The total cost is not rounded to two decimal places.'
64.95
```
failure_message: The total cost is not rounded correctly to two decimal places.
success_message: The total cost is rounded correctly to two decimal places.
```
>>> import io
>>> import sys
>>> captured_output = io.StringIO()
>>> original_stdout = sys.stdout
>>> sys.stdout = captured_output
>>> question_2()
>>> sys.stdout = sys.stdout
>>> printed_value = captured_output.getvalue().strip()
>>> expected_print_value = round(12.99 * 5, 2)
>>> assert printed_value == str(expected_print_value), f"The print function should output '{expected_print_value}', but got '{printed_value}'."
```
failure_message: The print function was not called with the correct total cost.
success_message: The print function was called with the correct total cost.


Tests for HW1-Q3-Michaels-Paper-Inventory (points: None)
```
>>> out = question_3()
>>> assert out['total_boxes'] == 38, 'Premium (15) + Standard (23) boxes should equal 38'
```
failure_message: The total boxes calculation is incorrect.
success_message: Premium and Standard boxes are correctly added together.
```
>>> out = question_3()
>>> assert out['total_reams'] == 48, '12 boxes * 4 reams should equal 48 total reams'
```
failure_message: Total reams calculation is incorrect.
success_message: Total reams calculation is correct.
```
>>> out = question_3()
>>> assert out['possible_orders'] == 8, '48 boxes divided by 6 boxes per order should equal 8 possible orders'
```
failure_message: Number of possible orders calculation is incorrect.
success_message: Number of possible orders is calculated correctly.
```
>>> out = question_3()
>>> assert out['leftover_boxes'] == 2, '17 boxes modulo 5 salespeople should equal 2 leftover boxes'
```
failure_message: Leftover boxes calculation is incorrect.
success_message: Leftover boxes calculation is correct.
```
>>> import io
>>> import sys
>>> captured_output = io.StringIO()
>>> original_stdout = sys.stdout
>>> sys.stdout = captured_output
>>> question_3()
>>> sys.stdout = original_stdout
>>> printed_value = captured_output.getvalue().strip()
>>> expected_print_value = 'Total boxes: 38\nTotal reams: 48\nPossible orders: 8\nLeftover boxes: 2'
>>> assert printed_value == expected_print_value, f"The print function should output '{expected_print_value}', but got '{printed_value}'."
```
failure_message: The printed output does not match the expected inventory values.
success_message: All inventory values were printed correctly.
