```python
import os
os.environ['JUPYTERHUB_USER'] = '/usr/bin/python3'

# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("2_week_1_attendance", "week_2", "attendance", assignment_points = 1.0, assignment_tag = 'week2-attendance')

# Initialize Otter
import otter
grader = otter.Notebook("2_week_1_attendance.ipynb")
```

# ðŸ—“ï¸ðŸ“š Week 2: Attendance âœ…

![](./assets/figures/scrabble-attendance.webp)

To get credit for attendance make sure you run all blocks of code. 

```python
I_attended_lab = True
```

```python
grader.check("week-1-lab-attendance")
```

## ðŸŽ‰ Well Done!

Keep up the good work.

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week2-attendance", "2_week_1_attendance")
```

Tests for week-1-lab-attendance (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(1.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(1.0)
>>> log_variable('total-points', f'Reading-Week-X, 2_week_1_attendance', 1.0)
>>> question_id = 'week-1-lab-attendance-1'
>>> max_score = 1.0
>>> score = 0
>>> assert I_attended_lab, 'You must attend the lab to get credit for this assignment.'
>>> if I_attended_lab:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_week_1_attendance', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Output does not include 'Each friend gets'.
success_message: Correctly included the phrase 'Each friend gets'.
