```python
from pykubegrader.tokens.validate_token import validate_token
validate_token('type the key provided by your instructor here', assignment = 'week8-quiz')

```

```python
validate_token(assignment = 'week8-quiz')

# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_quiz_q", "week_8", "quiz", assignment_points = 12.0, assignment_tag = 'week8-quiz')

# Initialize Otter
import otter
grader = otter.Notebook("1_quiz_q.ipynb")
```

# â“Quiz: Implicit Equation Plotting

## ðŸ“Œ Mathematical Equations

$$
\text{eq1:} \quad \left(\frac{x}{7}\right)^2 \sqrt{\frac{\left| |x| - 3 \right|}{|x| - 3}} + \left(\frac{y}{3}\right)^2 \sqrt{\frac{\left| y + \frac{3}{7}\sqrt{33} \right|}{y + \frac{3}{7}\sqrt{33}}} - 1
$$

$$
\text{eq2:} \quad \left| \frac{x}{2} \right| - \frac{3\sqrt{33} - 7}{112}x^2 - 3 + \sqrt{1 - \left(| |x| - 2 | - 1\right)^2} - y
$$

$$
\text{eq3:} \quad 9\sqrt{\frac{\left| (|x| - 1)(|x| - 0.75) \right|}{(1 - |x|)(|x| - 0.75)}} - 8|x| - y
$$


```python
# Importing the necessary libraries
# The following libraries are used in the code below
# import sqrt, meshgrid, arange from numpy
# pyplot from matplotlib as plt
# division from __future__ to make sure that the division is float division
# The code below is written in Python 3
...

# Define the grid range
xs = arange(-7.25, 7.25, 0.01)
ys = arange(-5, 5, 0.01)
x, y = meshgrid(xs, ys)

# Predefine the equations
eq4 = (3 * abs(x) + 0.75 * sqrt(abs((abs(x) - 0.75) * (abs(x) - 0.5)) / ((0.75 - abs(x)) * (abs(x) - 0.5))) - y)
eq5 = (2.25 * sqrt(abs((x - 0.5) * (x + 0.5)) / ((0.5 - x) * (0.5 + x))) - y)
eq6 = (6 * sqrt(10) / 7 + (1.5 - 0.5 * abs(x)) * sqrt(abs(abs(x) - 1) / (abs(x) - 1)) - (6 * sqrt(10) / 14) * sqrt(4 - (abs(x) - 1)**2) - y)
```

```python
grader.check("q1-import-libraries")
```


## Question 1 (Points: 2.0): ðŸ“Œ Mathematical Equations
### Computing the first equation
$$
\text{eq1:} \quad \left(\frac{x}{7}\right)^2 \sqrt{\frac{\left| |x| - 3 \right|}{|x| - 3}} + \left(\frac{y}{3}\right)^2 \sqrt{\frac{\left| y + \frac{3}{7}\sqrt{33} \right|}{y + \frac{3}{7}\sqrt{33}}} - 1
$$


```python
# Define the function to compute the equation 1 as compute_equation_1 with x and y as parameters
# The function should return the equation 1 as eq1
...
    """Compute the equation for contour plotting."""
    ...
    # Return the equation 1 as eq1
    ...
# List of equations to plot
# Compute the equation 1 as plot_equation_1
...
# List of equations to plot as plot_equations
...
# Generate contour plots for the equations as eq in plot_equations using plt.contour
# The contour should be generated for x, y and eq with the level set to 0
...
```

```python
grader.check("q2-compute-the-first-equation")
```

## Question 2 (Points: 2.0): Compute for equation 2
$$
\text{eq2:} \quad \left| \frac{x}{2} \right| - \frac{3\sqrt{33} - 7}{112}x^2 - 3 + \sqrt{1 - \left(| |x| - 2 | - 1\right)^2} - y
$$



```python
# Define the function to compute the equation 2 as compute_equation_2 with x and y as parameters
# The function should return the equation 2 as eq2
...
    """Compute the equation for contour plotting."""
    ...
    # Return the equation 2 as eq2
    ...

# Compute the equation 2 as plot_equation_2
...
# List of equations to plot as plot_equations with plot_equation_1 and plot_equation_2 
...
# Generate contour plots for the equations as eq in plot_equations using plt.contour
# The contour should be generated for x, y and eq with the level set to 0
...

```

```python
grader.check("q3-compute-the-second-equation")
```

## Question 4 (Points: 6.0): Question 3 (Points: 2.0): Compute Equation 3
$$
\text{eq3:} \quad 9\sqrt{\frac{\left| (|x| - 1)(|x| - 0.75) \right|}{(1 - |x|)(|x| - 0.75)}} - 8|x| - y
$$

```python
# Define the function to compute the equation 3 as compute_equation_3 with x and y as parameters
# The function should return the equation 3 as eq3
...
    """Compute the equation for contour plotting."""
    ...
    # Return the equation 3 as eq3
    ...

# Compute the equation 3 as plot_equation_3
...
# List of equations to plot as plot_equations with plot_equation_1, plot_equation_2 and plot_equation_3
...
# Generate contour plots for the equations as eq in plot_equations using plt.contour
# The contour should be generated for x, y and eq with the level set to 0
...

```

```python
grader.check("q4-compute-the-third-equation")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.tokens.validate_token import validate_token
validate_token(assignment = 'week8-quiz')


from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week8-quiz", "1_quiz_q")
```

Tests for q1-import-libraries (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> import sys
>>> imported_modules = set(sys.modules.keys())
>>> max_question_points = str(2.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(12.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> log_variable('total-points', f'week8-quiz, 1_quiz_q', 12.0)
>>> question_id = 'q1-import-libraries-1'
>>> max_score = 2.0
>>> score = 0
>>> required_libraries = {'matplotlib.pyplot', 'numpy', 'math'}
>>> missing_libraries = required_libraries - imported_modules
>>> assert not missing_libraries, f'Missing imports: {missing_libraries}. Ensure all required libraries are imported.'
>>> if not missing_libraries:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(imported_modules))
```
failure_message: Failed: One or more required libraries or typing imports are missing.
success_message: Success: All required libraries and typing imports are correctly imported!


Tests for q2-compute-the-first-equation (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> max_question_points = str(2.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(12.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> log_variable('total-points', f'week8-quiz, 1_quiz_q', 12.0)
>>> question_id = 'q2-compute-the-first-equation-1'
>>> max_score = 2.0
>>> score = 0
>>> eq1_edge_cases = [compute_equation_1(x, 3) for x in [-7, 0, 7]]
>>> condition = all((isinstance(val, float) for val in eq1_edge_cases))
>>> assert condition
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(eq1_edge_cases))
```
failure_message: Failed: `compute_equation_1()` does not handle edge cases correctly.
success_message: Success: `compute_equation_1()` handles edge cases correctly!


Tests for q3-compute-the-second-equation (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> max_question_points = str(2.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(12.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> log_variable('total-points', f'week8-quiz, 1_quiz_q', 12.0)
>>> question_id = 'q3-compute-the-second-equation-1'
>>> max_score = 2.0
>>> score = 0
>>> eq2_edge_cases = [compute_equation_2(x, -3) for x in [-4, 0, 4]]
>>> condition = all((isinstance(val, float) for val in eq2_edge_cases))
>>> assert condition
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(eq2_edge_cases))
```
failure_message: Failed: `compute_equation_2()` does not handle edge cases correctly.
success_message: Success: `compute_equation_2()` handles edge cases correctly!


Tests for q4-compute-the-third-equation (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> max_question_points = str(6.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(12.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> log_variable('total-points', f'week8-quiz, 1_quiz_q', 12.0)
>>> question_id = 'q4-compute-the-third-equation-1'
>>> max_score = 2.0
>>> score = 0
>>> num_contours = len(plot_equations)
>>> condition = num_contours >= 0
>>> assert condition
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(plot_equations))
>>> responses = update_responses(question_id, str(num_contours))
```
failure_message: Failed: Contour plots are not being generated properly.
success_message: Success: Contour plots are correctly generated!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> question_id = 'q4-compute-the-third-equation-2'
>>> max_score = 2.0
>>> score = 0
>>> output_types = [isinstance(eq, type(x)) for eq in plot_equations]
>>> condition = all(output_types)
>>> assert condition
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(output_types))
```
failure_message: Failed: Contour equations do not return arrays!
success_message: Success: Contour equations return arrays!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> question_id = 'q4-compute-the-third-equation-3'
>>> max_score = 2.0
>>> score = 0
>>> num_equations = len([plot_equation_1, plot_equation_2, plot_equation_3, eq4, eq5, eq6])
>>> condition = num_equations == 6
>>> assert condition
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(num_equations))
```
failure_message: Failed: Incorrect number of equations defined!
success_message: Success: Correct number of equations are defined!
