```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_practice_quiz_data_types_q", "week_2", "practicequiz", assignment_points = 14.5, assignment_tag = 'week2-practicequiz')

# Initialize Otter
import otter
grader = otter.Notebook("1_practice_quiz_data_types_q.ipynb")
```

# â“ Python Basics: Data Structures and Types Quiz ðŸ

![Python Quiz](./assets/figures/python-string-practice-quiz.png)

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_practice_quiz_data_types_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_practice_quiz_data_types_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_practice_quiz_data_types_q import Question3
Question3().show()

```

## Wastewater Treatment Analysis ðŸŒŠ

### Instructions:

Create a Python dictionary to analyze wastewater treatment data. Your solution must include:

1. Create a dictionary named `treatment_data` with these EXACT keys and values:
   - `facility`: String "Plant-2000"
   - `capacity`: Float 2.0
   - `test_ph_levels`: List [6.8, 7.2, 7.1, 7.4, 6.9]
   - `performance_metrics`: Dictionary containing:
     - `flow_rate`: Float 185.5
     - `efficiency`: Float 92.8  # percent

2. Calculate these values (store each in its own variable):
   - `avg_ph`: Average of test_ph_levels (use `sum` and `len`)
   - `max_ph`: Maximum pH from test_ph_levels (use `max`)
   - `capacity_utilization`: Flow rate divided by capacity

```python
# Create the treatment data dictionary with exact values
...


# Calculate required values
...

# We have provided the code to print the results for you
# Print results for verification (optional but helpful)
print(f"Average pH: {avg_ph:.1f}")
print(f"Maximum pH: {max_ph}")
print(f"Capacity Utilization: {capacity_utilization:.1f} L/s")
```

```python
grader.check("question-Wastewater-Data")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week2-practicequiz", "1_practice_quiz_data_types_q")
```

Tests for question-Wastewater-Data (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(6.5)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(6.5)
>>> log_variable('total-points', f'Reading-Week-X, 1_practice_quiz_data_types_q', 6.5)
>>> question_id = 'question-Wastewater-Data-1'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, isinstance(treatment_data, dict))
>>> assert isinstance(treatment_data, dict), 'treatment_data must be a dictionary'
>>> if isinstance(treatment_data, dict):
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_quiz_data_types_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Not a dictionary - make sure you create treatment_data as a dictionary
success_message: Dictionary type is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Wastewater-Data-2'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, str(type(treatment_data['facility'])))
>>> assert isinstance(treatment_data['facility'], str), 'facility must be a string'
>>> if isinstance(treatment_data['facility'], str):
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_quiz_data_types_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The 'facility' value must be a string (str)
success_message: Facility data type is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Wastewater-Data-3'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, str(type(treatment_data['capacity'])))
>>> assert isinstance(treatment_data['capacity'], float), 'capacity must be a float'
>>> if isinstance(treatment_data['capacity'], float):
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_quiz_data_types_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The 'capacity' value must be a float (not an int)
success_message: Capacity data type is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Wastewater-Data-4'
>>> max_score = 0.5
>>> score = 0
>>> assert isinstance(treatment_data['test_ph_levels'], list), 'test_ph_levels must be a list'
>>> if isinstance(treatment_data['test_ph_levels'], list):
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_quiz_data_types_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The 'test_ph_levels' value must be a list
success_message: Test pH levels data type is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Wastewater-Data-5'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, str(type(treatment_data['performance_metrics'])))
>>> assert isinstance(treatment_data['performance_metrics'], dict), 'performance_metrics must be a dictionary'
>>> if isinstance(treatment_data['performance_metrics'], dict):
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_quiz_data_types_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The 'performance_metrics' value must be a dictionary
success_message: Performance metrics data type is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Wastewater-Data-6'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, treatment_data['facility'])
>>> assert treatment_data['facility'] == 'Plant-2000', "facility must be exactly 'Plant-2000'"
>>> if treatment_data['facility'] == 'Plant-2000':
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_quiz_data_types_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The facility must be exactly 'Plant-2000' (check spelling and capitalization)
success_message: Facility value is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Wastewater-Data-7'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, abs(treatment_data['capacity'] - 2.0))
>>> assert abs(treatment_data['capacity'] - 2.0) < 0.001, 'capacity must be exactly 2.0'
>>> if abs(treatment_data['capacity'] - 2.0) < 0.001:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_quiz_data_types_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The capacity must be exactly 2.0 (as a float)
success_message: Capacity value is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Wastewater-Data-8'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, treatment_data['test_ph_levels'])
>>> conditional = treatment_data['test_ph_levels'] == [6.8, 7.2, 7.1, 7.4, 6.9]
>>> assert conditional, 'test_ph_levels must match exactly'
>>> if conditional:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_quiz_data_types_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The test_ph_levels list must be exactly [6.8, 7.2, 7.1, 7.4, 6.9]
success_message: Test pH levels list is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Wastewater-Data-9'
>>> max_score = 0.5
>>> score = 0
>>> assert abs(treatment_data['performance_metrics']['flow_rate'] - 185.5) < 0.001, 'flow_rate must be exactly 185.5'
>>> if abs(treatment_data['performance_metrics']['flow_rate'] - 185.5) < 0.001:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_quiz_data_types_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The flow_rate must be exactly 185.5 (as a float)
success_message: Flow rate value is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Wastewater-Data-10'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, abs(treatment_data['performance_metrics']['efficiency']))
>>> assert abs(treatment_data['performance_metrics']['efficiency'] - 92.8) < 0.001, 'efficiency must be exactly 92.8'
>>> if abs(treatment_data['performance_metrics']['efficiency'] - 92.8) < 0.001:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_quiz_data_types_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The efficiency must be exactly 92.8 (as a float)
success_message: Efficiency value is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Wastewater-Data-11'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, avg_ph)
>>> assert isinstance(avg_ph, float), 'avg_ph must be a float'
>>> assert abs(avg_ph - 7.08) < 0.001, 'avg_ph calculation incorrect - should be 7.08'
>>> if isinstance(avg_ph, float) and abs(avg_ph - 7.08) < 0.001:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_quiz_data_types_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: avg_ph must be a float type (use division with sum and len)
success_message: Average pH data type is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Wastewater-Data-12'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, max_ph)
>>> assert isinstance(max_ph, float), 'max_ph must be a float'
>>> assert max_ph == 7.4, 'max_ph calculation incorrect - should be 7.4'
>>> if isinstance(max_ph, float) and max_ph == 7.4:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_quiz_data_types_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: max_ph must be a float with value 7.4 (use max() function)
success_message: Maximum pH data type and value are correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Wastewater-Data-13'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, capacity_utilization)
>>> assert isinstance(capacity_utilization, float), 'capacity_utilization must be a float'
>>> assert abs(capacity_utilization - 92.75) < 0.1, 'capacity_utilization calculation incorrect - should be 92.75'
>>> if isinstance(capacity_utilization, float) and abs(capacity_utilization - 92.75) < 0.1:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practice_quiz_data_types_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: capacity_utilization must be a float equal to 92.75 (flow_rate/capacity)
success_message: Capacity utilization calculation is correct
