```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("5_classes_q", "week_7", "lecturenotgraded", assignment_points = 21.0, assignment_tag = 'week7-lecturenotgraded')

# Initialize Otter
import otter
grader = otter.Notebook("5_classes_q.ipynb")
```

# 💻 Engineering Material Simulation 🚀🔧

![](./assets/figures/ai-simulation-laptop-1536x919-1.webp)


## Problem Statement
You are designing a Python-based simulation for tracking the stress-strain behavior of different materials under mechanical loading. The simulation should be object-oriented and include the following functionalities:


1. Class and Instances:
    - Define a class `Material` to represent different engineering materials.
    - Each material instance should have properties such as `name`, `elastic_modulus`, `yield_strength`, and `strain` (default: `0.0`).


2. Change Methods with Error Handling:
    - Implement a method `apply_strain(self, strain_value: float)` to modify the strain.
    - Ensure that the strain does not go below zero and does not exceed a predefined fracture strain (`0.05` for all materials). If an invalid strain is applied, raise a `ValueError`.


3. Magic Methods:
    - Implement the `str` magic method to print a readable summary of the material properties.
    - Implement the `add` method to allow combining two materials into a new `CompositeMaterial` object with an average of their properties.


4. Custom Iterator:
    - Create a class `MaterialIterator` to iterate through a list of materials and return their names one by one.


## Task
Write the `Material` and `MaterialIterator` classes, ensuring they adhere to the problem constraints. Then, demonstrate the following:
- Create instances for `Steel` and `Aluminum`.
- Apply valid and invalid strain values and handle errors appropriately.
- Print the material properties using `str`.
- Combine `Steel` and `Aluminum` using `+` and print the new composite material.
- Use `MaterialIterator` to iterate through a list of materials and display their names.


```python
# Step 1: Define the Material class
class Material:

    # Step 2: add the __init__ method to initialize the Material class, with the following parameters:
    # name: (str) Name of the material
    # elastic_modulus: (float) Elastic modulus of the material in GPa
    # yield_strength: (float) Yield strength of the material in MPa
    # assign these parameters to the corresponding class attributes
    # Add placeholders for strain and fracture_strain attributes with default values of 0.0 and 0.05, respectively.
    ...
        """
        ...
        """
        ...

    # Step 3: add the apply_strain method to the Material class, with the following parameters:
    # strain_value: (float) Strain value to apply
    # Check if the strain_value is negative and raise a ValueError with the message "Strain cannot be negative!"
    # Check if the strain_value exceeds the fracture_strain attribute and raise a ValueError with the message "Strain exceeds fracture limit!"
    # If the strain_value is valid, update the strain attribute with the strain_value
    ...
        """
        Apply strain to the material while ensuring it remains within valid limits.
        :param strain_value: Strain value to apply
        """
        ...

    # Step 4: add the __str__ method to the Material class to return a readable string representation of the material.
    # The string should include the name, elastic modulus, yield strength, and current strain of the material.
    # here is an example of the expected output:
    # Material: Steel
    # Elastic Modulus: 200 GPa
    # Yield Strength: 250 MPa
    # Current Strain: 0.0200
    ...
        """
        Magic method to return a readable string representation of the material.
        """
        ...

    # Step 5: add the __add__ method to the Material class to combine two materials into a new composite material.
    # The method should return a new Material instance with averaged properties of the two materials.
    # The name of the composite material should be "Composite(Material1 + Material2)"
    # The elastic modulus and yield strength of the composite material should be the average of the two materials.
    # If the `other`` object is not an instance of Material, raise a TypeError with the message "Can only add another Material instance!"
    ...
        """
        Magic method to combine two materials into a new composite material.
        :param other: Another Material instance
        :return: CompositeMaterial with averaged properties
        """
        # Check if the other object is a Material instance if not return a TypeError with the message "Can only add another Material instance!"
        ...

        # Create a composite material with averaged properties
        # make the composite name as "Composite(Material1 + Material2)"
        ...

        # Calculate the average of elastic modulus and yield strength as composite_modulus and composite_yield_strength
        ...

        # Return a new Material instance with the composite properties
        ...


# Step 2: Define the MaterialIterator class


class MaterialIterator:

    # Step 3: add the __init__ method to initialize the MaterialIterator class, with the following parameters:
    # materials: (list) List of Material instances
    # assign the materials parameter to the materials attribute
    # initialize the index attribute with a value of 0
    ...
        """
        Initialize the iterator with a list of materials.
        :param materials: List of Material instances
        """
        ...

    # Step 4: add the __iter__ method to the MaterialIterator class to return the iterator object itself.
    ...

    # Step 5: add the __next__ method to the MaterialIterator class to return the next material name or stop iteration.
    # If the index is greater than or equal to the length of materials, raise a StopIteration exception.
    # Get the name of the current material using the index attribute.
    # Increment the index by 1 to move to the next material.
    # Return the material name.
    ...


# Step 3: Demonstrate Functionality


# Create instances of materials with name steel with elastic modulus 200 GPa and yield strength 250 MPa
...

# Create another instance of material with name aluminum with elastic modulus 70 GPa and yield strength 150 MPa
...

# Apply valid strain
try:
    steel.apply_strain(0.02)
except ValueError as e:
    print(f"Error: {e}")

# Attempt to apply invalid strain
try:
    aluminum.apply_strain(0.06)  # Exceeds fracture limit
except ValueError as e:
    print(f"Error: {e}")

# Print material properties using __str__
print(steel)
print(aluminum)

# Combine materials using __add__, and save to a new variable `composite``
composite = ...
print(composite)  # Composite material

# Use MaterialIterator to iterate through a list of materials
materials = [steel, aluminum, composite]
iterator = ...

print("\nIterating through materials:")
for material_name in iterator:
    print(material_name)
```

```python
grader.check("Materials-simulation-classes")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week7-lecturenotgraded", "5_classes_q")
```

Tests for Materials-simulation-classes (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(21.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(21.0)
>>> log_variable('total-points', f'week7-lecturenotgraded, 5_classes_q', 21.0)
>>> question_id = 'Materials-simulation-classes-1'
>>> max_score = 1.0
>>> score = 0
>>> class_exists = 'Material' in globals() and 'MaterialIterator' in globals()
>>> assert class_exists, 'Both `Material` and `MaterialIterator` classes must be defined.'
>>> if class_exists:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_classes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(class_exists))
```
failure_message: Failed: One or both classes (`Material`, `MaterialIterator`) are missing.
success_message: Success: `Material` and `MaterialIterator` classes are defined!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Materials-simulation-classes-2'
>>> max_score = 2.0
>>> score = 0
>>> test_material = Material('Titanium', 120, 400)
>>> material_attributes = {'name': test_material.name, 'elastic_modulus': test_material.elastic_modulus, 'yield_strength': test_material.yield_strength, 'strain': test_material.strain, 'fracture_strain': test_material.fracture_strain}
>>> expected_attributes = {'name': 'Titanium', 'elastic_modulus': 120, 'yield_strength': 400, 'strain': 0.0, 'fracture_strain': 0.05}
>>> assert material_attributes == expected_attributes, f'Expected {expected_attributes}, but got {material_attributes}.'
>>> if material_attributes == expected_attributes:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_classes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(material_attributes))
```
failure_message: Failed: `Material` does not initialize attributes correctly.
success_message: Success: `Material` initializes attributes correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Materials-simulation-classes-3'
>>> max_score = 2.0
>>> score = 0
>>> test_material = Material('Copper', 110, 350)
>>> test_material.apply_strain(0.02)
>>> updated_strain = test_material.strain
>>> expected_strain = 0.02
>>> assert updated_strain == expected_strain, f'Expected strain to be {expected_strain}, but got {updated_strain}.'
>>> if updated_strain == expected_strain:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_classes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(updated_strain))
```
failure_message: Failed: `apply_strain()` does not update strain correctly.
success_message: Success: `apply_strain()` correctly updates strain!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import pytest
>>> question_id = 'Materials-simulation-classes-4'
>>> max_score = 3.0
>>> score = 0
>>> test_material = Material('Brass', 90, 300)
>>> error_raised = False
>>> try:
...     test_material.apply_strain(-0.01)
... except ValueError as e:
...     error_raised = str(e) == 'Strain cannot be negative!'
>>> assert error_raised, 'Expected `ValueError: Strain cannot be negative!`, but no error was raised.'
>>> error_raised = False
>>> try:
...     test_material.apply_strain(0.06)
... except ValueError as e:
...     error_raised = str(e) == 'Strain exceeds fracture limit!'
>>> assert error_raised, 'Expected `ValueError: Strain exceeds fracture limit!`, but no error was raised.'
>>> if error_raised and error_raised:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_classes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(error_raised))
```
failure_message: Failed: `apply_strain()` does not handle invalid strains correctly.
success_message: Success: `apply_strain()` correctly handles invalid strains!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Materials-simulation-classes-5'
>>> max_score = 2.0
>>> score = 0
>>> test_material = Material('Bronze', 110, 250)
>>> material_string = str(test_material)
>>> expected_string = 'Material: Bronze\nElastic Modulus: 110 GPa\nYield Strength: 250 MPa\nCurrent Strain: 0.0000'
>>> assert material_string == expected_string, f'Expected:\n{expected_string}\nbut got:\n{material_string}'
>>> if material_string == expected_string:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_classes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(material_string))
```
failure_message: Failed: `__str__()` does not return the expected string.
success_message: Success: `__str__()` correctly represents a material!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Materials-simulation-classes-6'
>>> max_score = 3.0
>>> score = 0
>>> material1 = Material('Nickel', 150, 500)
>>> material2 = Material('Zinc', 90, 200)
>>> composite = material1 + material2
>>> composite_properties = {'name': composite.name, 'elastic_modulus': composite.elastic_modulus, 'yield_strength': composite.yield_strength}
>>> expected_properties = {'name': 'Composite(Nickel + Zinc)', 'elastic_modulus': 120, 'yield_strength': 350}
>>> assert composite_properties == expected_properties, f'Expected {expected_properties}, but got {composite_properties}.'
>>> if composite_properties == expected_properties:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_classes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(composite_properties))
```
failure_message: Failed: `__add__()` does not combine materials correctly.
success_message: Success: `__add__()` correctly creates composite materials!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Materials-simulation-classes-7'
>>> max_score = 1.0
>>> score = 0
>>> test_material = Material('Titanium', 120, 400)
>>> test_material.apply_strain(0.005)
>>> updated_strain = test_material.strain
>>> expected_strain = 0.005
>>> assert updated_strain == expected_strain, f'Expected strain to be {expected_strain}, but got {updated_strain}.'
>>> if updated_strain == expected_strain:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_classes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(updated_strain))
```
failure_message: Failed: `apply_strain()` does not correctly update strain for small values.
success_message: Success: `apply_strain()` correctly updates strain for small values!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Materials-simulation-classes-8'
>>> max_score = 3.0
>>> score = 0
>>> materials = [Material('Iron', 200, 250), Material('Lead', 40, 150), Material('Silver', 70, 300)]
>>> iterator = MaterialIterator(materials)
>>> iterated_names = list(iterator)
>>> expected_names = ['Iron', 'Lead', 'Silver']
>>> assert iterated_names == expected_names, f'Expected {expected_names}, but got {iterated_names}.'
>>> if iterated_names == expected_names:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_classes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(iterated_names))
```
failure_message: Failed: `MaterialIterator` does not iterate correctly.
success_message: Success: `MaterialIterator` correctly iterates through materials!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Materials-simulation-classes-9'
>>> max_score = 1.0
>>> score = 0
>>> test_material = Material('Graphene', 200, 500)
>>> test_material.apply_strain(0.05)
>>> updated_strain = test_material.strain
>>> expected_strain = 0.05
>>> assert updated_strain == expected_strain, f'Expected strain to be {expected_strain}, but got {updated_strain}.'
>>> if updated_strain == expected_strain:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_classes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(updated_strain))
```
failure_message: Failed: `apply_strain()` does not handle strain at the fracture limit.
success_message: Success: `apply_strain()` correctly handles strain at the fracture limit!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Materials-simulation-classes-10'
>>> max_score = 1.0
>>> score = 0
>>> copper = Material('Copper', 120, 400)
>>> zinc = Material('Zinc', 90, 300)
>>> composite = copper + zinc
>>> composite_string = str(composite)
>>> expected_string = 'Material: Composite(Copper + Zinc)\nElastic Modulus: 105.0 GPa\nYield Strength: 350.0 MPa\nCurrent Strain: 0.0000'
>>> assert composite_string == expected_string, f'Expected:\n{expected_string}\nbut got:\n{composite_string}'
>>> if composite_string == expected_string:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_classes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(composite_string))
```
failure_message: Failed: `__str__()` does not correctly represent composite materials.
success_message: Success: `__str__()` correctly represents composite materials!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Materials-simulation-classes-11'
>>> max_score = 1.0
>>> score = 0
>>> nickel = Material('Nickel', 180, 450)
>>> iron = Material('Iron', 200, 550)
>>> composite = nickel + iron
>>> composite_type = isinstance(composite, Material)
>>> assert composite_type, 'Expected `__add__()` to return an instance of `Material`, but it returned a different type.'
>>> if composite_type:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_classes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(composite_type))
```
failure_message: Failed: `__add__()` does not return a `Material` instance.
success_message: Success: `__add__()` correctly returns a `Material` instance!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Materials-simulation-classes-12'
>>> max_score = 1.0
>>> score = 0
>>> single_material = [Material('Silver', 70, 300)]
>>> iterator = MaterialIterator(single_material)
>>> iterated_name = next(iterator)
>>> expected_name = 'Silver'
>>> assert iterated_name == expected_name, f'Expected `{expected_name}`, but got `{iterated_name}`.'
>>> if iterated_name == expected_name:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_classes_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(iterated_name))
```
failure_message: Failed: `MaterialIterator` does not correctly iterate over a single material.
success_message: Success: `MaterialIterator` correctly iterates over a single material!
