```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("7_python_tools_q", "week_3", "readings", assignment_points = 15.0, assignment_tag = 'week3-readings')

# Initialize Otter
import otter
grader = otter.Notebook("7_python_tools_q.ipynb")
```

# â“ðŸ Python Tools for Engineering ðŸ› ï¸

![](./assets/figures/python-tools-for-engineering.webp)

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._7_python_tools_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._7_python_tools_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._7_python_tools_q import Question3
Question3().show()

```

## NumPy Matrix Determinant

### Instructions:
Write a Python function using NumPy that computes the determinant of a 2x2 matrix. The function should take a 2D list as input, convert it to a NumPy array, and return the determinant.

We have provided you with the structure for the function `matrix_determinant(matrix)` since you do not know that yet. You need to complete the function by adding the necessary code to calculate the determinant.

Your solution must include:
1. A function `matrix_determinant(matrix)` that:
   - Accepts a 2D list (e.g., `[[1, 2], [3, 4]]`).
   - Converts it into a NumPy array. Using the `np.array()` function. You can save this as the variable `np_matrix`.
   - Calculates and returns the determinant. Using the `np.linalg.det()` function. You must save this as the variable `determinant`.

```{tip}
If you do not know how any of the functions mentioned above work, you can look them up in the documentation help(function_name). This is a good test taking strategy as you will have access to the documentation during the quizzes and tests.
```
2. Test cases to validate the correctness of your implementation.

```python
# Import numpy here, make sure to use the common alias np.
...

def matrix_determinant(matrix):
    # Convert the matrix to a numpy array and return the determinant
    # make sure to save it to the variable `np_matrix`
    ...
    
    # Compute the determinant of the matrix
    # make sure to save it to the variable `determinant`
    ...
    return determinant


# Example test case
result = matrix_determinant([[1, 2], [3, 4]])
print(f"Determinant: {result}")
```

```python
grader.check("numpy-matrix-determinant")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week3-readings", "7_python_tools_q")
```

Tests for numpy-matrix-determinant (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(5.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(5.0)
>>> log_variable('total-points', f'Reading-Week-X, 7_python_tools_q', 5.0)
>>> question_id = 'numpy-matrix-determinant-1'
>>> max_score = 1.0
>>> score = 0
>>> condition = 'np' in locals()
>>> assert condition, "You must import numpy using the alias 'np'."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('7_python_tools_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(condition))
```
failure_message: Failed: You must import numpy using the alias 'np'
success_message: Success: 'numpy' is imported as 'np'
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> import inspect
>>> question_id = 'numpy-matrix-determinant-2'
>>> max_score = 1.0
>>> score = 0
>>> code = inspect.getsource(matrix_determinant)
>>> np_matrix_pattern = 'np_matrix\\s*=\\s*np\\.array\\s*\\(\\s*matrix\\s*\\)'
>>> np_matrix_creation_found = bool(re.search(np_matrix_pattern, code))
>>> assert np_matrix_creation_found, 'You must define `np_matrix` with `np_matrix = np.array(matrix)`.'
>>> if np_matrix_creation_found:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('7_python_tools_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(np_matrix_creation_found))
```
failure_message: Failed: You must create 'np_matrix' using np.array(matrix)
success_message: Success: 'np_matrix' is created from np.array(matrix)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> import inspect
>>> question_id = 'numpy-matrix-determinant-3'
>>> max_score = 1.0
>>> score = 0
>>> code = inspect.getsource(matrix_determinant)
>>> det_pattern = 'determinant\\s*=\\s*np\\.linalg\\.det\\s*\\(\\s*np_matrix\\s*\\)'
>>> det_pattern_found = bool(re.search(det_pattern, code))
>>> assert det_pattern_found, 'You must define `determinant` with `determinant = np.linalg.det(np_matrix)`.'
>>> if det_pattern_found:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('7_python_tools_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(det_pattern_found))
```
failure_message: Failed: You must compute 'determinant' using np.linalg.det(np_matrix)
success_message: Success: 'determinant' is computed from np.linalg.det(np_matrix)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'numpy-matrix-determinant-4'
>>> max_score = 2.0
>>> score = 0
>>> matrix_1 = [[1, 8], [3, 7]]
>>> expected_det_1 = -17
>>> computed_det = matrix_determinant(matrix_1)
>>> assert abs(computed_det - expected_det_1) < 1e-12, f'Expected determinant {expected_det_1}, but got {computed_det}'
>>> if abs(computed_det - expected_det_1) < 1e-12:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('7_python_tools_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(computed_det))
```
failure_message: Failed: The determinant is incorrect
success_message: Success: Function returns the correct determinant
