```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("17_operators_q", "week_1", "readings", assignment_points = 18.0, assignment_tag = 'week1-readings')

# Initialize Otter
import otter
grader = otter.Notebook("17_operators_q.ipynb")
```

# â“Python Operators

![](./assets/figures/dragon-dining.webp)

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._17_operators_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._17_operators_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._17_operators_q import Question3
Question3().show()

```

## Instructions

**"Mario the Dragon and the Handschumacher Dining Center Feast"**

Mario, the dragon mascot of Drexel, loves the Handschumacher Dining Center. Every week, he calculates his meal consumption to plan his legendary feasts. Mario knows he eats **21 meals per week** (3 meals per day for 7 days). However, Mario is curious about how his meal plan aligns with his love of sharing.

Write a function named `mario_meal_plan()` that:
1. **Calculates the total number of meals Mario eats in a week.**
2. **Determines if the total meals can be split evenly among his 6 dragon friends** (including himself).  
3. **Returns a string with the results**. If the meals are evenly distributed, the string should say:  
   `"Mario's meals divide evenly among 6 dragons, with each dragon eating X meals."`  
   Otherwise, the string should say:  
   `"Mario's meals do not divide evenly among 6 dragons, leaving Y extra meals."`  

### Example Outputs:
- `"Mario's meals divide evenly among 6 dragons, with each dragon eating 3 meals."`
- `"Mario's meals do not divide evenly among 6 dragons, leaving 1 extra meal."`


```python
# Calculate total meals Mario eats in a week, save as a variable `total_meals`
total_meals = ...

# Number of dragons including Mario, save as a variable `num_dragons`
num_dragons = ...

# Calculate meals per dragon and any remainder 
# names `meals_per_dragon` and `leftover_meals` respectively
...

# Do Not Change the print statement
# Determine result message
if leftover_meals == 0:
    print(f"Mario's meals divide evenly among 6 dragons, with each dragon eating {meals_per_dragon} meals.")
else:
    print(f"Mario's meals do not divide evenly among 6 dragons, leaving {leftover_meals} extra meals.")
```

```python
grader.check("question-operators-mario-dining")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week1-readings", "17_operators_q")
```

Tests for question-operators-mario-dining (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(4.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(4.0)
>>> log_variable('total-points', f'Reading-Week-X, 17_operators_q', 4.0)
>>> question_id = 'question-operators-mario-dining-1'
>>> max_score = 0.5
>>> score = 0
>>> assert total_meals == 21, 'The variable `total_meals` should be 21.'
>>> if total_meals == 21:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('17_operators_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The variable `total_meals` is incorrect.
success_message: Correctly assigned the variable `total_meals`.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-operators-mario-dining-2'
>>> max_score = 0.5
>>> score = 0
>>> assert num_dragons == 6, 'The variable `num_dragons` should be 6.'
>>> if num_dragons == 6:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('17_operators_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The variable `num_dragons` is incorrect.
success_message: Correctly assigned the variable `num_dragons`.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-operators-mario-dining-3'
>>> max_score = 1.0
>>> score = 0
>>> assert meals_per_dragon == 3, 'The variable `meals_per_dragon` should be 3.'
>>> if meals_per_dragon == 3:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('17_operators_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The variable `meals_per_dragon` is incorrect.
success_message: Correctly calculated `meals_per_dragon`.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-operators-mario-dining-4'
>>> max_score = 1.0
>>> score = 0
>>> assert leftover_meals == 3, 'The variable `leftover_meals` should be 3.'
>>> if leftover_meals == 3:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('17_operators_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The variable `leftover_meals` is incorrect.
success_message: Correctly calculated `leftover_meals`.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-operators-mario-dining-5'
>>> max_score = 1.0
>>> score = 0
>>> assert f"Mario's meals do not divide evenly among 6 dragons, leaving {leftover_meals} extra meals." == "Mario's meals do not divide evenly among 6 dragons, leaving 3 extra meals.", 'The print statement is incorrect.'
>>> if f"Mario's meals do not divide evenly among 6 dragons, leaving {leftover_meals} extra meals." == "Mario's meals do not divide evenly among 6 dragons, leaving 3 extra meals.":
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('17_operators_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: print statement is incorrect
success_message: print statement is correct
