```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_kanye_q", "week_5", "practicequiz", assignment_points = 16.0, assignment_tag = 'week5-practicequiz')

# Initialize Otter
import otter
grader = otter.Notebook("1_kanye_q.ipynb")
```

# ❓ Practice Quiz: Kanye Says the Darndest Things! 🤔

<video width="640" height="360" controls>
  <source src="../../_static/bo_burnam_movie.mp4" type="video/mp4">
  Your browser does not support the video tag.
</video>

Kanye is known for making **outrageous, unpredictable** statements. Your job? 

Write functions that **track, analyze, and fact-check** Kanye’s wildest quotes! 🎤🔥

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_kanye_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_kanye_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_kanye_q import Question3
Question3().show()

```

## Question 1 (Points: 7.0): Kanye’s Wildest Quotes Tracker 🎤🔥


 Kanye’s interviews are **full of unpredictable statements**. We need a function to **track** what he says and **detect the wildest moments**.

 The program should:

 1. Accept an arbitrary number of **Kanye quotes** using `*args`.

 2. Detect if any quote exceeds **500 characters**, and print a warning if detected.

 3. Calculate the **total length of all Kanye’s statements** and return it.

 4. Allow an optional `alert_threshold` parameter using `**kwargs` to adjust the warning limit dynamically.


 ### Implementation Details

 1. Define a function `track_kanye(*args, **kwargs)`.

 2. Loop through all `*args` values and:

    - Print a **warning message** if any quote exceeds `500 characters` or the provided `alert_threshold` from `**kwargs`.

 3. Calculate the **total length of all quotes**.

 4. Return the total character count.

```python
# Define the function to monitor Kanye quotes
...

    # Extract the alert threshold from the keyword arguments
    # Recall kwargs is a dictionary, and thus we can use the get method to extract the value of the key "alert_threshold"
    ...
    
    # Initialize an empty list to store warnings
    ...
    
    # Iterate over the quotes and check if any quote exceeds the alert threshold
    # If a quote exceeds the alert threshold, append a warning to the warnings list
    # we have provided a sample warning message for you to use, look at the call to determine the correct variable to use
    ...
            warnings.append(f"⚠️ Warning: Extremely long Kanye quote detected ({len(quote)} characters)!")

    # Calculate the total length of all quotes
    # We can use a generator expression to calculate the total length of all quotes
    # Store the total length in a variable called total_length
    ...
    
    # Return the total length of all quotes and the warnings
    ...

```

```python
grader.check("Kanye-Quote-Analysis")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week5-practicequiz", "1_kanye_q")
```

Tests for Kanye-Quote-Analysis (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(7.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(7.0)
>>> log_variable('total-points', f'Reading-Week-X, 1_kanye_q', 7.0)
>>> question_id = 'Kanye-Quote-Analysis-1'
>>> max_score = 2.0
>>> score = 0
>>> exec(base64.b64decode('dG90YWxfbGVuZ3RoLCBfID0gdHJhY2tfa2FueWUoIkkgYW0gYSBjcmVhdGl2ZSBnZW5pdXMuIiwgIkkgYW0gdGhlIGJlc3QgYXJ0aXN0IG9mIGFsbCB0aW1lLiIp').decode())
>>> exec(base64.b64decode('dG90YWxfbGVuZ3RoXyA9IGxlbigiSSBhbSBhIGNyZWF0aXZlIGdlbml1cy4iKSArIGxlbigiSSBhbSB0aGUgYmVzdCBhcnRpc3Qgb2YgYWxsIHRpbWUuIik=').decode())
>>> condition = total_length_ == total_length
>>> assert condition
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_kanye_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Function does not correctly calculate total character count.
success_message: Success: Function correctly calculates total character count!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Kanye-Quote-Analysis-2'
>>> max_score = 2.0
>>> score = 0
>>> _, warnings = track_kanye('This is a very long quote that keeps going on and on ' * 10)
>>> expected_warning = '⚠️ Warning: Extremely long Kanye quote detected'
>>> assert any((expected_warning in warning for warning in warnings)), 'Expected warning but did not find it.'
>>> if any((expected_warning in warning for warning in warnings)):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_kanye_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Function does not correctly detect long Kanye quotes.
success_message: Success: Function correctly detects long Kanye quotes!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Kanye-Quote-Analysis-3'
>>> max_score = 1.0
>>> score = 0
>>> _, warnings = track_kanye('Short', 'This is an extra long quote' * 20, alert_threshold=1000)
>>> assert not warnings, 'Expected no warnings for a custom alert threshold of 1000.'
>>> if not warnings:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_kanye_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Function does not apply custom alert threshold correctly.
success_message: Success: Function correctly applies custom alert threshold!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Kanye-Quote-Analysis-4'
>>> max_score = 2.0
>>> score = 0
>>> total_length, warnings = track_kanye()
>>> assert total_length == 0, 'Expected total character count to be 0 for empty input.'
>>> assert not warnings, 'Expected no warnings for empty input.'
>>> if total_length == 0 and (not warnings):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_kanye_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Function does not handle empty input correctly.
success_message: Success: Function correctly handles edge cases (no quotes)!
