```python
from pykubegrader.tokens.validate_token import validate_token
validate_token('type the key provided by your instructor here')

```

```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("4_skull_q", "week_2", "lecture", assignment_points = 8.0, assignment_tag = 'week2-lecture')

# Initialize Otter
import otter
grader = otter.Notebook("4_skull_q.ipynb")
```

# ❓ 🦴 Hytrl Skull Collection Problem

Imagine you are cataloging skull specimens in the Hytrl Skull Collection. Each skull has a unique catalog number, type (e.g., carnivore or herbivore), and weight in kilograms.

## Your Task:

1. **Define the Variables**:
   - Create a variable `catalog_number` and assign it a string value `"HSC-101"`.
   - Create a variable `skull_type` and assign it the string `"Carnivore"`.
   - Create a variable `weight_kg` and assign it a float value `4.25`.

2. **Format and Print Information**:
   - Use an f-string to combine these variables into a descriptive sentence.
   - The output should be:  
     `"The skull with catalog number HSC-101 is a Carnivore and weighs 4.25 kg."`

3. **Data Type Check**:
   - Use the `type()` function to verify and print the data type of each variable (`catalog_number`, `skull_type`, and `weight_kg`).
   - For each variable, print a sentence like:  
     `"The data type of catalog_number is <class 'str'>."`

4. **Comments**:
   - For each of the steps add a comment explaining what the code is doing. We are using autograders so the comments have to match an expected schema.
   
   The comments for each step should be: 
   1. `# Define the Variables`
   2. `# Format and Print Information`
   3. `# Data Type Check`


```python
def skull_catalog():
    # Step 1
    ...

    # Step 2
    ...

    # Step 3
    ...


# This will call the function and display the output
# Do not change the lines below
skull_catalog()
```

```python
grader.check("activity-Hytrl Skull Collection")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.tokens.validate_token import validate_token
validate_token()


from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week2-lecture", "4_skull_q")
```

Tests for activity-Hytrl Skull Collection (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> from io import StringIO
>>> import sys
>>> max_question_points = str(8.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(8.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token()
>>> log_variable('total-points', f'Reading-Week-X, 4_skull_q', 8.0)
>>> question_id = 'activity-Hytrl Skull Collection-1'
>>> max_score = 1.0
>>> score = 0
>>> captured_output = StringIO()
>>> sys.stdout = captured_output
>>> skull_catalog()
>>> sys.stdout = sys.__stdout__
>>> output = captured_output.getvalue().strip()
>>> condition = 'The skull with catalog number'.lower() in output.lower()
>>> assert condition, "Output must contain 'The skull with catalog number'."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_skull_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Output does not include 'The skull with catalog number'.
success_message: Correctly included the phrase 'The skull with catalog number'.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> import inspect
>>> question_id = 'activity-Hytrl Skull Collection-2'
>>> max_score = 1.0
>>> score = 0
>>> function_source = inspect.getsource(skull_catalog)
>>> comments = re.findall('#', function_source)
>>> condition = len(comments) >= 8
>>> assert condition, 'Solution must include at least 8 comments explaining the code.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_skull_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Missing required comments in the solution code.
success_message: All required comments are present.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> import inspect
>>> question_id = 'activity-Hytrl Skull Collection-3'
>>> max_score = 1.0
>>> score = 0
>>> source_code = inspect.getsource(skull_catalog)
>>> condition = re.search('#\\s*Define\\s*the\\s*Variables', source_code, re.IGNORECASE)
>>> assert condition, "The '# Define the Variables' comment is missing or does not describe the section properly."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_skull_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The '# Define the Variables' comment is missing or incorrect.
success_message: The '# Define the Variables' comment is present.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> import inspect
>>> question_id = 'activity-Hytrl Skull Collection-4'
>>> max_score = 1.0
>>> score = 0
>>> source_code = inspect.getsource(skull_catalog)
>>> condition = re.search('#\\s*Format\\s*and\\s*Print\\s*Information', source_code, re.IGNORECASE)
>>> assert condition, "The '# Format and Print Information' comment is missing or does not describe the section properly."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_skull_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The '# Format and Print Information' comment is missing or incorrect.
success_message: The '# Format and Print Information' comment is present.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> import inspect
>>> question_id = 'activity-Hytrl Skull Collection-5'
>>> max_score = 1.0
>>> score = 0
>>> source_code = inspect.getsource(skull_catalog)
>>> condition = re.search('#\\s*Data\\s*Type\\s*Check', source_code, re.IGNORECASE)
>>> assert condition, "The '# Data Type Check' comment is missing or does not describe the section properly."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_skull_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The '# Data Type Check' comment is missing or incorrect.
success_message: The '# Data Type Check' comment is present.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> from io import StringIO
>>> import sys
>>> question_id = 'activity-Hytrl Skull Collection-6'
>>> max_score = 1.0
>>> score = 0
>>> captured_output = StringIO()
>>> sys.stdout = captured_output
>>> skull_catalog()
>>> sys.stdout = sys.__stdout__
>>> output = captured_output.getvalue().strip()
>>> condition = "<class 'str'>" in output
>>> assert condition, "The datatype of `catalog_number` should be <class 'str'>."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_skull_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The datatype of `catalog_number` is incorrect or not verified properly.
success_message: Correctly verified the datatype of `catalog_number`.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> from io import StringIO
>>> import sys
>>> question_id = 'activity-Hytrl Skull Collection-7'
>>> max_score = 1.0
>>> score = 0
>>> captured_output = StringIO()
>>> sys.stdout = captured_output
>>> skull_catalog()
>>> sys.stdout = sys.__stdout__
>>> output = captured_output.getvalue().strip()
>>> condition = "<class 'str'>" in output
>>> assert condition, "The datatype of `skull_type` should be <class 'str'>."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_skull_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The datatype of `skull_type` is incorrect or not verified properly.
success_message: Correctly verified the datatype of `skull_type`.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> from io import StringIO
>>> import sys
>>> question_id = 'activity-Hytrl Skull Collection-8'
>>> max_score = 1.0
>>> score = 0
>>> captured_output = StringIO()
>>> sys.stdout = captured_output
>>> skull_catalog()
>>> sys.stdout = sys.__stdout__
>>> output = captured_output.getvalue().strip()
>>> condition = "<class 'float'>" in output
>>> assert condition, "The datatype of `weight_kg` should be <class 'float'>."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('4_skull_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The datatype of `weight_kg` is incorrect or not verified properly.
success_message: Correctly verified the datatype of `weight_kg`.
