```python
from pykubegrader.tokens.validate_token import validate_token
validate_token('type the key provided by your instructor here')

```

```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("9_dictionaries_q", "week_2", "lecture", assignment_points = 7.0, assignment_tag = 'week2-lecture')

# Initialize Otter
import otter
grader = otter.Notebook("9_dictionaries_q.ipynb")
```

# â“ ðŸ©º Chevalier Jackson Collection Problem

The Chevalier Jackson Collection contains thousands of objects extracted from patients' airways and esophagi. These include items such as safety pins, coins, buttons, needles, toy parts, dentures, animal bones, keys, and poker chips.

## Your Task:

1. **Define an Object Dictionary**:
   - Create a dictionary called `object_info` to represent a single item in the collection.
   - The dictionary should include the following keys:
     - `catalog_number`: A string, e.g., "CJC-001".
     - `description`: A string, e.g., "Open safety pin".
     - `type`: A string, e.g., "Foreign Body".
     - `year_removed`: An integer, e.g., 1923.

   Example:
   ```python
   object_info = {
       "catalog_number": "CJC-010",
       "description": "poker chip",
       "type": "Foreign Body",
       "year_removed": 2002
   }
   ```

2. **Access and Print Information**:
   - Use dictionary key access to print the following sentence using the data in `object_info`:
     - "The item CJC-010 is an poker chip classified as a Foreign Body, removed in 2002."
     - The statement should be saved to the variable `print_statement`, we have provided the code to print the statement for you.

3. **Modify the Dictionary**:
   - Update the `year_removed` to `1925`.
   - Print the updated dictionary.

4. **Add a New Key**:
   - Add a new key called `condition` with the value "Good" to the dictionary.
   - Print the updated dictionary.

```python

...

# Access and Print Information
print_statement = ...

# We have provided the code to print the statement
print(print_statement)

# Modify the Dictionary
...
# We have provided the code to print the updated dictionary
print("Updated Dictionary:", object_info)

# Add a New Key
...
# We have provided the code to print the updated dictionary
print("Updated Dictionary with Condition:", object_info)

```

```python
grader.check("chevalier-jackson-dictionary")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.tokens.validate_token import validate_token
validate_token()


from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week2-lecture", "9_dictionaries_q")
```

Tests for chevalier-jackson-dictionary (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(7.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(7.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token()
>>> log_variable('total-points', f'Reading-Week-X, 9_dictionaries_q', 7.0)
>>> question_id = 'chevalier-jackson-dictionary-1'
>>> max_score = 1.0
>>> score = 0
>>> assert 'object_info' in globals(), 'The object_info dictionary must be defined.'
>>> assert isinstance(object_info, dict), 'object_info must be a dictionary.'
>>> if 'object_info' in globals() and isinstance(object_info, dict):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('9_dictionaries_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The object_info dictionary must be defined and a dictionary.
success_message: The object_info dictionary is defined.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'chevalier-jackson-dictionary-2'
>>> max_score = 1.0
>>> score = 0
>>> assert object_info['catalog_number'] == 'CJC-001', 'The catalog_number value is incorrect.'
>>> if object_info['catalog_number'] == 'CJC-001':
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('9_dictionaries_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The catalog_number value is incorrect.
success_message: The catalog_number value is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'chevalier-jackson-dictionary-3'
>>> max_score = 1.0
>>> score = 0
>>> condition = object_info['description'] == 'Open safety pin'
>>> assert condition, 'The description value is incorrect.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('9_dictionaries_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The description value is incorrect.
success_message: The description value is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'chevalier-jackson-dictionary-4'
>>> max_score = 1.0
>>> score = 0
>>> assert object_info['type'] == 'Foreign Body', 'The type value is incorrect.'
>>> if object_info['type'] == 'Foreign Body':
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('9_dictionaries_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The type value is incorrect.
success_message: The type value is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> question_id = 'chevalier-jackson-dictionary-5'
>>> max_score = 1.0
>>> score = 0
>>> condition = re.search('.*CJC-001.*Open safety pin.*Foreign Body.*1923', print_statement, re.IGNORECASE)
>>> assert condition, 'The output for object_info access is incorrect.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('9_dictionaries_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The printed output for object_info access is incorrect.
success_message: Successfully accessed and printed information from the dictionary.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'chevalier-jackson-dictionary-6'
>>> max_score = 1.0
>>> score = 0
>>> condition = object_info['year_removed'] == 1925
>>> assert condition, 'The year_removed value is not updated to 1925.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('9_dictionaries_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The year_removed value was not updated correctly.
success_message: Successfully updated the year_removed value.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'chevalier-jackson-dictionary-7'
>>> max_score = 1.0
>>> score = 0
>>> assert 'condition' in object_info, 'The condition key was not added.'
>>> assert object_info['condition'] == 'Good', 'The condition value is incorrect.'
>>> if 'condition' in object_info and object_info['condition'] == 'Good':
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('9_dictionaries_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The condition key was not added correctly.
success_message: Successfully added the condition key.
