```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("5_seaborn_q", "week_8", "readings", assignment_points = 20.0, assignment_tag = 'week8-readings')

# Initialize Otter
import otter
grader = otter.Notebook("5_seaborn_q.ipynb")
```

# â“ Seaborn Graphing Made Simple

## Question 1 (Points: 8.0): Understanding Diamond Prices: A Data-Driven Exploration  

![](./assets/figures/agar+ring.jpg)

Diamonds are often considered a symbol of luxury, and their prices are influenced by multiple factors such as cut, clarity, carat, and color. But how do these factors correlate with each other, and what insights can we gain from visualizing their relationships?  

In this analysis, we will explore the relationship between diamond cut quality and price using a **violin plot**, which helps us understand the distribution of diamond prices across different cut grades. Additionally, we will generate a **correlation heatmap** to identify which numerical features of diamonds (such as carat weight, depth, and table size) are most strongly related to price.  

By analyzing these visualizations, we can answer key questions such as:  
- Does a higher-quality cut always result in a higher price?  
- Which diamond characteristics have the strongest correlation with price?  
- Are there any surprising patterns or trends in the data?  

Let's dive into the data and uncover the hidden patterns behind diamond pricing! ðŸ’ŽðŸ“Š. 

We will provide you with inline instructions to help you complete the analysis.

```python
# Import necessary libraries
# seaborn as sns
# matplotlib.pyplot as plt
# pandas as pd
...

# Load dataset
# Load the diamonds dataset from seaborn, to do this you can use the `sns.load_dataset` function.
# The diamonds dataset can be accessed by the name `diamonds`. See documentation [here](https://seaborn.pydata.org/generated/seaborn.load_dataset.html) for more information.
...

# Create figure with subplots
# Create a figure with two subplots, one for the violin plot and one for the correlation heatmap. The subplots should be in a 1x2 grid.
# Save the figure to the variable `fig`, and the axes to the variable `axes`.
# The figure should be of size 16x6.
...

# Violin plot with color palette
# Use the `sns.violinplot` function to create the violin plot.
# You will need to access the `cut` and `price` columns of the dataframe. This is done by calling the pandas dataframe like a dictionary, e.g. `df["cut"]` and `df["price"]`.
# The x-axis should be the `cut` column of the dataframe, and the y-axis should be the `price` column.
# The color palette should be `pastel`.
# The plot should be saved to the first subplot (`axes[0]`).
...

# Set the title, x-axis label, and y-axis label of the first subplot.
# Note the axes object is a list of two axes objects, the first one is `axes[0]` and the second one is `axes[1]`.
# The title should be "Distribution of Diamond Prices by Cut Quality".
# The x-axis label should be "Cut Quality".
# The y-axis label should be "Price (USD)".
...

# Correlation heatmap
# Create a correlation matrix of the dataframe.
# The correlation matrix should only include the numerical columns of the dataframe. You need to use the `corr` method of the dataframe. Make sure to set `numeric_only=True`.
# You can access the numerical columns of the dataframe by calling `df.select_dtypes(include=["number"])`.
# Save the correlation matrix to the variable `corr_matrix`.
...
# Use the `sns.heatmap` function to create the heatmap.
# The heatmap should be saved to the second subplot (`axes[1]`).
# The color palette should be `coolwarm`.
# The annotation should be the correlation matrix.
# The annotation should be formatted to 2 decimal places.
# look at the documentation for the `sns.heatmap` function for specific guidance.
...
# Set the title, x-axis label, and y-axis label of the second subplot.
# The title should be "Correlation Matrix of Diamond Features".
# The x-axis label should be "Features".
# The y-axis label should be "Features".
...

# Show plots
# Use the `plt.tight_layout` function to adjust the layout of the plots.
# Use the `plt.show` function to display the plots.
...
```

```python
grader.check("Alternative Distribution Plot with Diamonds Dataset")
```

## Question 2 (Points: 12.0): **Create and Analyze a Dendrogram for Wine Quality Data**

The **Wine dataset** contains **chemical properties of different wines** from three cultivars (wine types). Your goal is to:

1. Load the **UCI Wine dataset** from `sklearn.datasets`.
2. Use **only the numerical features** (excluding the wine class label).
3. Perform **hierarchical clustering** and generate a **dendrogram**.
4. Customize the dendrogram for better readability.
5. **Analyze the results**:
   - How many clusters are formed at a **cutoff height of 200**?
   - Based on clustering, which **two wine classes appear most similar**?

**Expected Output:**
- A **dendrogram** visualizing hierarchical clustering of the wine dataset.
- Answers to the analysis questions.

```python
# Import necessary libraries
# sklearn.datasets as load_wine
# sklearn.preprocessing as MinMaxScaler
# pandas.plotting as parallel_coordinates
...
# Load the wine dataset from sklearn, to do this you can use the `load_wine` function.
...
# Create a DataFrame from the wine dataset 
# The DataFrame should contain the feature names as columns and the target as the class column.
...
# Convert the target class to string for visualization
...
# Normalize numerical data using MinMaxScaler
...
# Normalize the numerical columns of the dataframe using the MinMaxScaler.
# The numerical columns are all columns except the `class` column.
...

# Add the target class for coloring the parallel coordinates plot 
...
# Plot parallel coordinates plot of the normalized data 
...
# Use the `parallel_coordinates` function from pandas to create the parallel coordinates plot.
# The dataframe should be the normalized dataframe.
# The class column should be the target class.
# The colormap should be `coolwarm`.
# The transparency should be 0.7.

...
# Set the title, x-axis label, and y-axis label of the plot.
# The title should be "Parallel Coordinates Plot of Wine Dataset".
# The x-axis label should be "Features".
# The y-axis label should be "Normalized Values".
...
```

```python
grader.check("Creating and Analyzing a Dendrogram with Wine Dataset")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week8-readings", "5_seaborn_q")
```

Tests for Alternative Distribution Plot with Diamonds Dataset (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import sys
>>> max_question_points = str(8.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(20.0)
>>> log_variable('total-points', f'week8-readings, 5_seaborn_q', 20.0)
>>> question_id = 'Alternative Distribution Plot with Diamonds Dataset-1'
>>> max_score = 1.0
>>> score = 0
>>> assert 'seaborn' in sys.modules, 'Seaborn is not imported.'
>>> assert 'matplotlib.pyplot' in sys.modules, 'Matplotlib is not imported.'
>>> assert 'pandas' in sys.modules, 'Pandas is not imported.'
>>> if 'seaborn' in sys.modules and 'matplotlib.pyplot' in sys.modules and ('pandas' in sys.modules):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_seaborn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Libraries are not imported correctly.
success_message: Success: Required libraries are imported correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Alternative Distribution Plot with Diamonds Dataset-2'
>>> max_score = 1.0
>>> score = 0
>>> assert 'df' in locals(), "DataFrame 'df' is not defined."
>>> assert isinstance(df, pd.DataFrame), 'df is not a pandas DataFrame.'
>>> assert not df.empty, 'df is empty.'
>>> if 'df' in locals() and isinstance(df, pd.DataFrame) and (not df.empty):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_seaborn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(df))
```
failure_message: Failed: Diamonds dataset is not loaded correctly.
success_message: Success: Diamonds dataset loaded correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> question_id = 'Alternative Distribution Plot with Diamonds Dataset-3'
>>> max_score = 1.0
>>> score = 0
>>> assert 'fig' in locals() and 'axes' in locals(), 'Figure or axes are not defined.'
>>> assert isinstance(fig, plt.Figure), 'fig is not a Matplotlib Figure.'
>>> assert isinstance(axes, (list, tuple, np.ndarray)) and len(axes) == 2, 'axes is not a list of two elements.'
>>> if ('fig' in locals() and 'axes' in locals()) and isinstance(fig, plt.Figure) and (isinstance(axes, (list, tuple, np.ndarray)) and len(axes) == 2):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_seaborn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(fig))
>>> responses = update_responses(question_id, str(axes))
```
failure_message: Failed: Figure and subplots are not created correctly.
success_message: Success: Figure and subplots created correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Alternative Distribution Plot with Diamonds Dataset-4'
>>> max_score = 1.0
>>> score = 0
>>> assert any((isinstance(child, plt.Axes) for child in fig.axes)), 'Violin plot is not created on the first subplot.'
>>> if any((isinstance(child, plt.Axes) for child in fig.axes)):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_seaborn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Violin plot is not created correctly.
success_message: Success: Violin plot created correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Alternative Distribution Plot with Diamonds Dataset-5'
>>> max_score = 1.0
>>> score = 0
>>> assert axes[0].get_title() == 'Distribution of Diamond Prices by Cut Quality', 'Violin plot title is incorrect.'
>>> assert axes[0].get_xlabel() == 'Cut Quality', 'Violin plot x-axis label is incorrect.'
>>> assert axes[0].get_ylabel() == 'Price (USD)', 'Violin plot y-axis label is incorrect.'
>>> if axes[0].get_title() == 'Distribution of Diamond Prices by Cut Quality' and axes[0].get_xlabel() == 'Cut Quality' and (axes[0].get_ylabel() == 'Price (USD)'):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_seaborn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Titles or labels for violin plot are incorrect.
success_message: Success: Titles and labels for violin plot set correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Alternative Distribution Plot with Diamonds Dataset-6'
>>> max_score = 1.0
>>> score = 0
>>> assert 'corr_matrix' in locals(), 'Correlation matrix is not defined.'
>>> assert isinstance(corr_matrix, pd.DataFrame), 'corr_matrix is not a DataFrame.'
>>> assert not corr_matrix.empty, 'Correlation matrix is empty.'
>>> if 'corr_matrix' in locals() and isinstance(corr_matrix, pd.DataFrame) and (not corr_matrix.empty):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_seaborn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(corr_matrix))
```
failure_message: Failed: Correlation matrix is incorrect.
success_message: Success: Correlation matrix created correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Alternative Distribution Plot with Diamonds Dataset-7'
>>> max_score = 1.0
>>> score = 0
>>> assert any((isinstance(child, plt.Axes) for child in fig.axes)), 'Heatmap is not created on the second subplot.'
>>> if any((isinstance(child, plt.Axes) for child in fig.axes)):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_seaborn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Correlation heatmap is not created correctly.
success_message: Success: Correlation heatmap created correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Alternative Distribution Plot with Diamonds Dataset-8'
>>> max_score = 1.0
>>> score = 0
>>> assert axes[1].get_title() == 'Correlation Matrix of Diamond Features', 'Heatmap title is incorrect.'
>>> assert axes[1].get_xlabel() == 'Features', 'Heatmap x-axis label is incorrect.'
>>> assert axes[1].get_ylabel() == 'Features', 'Heatmap y-axis label is incorrect.'
>>> if axes[1].get_title() == 'Correlation Matrix of Diamond Features' and axes[1].get_xlabel() == 'Features' and (axes[1].get_ylabel() == 'Features'):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_seaborn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Titles or labels for heatmap are incorrect.
success_message: Success: Titles and labels for heatmap set correctly!


Tests for Creating and Analyzing a Dendrogram with Wine Dataset (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import sys
>>> max_question_points = str(12.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(20.0)
>>> log_variable('total-points', f'week8-readings, 5_seaborn_q', 20.0)
>>> question_id = 'Creating and Analyzing a Dendrogram with Wine Dataset-1'
>>> max_score = 1.0
>>> score = 0
>>> assert 'sklearn.datasets' in sys.modules
>>> assert 'sklearn.preprocessing' in sys.modules
>>> assert 'pandas.plotting' in sys.modules
>>> if 'sklearn.datasets' in sys.modules and 'sklearn.preprocessing' in sys.modules and ('pandas.plotting' in sys.modules):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_seaborn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Some required libraries are missing.
success_message: Libraries imported correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> question_id = 'Creating and Analyzing a Dendrogram with Wine Dataset-2'
>>> max_score = 2.0
>>> score = 0
>>> assert 'data' in wine and 'target' in wine
>>> assert isinstance(wine.data, (list, tuple, np.ndarray))
>>> if ('data' in wine and 'target' in wine) and isinstance(wine.data, (list, tuple, np.ndarray)):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_seaborn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Wine dataset is not loaded correctly.
success_message: Wine dataset loaded correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Creating and Analyzing a Dendrogram with Wine Dataset-3'
>>> max_score = 2.0
>>> score = 0
>>> assert isinstance(df, pd.DataFrame)
>>> assert set(wine.feature_names).issubset(df.columns)
>>> assert 'class' in df.columns
>>> if isinstance(df, pd.DataFrame) and set(wine.feature_names).issubset(df.columns) and ('class' in df.columns):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_seaborn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Wine dataset DataFrame conversion is incorrect.
success_message: Wine dataset converted into a DataFrame correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Creating and Analyzing a Dendrogram with Wine Dataset-4'
>>> max_score = 2.0
>>> score = 0
>>> assert df['class'].dtype == object
>>> if df['class'].dtype == object:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_seaborn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Target column is not in string format.
success_message: Target column converted to string format.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Creating and Analyzing a Dendrogram with Wine Dataset-5'
>>> max_score = 2.0
>>> score = 0
>>> assert isinstance(scaler, MinMaxScaler)
>>> if isinstance(scaler, MinMaxScaler):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_seaborn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: MinMaxScaler is not initialized correctly.
success_message: MinMaxScaler initialized correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> question_id = 'Creating and Analyzing a Dendrogram with Wine Dataset-6'
>>> max_score = 3.0
>>> score = 0
>>> assert df_normalized.shape == df.shape
>>> assert 'class' in df_normalized.columns
>>> a_0 = df_normalized.drop(columns=['class']) >= 0
>>> a_1 = df_normalized.drop(columns=['class']) <= 1
>>> assert np.all(a_0) and np.all(a_1)
>>> if df_normalized.shape == df.shape and 'class' in df_normalized.columns and (np.all(a_0) and np.all(a_1)):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('5_seaborn_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Numerical data normalization is incorrect.
success_message: Numerical data correctly normalized.
