```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("2_refining_plot_q", "week_8", "lecturenotgraded", assignment_points = 44.0, assignment_tag = 'week8-lecturenotgraded')

# Initialize Otter
import otter
grader = otter.Notebook("2_refining_plot_q.ipynb")
```

# 💻 Activity: Refining a Simple Plot

In many labs and classes, you'll need to meaningfully visualize data for your audience.

Successful communication of data is one of the most important aspects of engineering. If you can't get important ideas out of your head in a way that makes them digestible to others then you are unlikely to succeed as an engineer. 

Please follow along on your computers as we refine a visualization of sine and cosine.

## Create data

The first step is to get data for the sine and cosine functions.

Note that we import `pyplot` from `matplotlib` to use `pyplot` to visualize 2-D data.

```python
import numpy as np
import matplotlib.pyplot as plt

# creates a linear spaced array
# create a linear spaced array of 256 values ranging from -π to +π (included), using the np.linspace() function, set the endpoint to True
# assign the result to the variable X
...

# computes the sine and cosine
# compute the cosine and sine of the array X, and assign the results to the variables C and S
...

```

```python
grader.check("q1-create-data")
```

`X` is now a NumPy array with 256 values ranging from -π to +π (included). `C` is the cosine (256 values), and `S` is the sine (256 values).

```python
# plot the cosine against X
# use the plt.plot() function to plot the cosine against X
# save as the variable plt_cosine
...

# plot the sine against X
# use the plt.plot() function to plot the sine against X
# save as the variable plt_sine
...

# display the plot using the plt.show() function
...

```

```python
grader.check("q2-plot-data")
```

## Question 2 (Points: 2.0): Customizing Plots

Matplotlib was designed to make 2-dimensional plotting easy! We will generate simple graphics today, but you can use these links to learn about more powerful applications of matplotlib:

[Customizing matplotlib ](http://matplotlib.sourceforge.net/users/customizing.html)

[Controlling line properties](https://matplotlib.org/stable/tutorials/introductory/pyplot.html#controlling-line-properties)


Let's adjust color so cosine is blue and the sine is red. Let's also increase line width to 2.5.

We want the graphic to display well on a laptop screen, so let's also adjust the aspect ratio (width:height) using `figsize`.

```python
# create a figure with a size of 10x6 and a dpi of 300
# save as the variable plt_figure
# BEGIN SOLUTION
plt_figure = plt.figure(figsize=(10, 6), dpi=300)
# END SOLUTION

# plot the cosine against X
# use the plt.plot() function to plot the cosine against X
# set the color to blue, the linewidth to 2.5, and the linestyle to "-", this will make a solid blue line
# save as the variable plt_cosine
# BEGIN SOLUTION
plt_cosine = plt.plot(X, C, color="blue", linewidth=2.5, linestyle="-")
# END SOLUTION

# plot the sine against X
# use the plt.plot() function to plot the sine against X
# set the color to red, the linewidth to 2.5, and the linestyle to "-", this will make a solid red line
# save as the variable plt_sine
# BEGIN SOLUTION
plt_sine = plt.plot(X, S, color="red", linewidth=2.5, linestyle="-")
# END SOLUTION
```

```python
grader.check("q3-create-figure")
```

## Question 3 (Points: 2.0): Setting Limits

[xlim() command ](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.xlim.html)

[ylim() command](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.ylim.html)

The current limits of the figure are a bit too tight. To better visualize the functions we want to extend the range of the x- and y-axes.

Extend the ranges to be 10% more negative at the minimum and 10% greater at the maximum.

```python
# create a new figure with a size of 8x5 and a dpi of 300
# save as the variable plt_figure
...

# create a subplot with 1 row and 1 column, and the first plot, this is done by plt.subplot(111)
...

# create a linear spaced array of 256 values ranging from -π to +π (included), using the np.linspace() function, set the endpoint to True
# assign the result to the variable X
...

# compute the cosine and sine of the array X, and assign the results to the variables C and S
...

# copy these lines from above
# use the plt.plot() function to plot the cosine against X
# set the color to blue, the linewidth to 2.5, and the linestyle to "-", this will make a solid blue line
...

# use the plt.plot() function to plot the sine against X
# set the color to red, the linewidth to 2.5, and the linestyle to "-", this will make a solid red line
...

# set the limits of the x-axis and y-axis
# set the range of the x-axis to be 10% more negative at the minimum and 10% greater at the maximum, hint use X.min() and X.max()
# set the range of the y-axis to be 10% more negative at the minimum and 10% greater at the maximum, hint use C.min() and C.max()
# save as the variable plt_xlim and plt_ylim
...

# display the plot using the plt.show() function
...
```

```python
grader.check("q4-set-limits")
```

## Question 4 (Points: 15.0): Setting Ticks

[xticks() command](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.xticks)

[yticks() command](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.yticks)

The current ticks are not ideal because they do not show the interesting values for sine and cosine, like ± π and ± (π/2). 

Use the `xticks` function to make ticks at -π, -π/2, 0, π/2, and π.

Use the `yticks` function to make ticks at -1, 0, and 1.

### Setting the Tick Labels

[Working with text](https://matplotlib.org/stable/tutorials/text/text_intro.html)

[set_xticklabels()](https://matplotlib.org/stable/api/_as_gen/matplotlib.axes.Axes.set_xticklabels)

[set_yticklabels()](https://matplotlib.org/stable/api/_as_gen/matplotlib.axes.Axes.set_yticklabels)


The ticks are now properly placed, but their label is not very explicit. 

We could guess that 3.142 is π, but it would be better to make it explicit. 

When we set tick values, we can also provide a corresponding label in the second argument list.

Note that we'll use [latex](https://www.overleaf.com/learn/latex/Learn_LaTeX_in_30_minutes) to allow for a nice rendering of the label.


```python
# Copy the cell above into the area with the ..., this will create another version of the same plot

...

# use the plt.xticks() function to make ticks at -π, -π/2, 0, π/2, and π, this is done by providing a list of values to the function
# use the plt.yticks() function to make ticks at -1, 0, and 1, this is done by providing a list of values to the function
...

# display the plot using the plt.show() function
...

```

```python
grader.check("q5-set-ticks")
```

#### Moving Spines

[Spines](https://matplotlib.org/stable/gallery/spines/spines.html)

Spines are the lines connecting the axis tick marks and noting the boundaries of the data area.

They can be placed at arbitrary positions. Until now, they were on the border of the axis.

We'll change that since we want to have them in the middle.

1. Since there are four of them (top/bottom/left/right), we'll discard the top and right by setting their color to none.
2. We'll move the bottom and left ones to the origin of the graph.


```python
# Copy the cell above into the area with the ..., this will create another version of the same plot

...

# set the right and top spines to none
...

# set the bottom spine to the origin where the data is 0
# move the bottom spline to the origin
...
```

```python
grader.check("q6-move-spines")
```

## Question 6 (Points: 6.0): Adding a Legend

[Legend guide](https://matplotlib.org/stable/tutorials/intermediate/legend_guide.html)

[legend() command](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.legend)


Legends help someone viewing a plot to understand what it shows. 

Let's add a legend in the upper left corner.

Insert the keyword argument `label` to the plot commands, like `label="cosine"`

```python
# Copy the cell above into the area with the ..., this will create another version of the same plot
...

# call the plt.legend() function to add a legend in the upper left corner, turn the frame off, by setting frameon to False
...

# display the plot using the plt.show() function
plt.show()

```

```python
grader.check("q7-add-legend")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week8-lecturenotgraded", "2_refining_plot_q")
```

Tests for q1-create-data (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(11.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(44.0)
>>> log_variable('total-points', f'week8-lecturenotgraded, 2_refining_plot_q', 44.0)
>>> question_id = 'q1-create-data-1'
>>> max_score = 2.0
>>> score = 0
>>> array_length = len(X)
>>> array_start = X[0]
>>> array_end = X[-1]
>>> assert array_length == 256, f'Expected array length 256, but got {array_length}.'
>>> assert np.isclose(array_start, -np.pi), f'Expected start value -π, but got {array_start}.'
>>> assert np.isclose(array_end, np.pi), f'Expected end value +π, but got {array_end}.'
>>> if (array_length == 256, f'Expected array length 256, but got {array_length}.') and (np.isclose(array_start, -np.pi), f'Expected start value -π, but got {array_start}.') and (np.isclose(array_end, np.pi), f'Expected end value +π, but got {array_end}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(array_length))
>>> responses = update_responses(question_id, str(array_start))
>>> responses = update_responses(question_id, str(array_end))
```
failure_message: Failed: Linear spaced array is not created as expected.
success_message: Success: Linear spaced array is created correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q1-create-data-2'
>>> max_score = 2.0
>>> score = 0
>>> cos_start = C[0]
>>> cos_end = C[-1]
>>> cos_mid = C[len(C) // 2]
>>> assert np.isclose(cos_start, -1.0), f'Expected cos(-π) = -1.0, but got {cos_start}.'
>>> assert np.isclose(cos_end, -1.0, atol=0.0001), f'Expected cos(+π) = -1.0, but got {cos_end}.'
>>> assert np.isclose(cos_mid, 1.0, atol=0.0001), f'Expected cos(0) = 1.0, but got {cos_mid}.'
>>> if (np.isclose(cos_start, -1.0), f'Expected cos(-π) = -1.0, but got {cos_start}.') and (np.isclose(cos_end, -1.0, atol=0.0001), f'Expected cos(+π) = -1.0, but got {cos_end}.') and (np.isclose(cos_mid, 1.0, atol=0.0001), f'Expected cos(0) = 1.0, but got {cos_mid}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(cos_start))
>>> responses = update_responses(question_id, str(cos_end))
>>> responses = update_responses(question_id, str(cos_mid))
```
failure_message: Failed: Cosine values are not computed as expected.
success_message: Success: Cosine values are computed correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q1-create-data-3'
>>> max_score = 2.0
>>> score = 0
>>> sin_start = S[0]
>>> sin_end = S[-1]
>>> sin_mid = S[len(S) // 2]
>>> assert np.isclose(sin_start, 0.01, atol=0.01), f'Expected sin(-π) = 0.0, but got {sin_start}.'
>>> assert np.isclose(sin_end, 0.01, atol=0.01), f'Expected sin(+π) = 0.0, but got {sin_end}.'
>>> assert np.isclose(sin_mid, 0.01, atol=0.01), f'Expected sin(0) = 0.0, but got {sin_mid}.'
>>> if (np.isclose(sin_start, 0.01, atol=0.01), f'Expected sin(-π) = 0.0, but got {sin_start}.') and (np.isclose(sin_end, 0.01, atol=0.01), f'Expected sin(+π) = 0.0, but got {sin_end}.') and (np.isclose(sin_mid, 0.01, atol=0.01), f'Expected sin(0) = 0.0, but got {sin_mid}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(sin_start))
>>> responses = update_responses(question_id, str(sin_end))
>>> responses = update_responses(question_id, str(sin_mid))
```
failure_message: Failed: Sine values are not computed as expected.
success_message: Success: Sine values are computed correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q1-create-data-4'
>>> max_score = 1.0
>>> score = 0
>>> X_shape, C_shape, S_shape = (X.shape, C.shape, S.shape)
>>> X_dtype, C_dtype, S_dtype = (X.dtype, C.dtype, S.dtype)
>>> assert X_shape == (256,), f'Expected shape (256,) for X, but got {X_shape}.'
>>> assert C_shape == (256,), f'Expected shape (256,) for C, but got {C_shape}.'
>>> assert S_shape == (256,), f'Expected shape (256,) for S, but got {S_shape}.'
>>> assert X_dtype == np.float64, f'Expected dtype float64 for X, but got {X_dtype}.'
>>> assert C_dtype == np.float64, f'Expected dtype float64 for C, but got {C_dtype}.'
>>> assert S_dtype == np.float64, f'Expected dtype float64 for S, but got {S_dtype}.'
>>> if X_shape == (256,) and C_shape == (256,) and (S_shape == (256,)) and (X_dtype == np.float64) and (C_dtype == np.float64) and (S_dtype == np.float64):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(X_shape))
>>> responses = update_responses(question_id, str(C_shape))
>>> responses = update_responses(question_id, str(S_shape))
>>> responses = update_responses(question_id, str(X_dtype))
>>> responses = update_responses(question_id, str(C_dtype))
>>> responses = update_responses(question_id, str(S_dtype))
```
failure_message: Failed: Arrays do not have correct shape or type.
success_message: Success: Arrays have correct shape and type!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q1-create-data-5'
>>> max_score = 2.0
>>> score = 0
>>> cos_periodicity = np.allclose(C[0], C[-1], rtol=0.001)
>>> sin_periodicity = np.allclose(S[0], S[-1], rtol=0.001)
>>> assert cos_periodicity, f'Expected periodicity in Cosine: C[0] ≈ C[-1], but got C[0] = {C[0]} and C[-1] = {C[-1]}.'
>>> assert sin_periodicity, f'Expected periodicity in Sine: S[0] ≈ S[-1], but got S[0] = {S[0]} and S[-1] = {S[-1]}.'
>>> if (cos_periodicity, f'Expected periodicity in Cosine: C[0] ≈ C[-1], but got C[0] = {C[0]} and C[-1] = {C[-1]}.') and (sin_periodicity, f'Expected periodicity in Sine: S[0] ≈ S[-1], but got S[0] = {S[0]} and S[-1] = {S[-1]}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(cos_periodicity))
>>> responses = update_responses(question_id, str(sin_periodicity))
```
failure_message: Failed: Cosine and Sine periodicity is incorrect.
success_message: Success: Cosine and Sine periodicity is correct!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q1-create-data-6'
>>> max_score = 2.0
>>> score = 0
>>> cos_even_symmetry = np.allclose(C, np.flip(C), rtol=0.001)
>>> sin_odd_symmetry = np.allclose(S, -np.flip(S), rtol=0.001)
>>> assert cos_even_symmetry, f'Expected Cosine to be even (C(x) = C(-x)), but symmetry check failed.'
>>> assert sin_odd_symmetry, f'Expected Sine to be odd (S(x) = -S(-x)), but symmetry check failed.'
>>> if (cos_even_symmetry, f'Expected Cosine to be even (C(x) = C(-x)), but symmetry check failed.') and (sin_odd_symmetry, f'Expected Sine to be odd (S(x) = -S(-x)), but symmetry check failed.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(cos_even_symmetry))
>>> responses = update_responses(question_id, str(sin_odd_symmetry))
```
failure_message: Failed: Symmetry in Cosine and Sine is incorrect.
success_message: Success: Symmetry in Cosine and Sine is correct!


Tests for q2-plot-data (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(2.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(44.0)
>>> log_variable('total-points', f'week8-lecturenotgraded, 2_refining_plot_q', 44.0)
>>> question_id = 'q2-plot-data-1'
>>> max_score = 2.0
>>> score = 0
>>> assert isinstance(plt_cosine, list) and isinstance(plt_sine, list), 'Plots are not created.'
>>> condition = plt_cosine[0].get_xdata().tolist() == X.tolist() and plt_cosine[0].get_ydata().tolist() == C.tolist()
>>> assert condition
>>> if (isinstance(plt_cosine, list) and isinstance(plt_sine, list)) and condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(plt_cosine))
>>> responses = update_responses(question_id, str(plt_sine))
```
failure_message: Failed: Plots are not created or displayed correctly.
success_message: Success: Plots are created and displayed correctly!


Tests for q3-create-figure (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import matplotlib.pyplot as plt
>>> max_question_points = str(2.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(44.0)
>>> log_variable('total-points', f'week8-lecturenotgraded, 2_refining_plot_q', 44.0)
>>> question_id = 'q3-create-figure-1'
>>> max_score = 2.0
>>> score = 0
>>> condition = plt_figure is not None and isinstance(plt_figure, plt.Figure) and (plt_figure.get_size_inches().tolist() == [10, 6]) and (plt_figure.get_dpi() == 300)
>>> assert condition, 'Figure is not created with the correct size (10x6) or dpi (300).'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(plt_figure))
```
failure_message: Failed: The figure is not created with the correct size and dpi.
success_message: Success: The figure is created with the correct size and dpi!


Tests for q4-set-limits (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import matplotlib.pyplot as plt
>>> max_question_points = str(15.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(44.0)
>>> log_variable('total-points', f'week8-lecturenotgraded, 2_refining_plot_q', 44.0)
>>> question_id = 'q4-set-limits-1'
>>> max_score = 2.0
>>> score = 0
>>> condition = plt_figure is not None and isinstance(plt_figure, plt.Figure) and (plt_figure.get_size_inches().tolist() == [8, 5]) and (plt_figure.dpi == 300)
>>> assert condition, 'Figure is not created with the correct size or dpi. Ensure plt.figure(figsize=(8, 5), dpi=300) is used.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(plt_figure))
```
failure_message: Failed: The figure is not created correctly with size 8x5 and dpi 300.
success_message: Success: The figure is created correctly with size 8x5 and dpi 300!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q4-set-limits-2'
>>> max_score = 1.0
>>> score = 0
>>> condition = ax is not None and isinstance(ax, plt.Axes)
>>> assert condition, 'Subplot is not created correctly. Ensure plt.subplot(111) is used and assigned to ax.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(ax))
```
failure_message: Failed: Subplot is not created correctly using plt.subplot(111).
success_message: Success: Subplot is created correctly using plt.subplot(111)!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> question_id = 'q4-set-limits-3'
>>> max_score = 2.0
>>> score = 0
>>> expected_X = np.linspace(-np.pi, np.pi, 256, endpoint=True)
>>> condition = X is not None and isinstance(X, np.ndarray) and (X.shape == (256,)) and np.allclose(X, expected_X)
>>> assert condition, 'X is not correctly calculated as a linear spaced array. Ensure np.linspace(-np.pi, np.pi, 256, endpoint=True) is used.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(X))
```
failure_message: Failed: X is not correctly computed as a linear spaced array.
success_message: Success: X is correctly computed as a linear spaced array!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q4-set-limits-4'
>>> max_score = 2.0
>>> score = 0
>>> expected_C = np.cos(X)
>>> expected_S = np.sin(X)
>>> condition = C is not None and S is not None and isinstance(C, np.ndarray) and isinstance(S, np.ndarray) and (C.shape == X.shape) and (S.shape == X.shape) and np.allclose(C, expected_C) and np.allclose(S, expected_S)
>>> assert condition, 'C and S are not calculated as cosine and sine of X. Ensure C, S = np.cos(X), np.sin(X) are used.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(C))
>>> responses = update_responses(question_id, str(S))
```
failure_message: Failed: C and S are not computed correctly as cosine and sine of X.
success_message: Success: C and S are computed correctly as cosine and sine of X!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q4-set-limits-5'
>>> max_score = 2.0
>>> score = 0
>>> line = plt_cosine[0]
>>> condition = line.get_color() == 'blue' and line.get_linewidth() == 2.5 and (line.get_linestyle() == '-')
>>> assert condition, "Cosine is not plotted with the correct style. Ensure color='blue', linewidth=2.5, linestyle='-' are used."
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(plt_cosine))
```
failure_message: Failed: Cosine is not plotted with the correct style.
success_message: Success: Cosine is plotted with the correct color, linewidth, and linestyle!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q4-set-limits-6'
>>> max_score = 2.0
>>> score = 0
>>> line = plt_sine[0]
>>> condition = line.get_color() == 'red' and line.get_linewidth() == 2.5 and (line.get_linestyle() == '-')
>>> assert condition, "Sine is not plotted with the correct style. Ensure color='red', linewidth=2.5, linestyle='-' are used."
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(plt_sine))
```
failure_message: Failed: Sine is not plotted with the correct style.
success_message: Success: Sine is plotted with the correct color, linewidth, and linestyle!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q4-set-limits-7'
>>> max_score = 2.0
>>> score = 0
>>> line = plt_sine[0]
>>> condition = plt_sine is not None and isinstance(plt_sine, list) and (len(plt_sine) > 0) and (line.get_xdata().tolist() == X.tolist()) and (line.get_ydata().tolist() == S.tolist()) and (line.get_color() == 'red') and (line.get_linewidth() == 2.5) and (line.get_linestyle() == '-')
>>> assert condition, "Sine plot is not correct. Ensure plt.plot(X, S, color='red', linewidth=2.5, linestyle='-') is used."
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(plt_sine))
>>> responses = update_responses(question_id, str(X))
>>> responses = update_responses(question_id, str(S))
```
failure_message: Failed: Sine is not plotted correctly with the specified style.
success_message: Success: Sine is plotted correctly with the specified style!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q4-set-limits-8'
>>> max_score = 2.0
>>> score = 0
>>> expected_xlim = (X.min() * 1.1, X.max() * 1.1)
>>> expected_ylim = (C.min() * 1.1, C.max() * 1.1)
>>> condition = plt_xlim == expected_xlim and plt_ylim == expected_ylim
>>> assert condition, f'Axis limits are not set correctly. Expected xlim: {expected_xlim} and ylim: {expected_ylim}.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(plt_xlim))
>>> responses = update_responses(question_id, str(plt_ylim))
```
failure_message: Failed: Axis limits are not set correctly.
success_message: Success: Axis limits are set correctly!


Tests for q5-set-ticks (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(6.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(44.0)
>>> log_variable('total-points', f'week8-lecturenotgraded, 2_refining_plot_q', 44.0)
>>> question_id = 'q5-set-ticks-1'
>>> max_score = 2.0
>>> score = 0
>>> expected_C = np.cos(X)
>>> expected_S = np.sin(X)
>>> condition = C is not None and S is not None and isinstance(C, np.ndarray) and isinstance(S, np.ndarray) and (C.shape == X.shape) and (S.shape == X.shape) and np.allclose(C, expected_C) and np.allclose(S, expected_S)
>>> assert condition, 'C and S are not calculated as cosine and sine of X. Ensure C, S = np.cos(X), np.sin(X) are used.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(C))
>>> responses = update_responses(question_id, str(S))
```
failure_message: Failed: C and S are not computed correctly as cosine and sine of X.
success_message: Success: C and S are computed correctly as cosine and sine of X!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q5-set-ticks-2'
>>> max_score = 2.0
>>> score = 0
>>> expected_xlim = (X.min() * 1.1, X.max() * 1.1)
>>> expected_ylim = (C.min() * 1.1, C.max() * 1.1)
>>> condition = plt_xlim == expected_xlim and plt_ylim == expected_ylim
>>> assert condition, f'Axis limits are not set correctly. Expected xlim: {expected_xlim} and ylim: {expected_ylim}.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(plt_xlim))
>>> responses = update_responses(question_id, str(plt_ylim))
```
failure_message: Failed: Axis limits are not set correctly.
success_message: Success: Axis limits are set correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q5-set-ticks-3'
>>> max_score = 2.0
>>> score = 0
>>> expected_xlim = (X.min() * 1.1, X.max() * 1.1)
>>> expected_ylim = (C.min() * 1.1, C.max() * 1.1)
>>> condition = plt_xlim == expected_xlim and plt_ylim == expected_ylim
>>> assert condition, f'Axis limits are not set correctly. Expected xlim: {expected_xlim} and ylim: {expected_ylim}.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(plt_xlim))
>>> responses = update_responses(question_id, str(plt_ylim))
```
failure_message: Failed: Axis limits are not set correctly.
success_message: Success: Axis limits are set correctly!


Tests for q6-move-spines (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(6.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(44.0)
>>> log_variable('total-points', f'week8-lecturenotgraded, 2_refining_plot_q', 44.0)
>>> question_id = 'q6-move-spines-1'
>>> max_score = 2.0
>>> score = 0
>>> bottom_spine_position = ax.spines['bottom'].get_position()[1]
>>> assert bottom_spine_position == 0, 'Bottom spine is not positioned at the data origin.'
>>> if bottom_spine_position == 0:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(bottom_spine_position))
```
failure_message: Failed: Bottom spine is not positioned at the data origin.
success_message: Success: Bottom spine is positioned at the data origin!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q6-move-spines-2'
>>> max_score = 2.0
>>> score = 0
>>> left_spine_position = ax.spines['left'].get_position()[1]
>>> condition = left_spine_position == 0
>>> assert condition, 'Left spine is not positioned at the data origin.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(left_spine_position))
```
failure_message: Failed: Left spine is not positioned at the data origin.
success_message: Success: Left spine is positioned at the data origin!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q6-move-spines-3'
>>> max_score = 2.0
>>> score = 0
>>> y_ticks_position = ax.yaxis.get_ticks_position()
>>> condition = y_ticks_position == 'left'
>>> assert condition, 'Y-axis ticks are not positioned on the left.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(y_ticks_position))
```
failure_message: Failed: Y-axis ticks are not positioned on the left.
success_message: Success: Y-axis ticks are positioned on the left!


Tests for q7-add-legend (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(2.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(44.0)
>>> log_variable('total-points', f'week8-lecturenotgraded, 2_refining_plot_q', 44.0)
>>> question_id = 'q7-add-legend-1'
>>> max_score = 2.0
>>> score = 0
>>> condition = True
>>> assert condition
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('2_refining_plot_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Legend is not defined correctly.
success_message: Success: Legend is defined correctly!
