```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("6_Curve_Fitting_q", "week_8", "lecturenotgraded", assignment_points = 45.0, assignment_tag = 'week8-lecturenotgraded')

# Initialize Otter
import otter
grader = otter.Notebook("6_Curve_Fitting_q.ipynb")
```

# 💻Activity: Curve Fitting

Engineers often have to deal with large amounts of raw data. 

To assist in data analysis it is common to fit a model to data to make useful insights easier.

The fitting process happens by performing computations to minimize an objective relationship. 

There are many optimization methods -- this is a rich field of science and engineering.

SciPy is a package for scientific computing in python that has many built in tools for optimization and fitting.

## Curve Fitting Example

Suppose we have some data on a sine function with some noise included, build a function to generate this data.


```python
# import numpy as np
...

# Do not change the random seed
np.random.seed(123)

# generate a x linear spaced vector from -5 to 5 with 50 points, use the np.linspace function
...

# generate y data which is 2.9 * np.sin(1.5 * x_data) , add to this some random noise using the np.random.normal function
...

```

```python
grader.check("generate-sine-data")
```

When dealing with data it is usually helpful to visualize the data as a graph

```python
# plot the raw data using matplotlib
...

# plot the x_data and y_data using matplotlib, use the plt.plot function, make the points blue squares using the "-s" marker, save to the variable plot_
...

```

```python
grader.check("generate-initial-plot")
```

We know the data lies approximately on a sine wave. We do not know the amplitudes or the period.

We can estimate those by least squares curve fitting. First we have to define the test function to fit. In this case, a sine with unknown amplitude and period:

```python
# write a function that takes x and a and b and returns a * np.sin(b * x), the sine function with unknown amplitude and period
...

```

```python
grader.check("make-sine-function")
```

We then use scipy.optimize.curve_fit() to find a and b:

```python
# import the optimize module from scipy
...

# Do not change the random seed
np.random.seed(123)

# use the curve_fit function in optimize to fit the sine_func to the data, use the p0 parameter to set the initial guess for the parameters for the initial guess we can use the values of [2,2 ]
...

# print the parameters
...

```

```python
grader.check("scipy-optimize-curve-fit")
```

## Question 4 (Points: 10.0): Visualizing our Results

```python
# Create a new figure and axis, make it using plt.subplots, and set the figsize to (10, 6)
# save the figure to the variable fig, and the axis to the variable ax
...

# Plot the raw data
# save the plot to the variable raw_data_plot, make the points blue circles using the "bo" marker, add a label "Raw data" using the optional label parameter
...

# Plot the fitted function 
# save the plot to the variable fitted_function_plot, make the line red using the "red" tag, add a label "Fitted function" using the optional label parameter
...

# Label the axes "X" and "Y", using the plt.xlabel and plt.ylabel functions
...

# Set y-axis limits, set the y-axis limits to be 1.2 times the minimum and maximum of the y_data using the plt.ylim function
...

# Add a legend, set the loc to "upper right" using the plt.legend function
...

```

```python
grader.check("sine-viz-plot")
```

## Question 5 (Points: 9.0): Example on a Exponential Function

```python
# import numpy as np
# import the curve_fit function from scipy.optimize
# import matplotlib.pyplot as plt
...

# Reset the random seed and recompute to check reproducibility
np.random.seed(123)

# Define the exponential function to fit, the expression is a * np.exp(-b * x) + c
...

# Generate noisy data, from 0 to 4 with 50 points, use the np.linspace function
...

# generate ydata which is exp_func(xdata, 2.5, 1.3, 0.5) + 0.2 * np.random.normal(size=len(xdata))
...

# Fit the data, save the fit parameters in popt, and the covariance matrix in pcov
...

# Plot the data and fitted function, make the points blue circles using the "bo" marker, make the line red using the "red" tag, add a label "Raw data" using the optional label parameter.
# save the plot to the variable raw_data_plot
...

# plot the fitted function using matplotlib, make the line red using the "red" tag, add a label "Fitted function" using the optional label parameter.
# you can pass the parameters to the exp_func function using the * operator to unpack the parameters
# save the plot to the variable fitted_function_plot
...

# add a label for the x-axis "X" using the plt.xlabel function
...

# add a label for the y-axis "Y" using the plt.ylabel function
...

# add a legend to the plot using the plt.legend function, the legend should be in the upper right corner using the "upper right" location
...

plt.show()


# calculate the residuals, the difference between the ydata and the fitted function, the residuals are the difference between the ydata and the fitted function
...

# calculate the sum of the squared residuals, this is a common metric for the quality of the fit, save to the variable ss_res
...

# calculate the total sum of squares, this is a common metric for the quality of the fit, save to the variable ss_tot
...

# calculate the r-squared value, this is a common metric for the quality of the fit, the r-squared value is 1 - (ss_res / ss_tot)
...

# print the r-squared value
...
```

```python
grader.check("exponential-function-fit")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week8-lecturenotgraded", "6_Curve_Fitting_q")
```

Tests for exponential-function-fit (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(6.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(45.0)
>>> log_variable('total-points', f'week8-lecturenotgraded, 6_Curve_Fitting_q', 45.0)
>>> question_id = 'exponential-function-fit-1'
>>> max_score = 1.0
>>> score = 0
>>> condition = 'np' in globals() and 'curve_fit' in globals() and ('plt' in globals())
>>> assert condition, 'Required modules are not imported correctly. Ensure numpy, curve_fit, and matplotlib.pyplot are imported.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(condition))
```
failure_message: Failed: Required modules are not imported correctly.
success_message: Success: All required modules are imported correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'exponential-function-fit-2'
>>> max_score = 1.0
>>> score = 0
>>> condition = 'xdata' in globals() and isinstance(xdata, np.ndarray) and (len(xdata) == 50) and np.isclose(xdata[0], 0) and np.isclose(xdata[-1], 4)
>>> assert condition, 'xdata is not generated correctly. Ensure np.linspace(0, 4, 50) is used.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(xdata))
```
failure_message: Failed: xdata is not generated correctly using np.linspace.
success_message: Success: xdata is generated correctly using np.linspace!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'exponential-function-fit-3'
>>> max_score = 2.0
>>> score = 0
>>> condition = 'popt' in globals() and 'pcov' in globals() and isinstance(popt, np.ndarray) and isinstance(pcov, np.ndarray) and (len(popt) == 3) and (popt[0] > 2) and (popt[0] < 3) and (popt[1] > 1) and (popt[1] < 2) and (popt[2] > 0) and (popt[2] < 1)
>>> assert condition, 'curve_fit is not used correctly. Ensure curve_fit(exp_func, xdata, ydata) is used.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(popt))
>>> responses = update_responses(question_id, str(pcov))
```
failure_message: Failed: Data is not fitted correctly using curve_fit.
success_message: Success: Data is fitted correctly using curve_fit!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'exponential-function-fit-4'
>>> max_score = 2.0
>>> score = 0
>>> residuals_test = ydata - exp_func(xdata, *popt)
>>> ss_res_test = np.sum(residuals_test ** 2)
>>> ss_tot_test = np.sum((ydata - np.mean(ydata)) ** 2)
>>> expected_r_squared = 1 - ss_res_test / ss_tot_test
>>> condition = 'r_squared' in globals() and np.isclose(r_squared, expected_r_squared)
>>> assert condition, f'R-squared value is not calculated correctly. Expected {expected_r_squared} but got {r_squared}.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(r_squared))
```
failure_message: Failed: R-squared value is not calculated correctly.
success_message: Success: R-squared value is calculated correctly!


Tests for generate-initial-plot (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(6.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(45.0)
>>> log_variable('total-points', f'week8-lecturenotgraded, 6_Curve_Fitting_q', 45.0)
>>> question_id = 'generate-initial-plot-1'
>>> max_score = 1.0
>>> score = 0
>>> condition = 'plt' in globals()
>>> assert condition, "You must import matplotlib.pyplot using the alias 'plt'."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(condition))
```
failure_message: Failed: You must import 'matplotlib.pyplot' using the alias 'plt'.
success_message: Success: 'matplotlib.pyplot' is imported as 'plt'!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'generate-initial-plot-2'
>>> max_score = 2.0
>>> score = 0
>>> condition = 'plot_' in globals() and isinstance(plot_, list) and (len(plot_) > 0) and (plot_[0].get_xdata().tolist() == x_data.tolist()) and (plot_[0].get_ydata().tolist() == y_data.tolist())
>>> assert condition, 'x_data and y_data are not correctly plotted. Ensure plt.plot(x_data, y_data) is used and assigned to plot_.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(x_data))
>>> responses = update_responses(question_id, str(y_data))
>>> responses = update_responses(question_id, str(plot_))
```
failure_message: Failed: x_data and y_data are not correctly plotted using plt.plot().
success_message: Success: x_data and y_data are correctly plotted using plt.plot()!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'generate-initial-plot-3'
>>> max_score = 3.0
>>> score = 0
>>> line = plot_[0]
>>> marker = line.get_marker()
>>> linestyle = line.get_linestyle()
>>> color = line.get_color()
>>> expected_marker = 's'
>>> expected_linestyle = '-'
>>> expected_color = '#1f77b4'
>>> condition = marker == expected_marker and linestyle == expected_linestyle and (color == expected_color)
>>> assert condition, "Plot style is incorrect. Ensure plt.plot(x_data, y_data, '-s') is used for blue squares."
>>> if condition:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(marker))
>>> responses = update_responses(question_id, str(linestyle))
>>> responses = update_responses(question_id, str(color))
```
failure_message: Failed: The plot does not correctly use blue squares as markers with '-s' style.
success_message: Success: The plot correctly uses blue squares as markers with '-s' style!


Tests for generate-sine-data (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(6.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(45.0)
>>> log_variable('total-points', f'week8-lecturenotgraded, 6_Curve_Fitting_q', 45.0)
>>> question_id = 'generate-sine-data-1'
>>> max_score = 1.0
>>> score = 0
>>> condition = 'np' in globals()
>>> assert condition, "You must import numpy using the alias 'np'."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(condition))
```
failure_message: Failed: You must import numpy using the alias 'np'.
success_message: Success: 'numpy' is imported as 'np'!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> question_id = 'generate-sine-data-2'
>>> max_score = 2.0
>>> score = 0
>>> expected_x_data = np.linspace(-5, 5, num=50)
>>> condition = 'x_data' in globals() and isinstance(x_data, np.ndarray) and (x_data.shape == (50,)) and np.allclose(x_data, expected_x_data)
>>> assert condition, 'x_data is not correctly generated. Ensure np.linspace(-5, 5, num=50) is used.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(x_data))
```
failure_message: Failed: x_data is not correctly generated using np.linspace.
success_message: Success: x_data is correctly generated using np.linspace!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'generate-sine-data-3'
>>> max_score = 3.0
>>> score = 0
>>> np.random.seed(123)
>>> expected_y_data = 2.9 * np.sin(1.5 * x_data) + np.random.normal(size=50)
>>> condition = 'y_data' in globals() and isinstance(y_data, np.ndarray) and (y_data.shape == (50,)) and np.allclose(y_data, expected_y_data)
>>> assert condition, 'y_data is not correctly generated. Ensure the formula 2.9 * np.sin(1.5 * x_data) + np.random.normal(size=50) is used with the fixed random seed.'
>>> if condition:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(y_data))
```
failure_message: Failed: y_data is not correctly generated with the expected noise pattern.
success_message: Success: y_data is correctly generated with the expected noise pattern!


Tests for make-sine-function (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(8.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(45.0)
>>> log_variable('total-points', f'week8-lecturenotgraded, 6_Curve_Fitting_q', 45.0)
>>> question_id = 'make-sine-function-1'
>>> max_score = 1.0
>>> score = 0
>>> condition = 'sine_func' in globals() and callable(sine_func)
>>> assert condition, "The function 'sine_func' is not defined or named incorrectly. Ensure the function name is 'sine_func'."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(condition))
```
failure_message: Failed: The function 'sine_func' is not defined or named incorrectly.
success_message: Success: The function 'sine_func' is defined correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'make-sine-function-2'
>>> max_score = 1.0
>>> score = 0
>>> x = np.array([0, np.pi / 2, np.pi])
>>> a = 2
>>> b = 1
>>> result = sine_func(x, a, b)
>>> condition = isinstance(result, np.ndarray)
>>> assert condition, 'The function does not return a NumPy array. Ensure you are using NumPy functions correctly.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(result))
```
failure_message: Failed: The function does not return a NumPy array.
success_message: Success: The function returns a NumPy array!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'make-sine-function-3'
>>> max_score = 2.0
>>> score = 0
>>> x = np.array([0, np.pi / 2, np.pi])
>>> a = 1
>>> b = 1
>>> expected_output = np.array([0, 1, 0])
>>> actual_output = sine_func(x, a, b)
>>> condition = np.allclose(actual_output, expected_output)
>>> assert condition, f'Expected {expected_output} but got {actual_output}. Ensure the function is calculating the sine correctly.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(x))
>>> responses = update_responses(question_id, str(a))
>>> responses = update_responses(question_id, str(b))
>>> responses = update_responses(question_id, str(expected_output))
>>> responses = update_responses(question_id, str(actual_output))
```
failure_message: Failed: The function does not calculate the sine correctly for basic inputs.
success_message: Success: The function calculates the sine correctly for basic inputs!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'make-sine-function-4'
>>> max_score = 2.0
>>> score = 0
>>> x = np.array([0, np.pi / 2, np.pi])
>>> a = 2
>>> b = 1
>>> expected_output = np.array([0, 2, 0])
>>> actual_output = sine_func(x, a, b)
>>> condition = np.allclose(actual_output, expected_output)
>>> assert condition, f"Expected {expected_output} but got {actual_output}. Ensure the function adjusts the amplitude using 'a'."
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(x))
>>> responses = update_responses(question_id, str(a))
>>> responses = update_responses(question_id, str(b))
>>> responses = update_responses(question_id, str(expected_output))
>>> responses = update_responses(question_id, str(actual_output))
```
failure_message: Failed: The function does not adjust the amplitude correctly.
success_message: Success: The function correctly adjusts the amplitude!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'make-sine-function-5'
>>> max_score = 2.0
>>> score = 0
>>> x = np.linspace(0, 2 * np.pi, 5)
>>> a = 3
>>> b = 1
>>> expected_output = 3 * np.sin(x)
>>> actual_output = sine_func(x, a, b)
>>> condition = np.allclose(actual_output, expected_output)
>>> assert condition, f'Expected {expected_output} but got {actual_output}. Ensure the function handles multiple input values.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(x))
>>> responses = update_responses(question_id, str(a))
>>> responses = update_responses(question_id, str(b))
>>> responses = update_responses(question_id, str(expected_output))
>>> responses = update_responses(question_id, str(actual_output))
```
failure_message: Failed: The function does not handle multiple input values correctly.
success_message: Success: The function handles multiple input values correctly!


Tests for scipy-optimize-curve-fit (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(10.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(45.0)
>>> log_variable('total-points', f'week8-lecturenotgraded, 6_Curve_Fitting_q', 45.0)
>>> question_id = 'scipy-optimize-curve-fit-1'
>>> max_score = 1.0
>>> score = 0
>>> condition = 'optimize' in globals()
>>> assert condition, "You must import 'optimize' from 'scipy'. Ensure the import statement is: 'from scipy import optimize'."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(condition))
```
failure_message: Failed: You must import 'optimize' from 'scipy'.
success_message: Success: 'optimize' is imported from 'scipy'!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'scipy-optimize-curve-fit-2'
>>> max_score = 2.0
>>> score = 0
>>> condition = 'params' in globals() and 'params_covariance' in globals() and isinstance(params, np.ndarray) and isinstance(params_covariance, np.ndarray) and (params.shape == (2,))
>>> assert condition, 'curve_fit is not used correctly. Ensure optimize.curve_fit() is used and returns params and params_covariance.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(params))
>>> responses = update_responses(question_id, str(params_covariance))
```
failure_message: Failed: curve_fit is not used correctly to fit the data.
success_message: Success: curve_fit is used correctly to fit the data!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'scipy-optimize-curve-fit-3'
>>> max_score = 2.0
>>> score = 0
>>> condition = 'params' in globals() and len(params) == 2 and np.allclose(optimize.curve_fit(sine_func, x_data, y_data, p0=[2, 2])[0], params)
>>> assert condition, 'The initial guess is not set correctly. Ensure p0=[2, 2] is used in curve_fit.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(params))
```
failure_message: Failed: The initial guess for curve fitting is not set correctly.
success_message: Success: The initial guess for curve fitting is set correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'scipy-optimize-curve-fit-4'
>>> max_score = 2.0
>>> score = 0
>>> condition = isinstance(params, np.ndarray) and len(params) == 2 and all((isinstance(p, (float, np.float64)) for p in params)) and (-10 < params[0] < 10) and (-10 < params[1] < 10)
>>> assert condition, 'Fitted parameters are not in the expected range or are not of type float.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(params))
```
failure_message: Failed: The fitted parameters are not in the expected range or not of correct type.
success_message: Success: The fitted parameters are in the expected range and of correct type!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'scipy-optimize-curve-fit-5'
>>> max_score = 2.0
>>> score = 0
>>> np.random.seed(123)
>>> recomputed_y_data = 2.9 * np.sin(1.5 * x_data) + np.random.normal(size=50)
>>> recomputed_params, _ = optimize.curve_fit(sine_func, x_data, recomputed_y_data, p0=[2, 2])
>>> condition = np.allclose(params, recomputed_params)
>>> assert condition, 'The curve fitting is not reproducible. Ensure the random seed is fixed with np.random.seed(123).'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(params))
```
failure_message: Failed: The curve fitting is not reproducible with the fixed random seed.
success_message: Success: The curve fitting is reproducible with the fixed random seed!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import re
>>> from io import StringIO
>>> import sys
>>> question_id = 'scipy-optimize-curve-fit-6'
>>> max_score = 1.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> print(params)
>>> output_value = output.getvalue().strip()
>>> expected_output = re.escape(np.array2string(params))
>>> condition = re.search(expected_output, output_value, re.IGNORECASE)
>>> sys.stdout = sys.__stdout__
>>> assert condition, f'Expected the printed parameters to be: {params}, but got: {output_value}. Ensure the parameters are printed using print(params).'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(output_value))
>>> responses = update_responses(question_id, str(params))
```
failure_message: Failed: The parameters are not printed correctly.
success_message: Success: The parameters are printed correctly!


Tests for sine-viz-plot (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(9.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(45.0)
>>> log_variable('total-points', f'week8-lecturenotgraded, 6_Curve_Fitting_q', 45.0)
>>> question_id = 'sine-viz-plot-1'
>>> max_score = 1.0
>>> score = 0
>>> condition = 'fig' in globals() and 'ax' in globals() and isinstance(fig, plt.Figure) and isinstance(ax, plt.Axes) and (fig.get_size_inches().tolist() == [10, 6])
>>> assert condition, 'Figure and axis are not created correctly. Ensure plt.subplots(figsize=(10, 6)) is used.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(fig))
>>> responses = update_responses(question_id, str(ax))
```
failure_message: Failed: Figure and axis are not created correctly with plt.subplots().
success_message: Success: Figure and axis are created correctly with plt.subplots()!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'sine-viz-plot-2'
>>> max_score = 2.0
>>> score = 0
>>> line = fitted_function_plot[0]
>>> condition = np.allclose(line.get_xdata(), x_data) and np.allclose(line.get_ydata(), sine_func(x_data, *params)) and (line.get_color() == 'red') and (line.get_label() == 'Fitted function')
>>> assert condition, "Fitted function is not plotted correctly. Ensure ax.plot(x_data, sine_func(x_data, *params), color='red', label='Fitted function') is used."
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(fitted_function_plot))
```
failure_message: Failed: Fitted function is not plotted correctly with a red line or labeled incorrectly.
success_message: Success: Fitted function is plotted correctly with a red line and labeled as 'Fitted function'!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'sine-viz-plot-3'
>>> max_score = 1.0
>>> score = 0
>>> x_label = ax.get_xlabel()
>>> condition = x_label == 'X'
>>> assert condition, "X-axis label is not set correctly. Ensure ax.set_xlabel('X') is used."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(x_label))
```
failure_message: Failed: X-axis label is not set correctly as 'X'.
success_message: Success: X-axis label is set correctly as 'X'!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'sine-viz-plot-4'
>>> max_score = 1.0
>>> score = 0
>>> y_label = ax.get_ylabel()
>>> condition = y_label == 'Y'
>>> assert condition, "Y-axis label is not set correctly. Ensure ax.set_ylabel('Y') is used."
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(y_label))
```
failure_message: Failed: Y-axis label is not set correctly as 'Y'.
success_message: Success: Y-axis label is set correctly as 'Y'!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'sine-viz-plot-5'
>>> max_score = 2.0
>>> score = 0
>>> expected_ymin = min(y_data) * 1.2
>>> expected_ymax = max(y_data) * 1.2
>>> y_limits = ax.get_ylim()
>>> condition = np.isclose(y_limits[0], expected_ymin) and np.isclose(y_limits[1], expected_ymax)
>>> assert condition, f'Y-axis limits are not set correctly. Expected ({expected_ymin}, {expected_ymax}). Ensure ax.set_ylim(min(y_data)*1.2, max(y_data)*1.2) is used.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(y_limits))
```
failure_message: Failed: Y-axis limits are not set correctly.
success_message: Success: Y-axis limits are set correctly as 1.2 times the minimum and maximum of y_data!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'sine-viz-plot-6'
>>> max_score = 2.0
>>> score = 0
>>> legend = ax.get_legend()
>>> legend_location = legend._loc
>>> expected_location = 1
>>> condition = legend is not None and legend_location == expected_location
>>> assert condition, "Legend is not displayed correctly. Ensure ax.legend(loc='upper right') is used."
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('6_Curve_Fitting_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(legend_location))
```
failure_message: Failed: Legend is not displayed correctly in the upper right corner.
success_message: Success: Legend is displayed correctly in the upper right corner!
