```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("13_Philly_Python_Challenge", "week_1", "lecture", assignment_points = 5.0, assignment_tag = 'week1-lecture')

# Initialize Otter
import otter
grader = otter.Notebook("13_Philly_Python_Challenge.ipynb")
```

# 🥨 Philly Pretzel Challenge!

![Philly Pretzel Challenge](./assets/figures/pretzel.webp)

## The Problem

In Philadelphia, soft pretzels are a beloved snack. Imagine a group of friends buying a box of pretzels to share:

* There are 23 pretzels in the box.
* There are 5 friends in the group.

Write a program to calculate:
1. How many pretzels each friend gets (integer division).
2. How many pretzels are left over (modulus operator).

Detailed Instructions:

1. **Define the Total Number of Pretzels**:
   - Create a variable named `pretzels` and assign it the value `23`.
   - This represents the total number of soft pretzels available in the box.

2. **Define the Number of Friends**:
   - Create a variable named `friends` and assign it the value `5`.
   - This represents the number of friends sharing the pretzels.

3. **Calculate How Many Pretzels Each Friend Gets**:
   - Use integer division (`//`) to divide `pretzels` by `friends`.
   - Integer division gives the whole number of pretzels each friend can get without considering leftovers.
   - Store the result in a variable named `pretzels_per_friend`.

4. **Calculate the Leftover Pretzels**:
   - Use the modulus operator (`%`) to divide `pretzels` by `friends` and determine the remainder.
   - The remainder represents the number of pretzels left after evenly distributing them among the friends.
   - Store the result in a variable named `leftover_pretzels`.

5. **Print the Results**:
   - Use the `print()` function to display how many pretzels each friend gets and how many are left over.
   - Use f-strings to include the calculated values in the output.

   ```python
   print(f'Each friend gets {pretzels_per_friend} pretzels.')
   print(f'There are {leftover_pretzels} pretzels left over.')
   ```

   - The output will look like this:
     ```
     Each friend gets 4 pretzels.
     There are 3 pretzels left over.
     ```

```python
def pretzel_calculator():
    # This line of code defines a function
    # You should write your code where the ... is located
    # Do not change code above this line 
    
    # 1. Define the total number of pretzels and friends
    ...

    # 2. Calculate How Many Pretzels Each Friend Gets
    ...

    # 3. Calculate How Many Pretzels Are Left Over
    ...

    # 4. Print the Results
    # Tip: Copy the code from above we provided it for you
    ...

    # Do not change code below this line
    return pretzels_per_friend, leftover_pretzels
```

```python
grader.check("activity-2-Philly Pretzel Challenge")
```

## 🎉 Well Done!

Congratulations on completing the Philly Python Challenge! Now go grab a soft pretzel to celebrate! 🥨

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week1-lecture", "13_Philly_Python_Challenge")
```

Tests for activity-2-Philly Pretzel Challenge (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> from io import StringIO
>>> import sys
>>> max_question_points = str(5.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(5.0)
>>> log_variable('total-points', f'Reading-Week-X, 13_Philly_Python_Challenge', 5.0)
>>> question_id = 'activity-2-Philly Pretzel Challenge-1'
>>> max_score = 1.0
>>> score = 0
>>> captured_output = StringIO()
>>> sys.stdout = captured_output
>>> pretzels_per_friend, leftover_pretzels = pretzel_calculator()
>>> sys.stdout = sys.__stdout__
>>> output = captured_output.getvalue().strip()
>>> assert 'Each friend gets' in output, "Output must contain 'Each friend gets'."
>>> if 'Each friend gets' in output:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('13_Philly_Python_Challenge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Output does not include 'Each friend gets'.
success_message: Correctly included the phrase 'Each friend gets'.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> from io import StringIO
>>> import sys
>>> question_id = 'activity-2-Philly Pretzel Challenge-2'
>>> max_score = 1.0
>>> score = 0
>>> captured_output = StringIO()
>>> sys.stdout = captured_output
>>> pretzels_per_friend, leftover_pretzels = pretzel_calculator()
>>> sys.stdout = sys.__stdout__
>>> output = captured_output.getvalue().strip()
>>> assert 'There are' in output, "Output must contain 'There are'."
>>> if 'There are' in output:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('13_Philly_Python_Challenge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Output does not include 'There are'.
success_message: Correctly included the phrase 'There are'.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> from io import StringIO
>>> import sys
>>> question_id = 'activity-2-Philly Pretzel Challenge-3'
>>> max_score = 1.0
>>> score = 0
>>> captured_output = StringIO()
>>> sys.stdout = captured_output
>>> pretzels_per_friend, leftover_pretzels = pretzel_calculator()
>>> sys.stdout = sys.__stdout__
>>> output = captured_output.getvalue().strip()
>>> assert 'pretzels left over' in output, "Output must contain 'pretzels left over'."
>>> if 'pretzels left over' in output:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('13_Philly_Python_Challenge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Output does not include 'pretzels left over'.
success_message: Correctly included the phrase 'pretzels left over'.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> from io import StringIO
>>> import sys
>>> question_id = 'activity-2-Philly Pretzel Challenge-4'
>>> max_score = 1.0
>>> score = 0
>>> captured_output = StringIO()
>>> sys.stdout = captured_output
>>> pretzels_per_friend, leftover_pretzels = pretzel_calculator()
>>> sys.stdout = sys.__stdout__
>>> output = captured_output.getvalue().strip()
>>> pattern = 'Each friend gets (\\d+) pretzels\\.\\s*There are (\\d+) pretzels left over\\.'
>>> match = re.match(pattern, output)
>>> assert match is not None, 'Output format is incorrect.'
>>> pretzels_per_friend = int(match.group(1))
>>> leftover_pretzels = int(match.group(2))
>>> if match is not None:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('13_Philly_Python_Challenge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Output format is incorrect or numbers could not be extracted.
success_message: Output format is correct and numbers extracted successfully.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import sys
>>> from io import StringIO
>>> question_id = 'activity-2-Philly Pretzel Challenge-5'
>>> max_score = 1.0
>>> score = 0
>>> captured_output = StringIO()
>>> sys.stdout = captured_output
>>> pretzels_per_friend, leftover_pretzels = pretzel_calculator()
>>> assert pretzels_per_friend == 4, 'The number of pretzels per friend should be 4.'
>>> assert leftover_pretzels == 3, 'The number of leftover pretzels should be 3.'
>>> if pretzels_per_friend == 4 and leftover_pretzels == 3:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('13_Philly_Python_Challenge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Pretzel calculations are incorrect.
success_message: Correctly calculated the pretzels per friend and leftovers.
