```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_lab_building_bridge", "week_8", "lab", assignment_points = 126.0, assignment_tag = 'week8-lab')

# Initialize Otter
import otter
grader = otter.Notebook("1_lab_building_bridge.ipynb")
```

# 🏗️ **Mission: Bridge Safety Analysis – The Engineer’s Challenge!** 🌉


## **🚀 Introduction: Welcome to the Team!**
Congratulations, Engineer! 🎉 You’ve just been hired as the lead **structural engineer** at **SteelWorks Inc.** Your first mission? **Analyze the safety of the Golden Gate Bridge under different loading conditions**—or risk total catastrophe! 😱

Our bridge must **withstand traffic loads, fierce winds, and the occasional flock of overly enthusiastic pigeons**. 🐦💨 As a skilled Python programmer, you’ll use **Object-Oriented Programming (OOP), scientific computing, and graphing techniques** to ensure that **our bridge doesn’t collapse into the ocean!** 🌊


## **🛠️ Engineering Context: Why This Matters?**
Every year, bridges are **designed, tested, and maintained** to ensure safety for **millions of people**. A **poorly designed bridge** can lead to **catastrophic failures**—remember the **Tacoma Narrows Bridge disaster**? 🌪️🌉 *Oops.*

By analyzing real-world **deflection curves, stress distributions, and wind loads**, you will:
💡 Learn **fundamental mechanics principles** using computational tools.
💡 Develop **engineering intuition** on load-bearing structures.
💡 Prepare for your future as an **engineering superhero!** 🦸‍♂️⚡


## **🎯 Your Challenge: Save the Bridge!**
You will **design a Python simulation** that includes:
🛑 **Beam Deflection Analysis**: How much will our bridge bend under load?
🛑 **Stress & Safety Factor Checks**: Will our materials **withstand** the forces?
🛑 **Wind Load Effects**: What happens when strong winds hit the bridge? 🌬️
🛑 **Graphical Reports**: Visualize stress, deflection, and safety margins.



## **🔧 Engineering Formulas Used**
We’ll use some **key engineering equations** in our Python models:


### **🌀 Beam Deflection Equation** (Traffic Load)
$$
y(x) = \frac{q x (L-x)}{24 E I} (L^2 - x^2)
$$


### **💪 Bending Stress in the Beam**
$$
\sigma(x) = \frac{M(x) c}{I}, \quad M(x) = \frac{q x (L - x)}{2}
$$


### **🌬️ Wind Load Contribution**
$$
M_{wind}(x) = W x (L - x)
$$

$$
M_{total}(x) = M(x) + M_{wind}(x)
$$


## **📈 Graphing & Visualization**
You’ll use **Matplotlib** to create stunning plots for:\
📊 **Bridge deflection curves** to see how much it bends.\
📊 **Stress distribution plots** to analyze failure points.\
📊 **Wind load impact graphs** to visualize dangerous weather effects.\


## **💾 Coding Concepts You’ll Apply**
🔹 **Object-Oriented Programming (OOP)** – Classes, Inheritance, Encapsulation.\
🔹 **Encapsulation with `@property` methods** – Dynamic safety factor updates.


## **Step 1: Abstract Base Class for Bridge Components**
We enforce structure using an **abstract class** for **bridge components**. Remember, this is a great way to protect you from yourself! Moreover, as an engineer, you need to make sure that you do not miss the safety factor! Therefore we are going to use getter and setter methods to ensure that we don't accidentally mess with, or include a safety factor that is less than 1.0!

Recall a safety factor is a factor of safety that is used to ensure that the bridge is safe. It is defined as the ratio of the yield strength to the stress in the bridge.


```python
# import the necessary libraries 
# from abc import ABC, abstractmethod
# import numpy as np
# import matplotlib.pyplot as plt
...

# Define the BridgeComponent class make sure that it inherits from ABC
...
    """Abstract base class for all bridge components.
    
    This class defines the common interface and properties that all bridge components must implement.
    It enforces safety requirements through validation of safety factors.
    
    Attributes:
        material (str): The material the component is made of
        length (float): Length of the component in meters
        yield_strength (float): The yield strength of the material in Pa
        _safety_factor (float): Safety factor that must be > 1.0
    """
    # Define the __init__ method
    # This is the constructor for the class look at the docstring to understand the arguments
    ...
        """Initialize a bridge component.
        
        Args:
            material (str): Material name/type
            length (float): Component length in meters
            yield_strength (float): Material yield strength in Pa
            safety_factor (float): Safety factor > 1.0
        """
        # Assign the atributes to the instance of the class, for the safety factor we should use a private variable commonly denoted by a single underscore in front of the variable name
        ...

    # Define the safety_factor using the @property decorator
    # This should return the value for the safety factor
    ...
    
    # Define the setter for the safety factor
    # This should set the value for the safety factor
    # the syntax is @property_name.setter
    # make sure that the setter checks if the value is greater than 1.0
    # if it is less than 1.0, it should raise a ValueError("Safety factor must be greater than 1.0")
    ...
    
    # Define the compute_stress method
    # This should be an abstract method that must be implemented by the subclasses
    # for now, we will just pass the method
    ...
    
    # Define the is_safe method
    # This should check if the stress is less than the yield strength divided by the safety factor, save the result in a variable called safe_limit
    # then return the result of the comparison between the stress and the safe_limit
    # the safe_limit can be a local variable or a property of the class
    ...
```

```python
grader.check("q1-bridge-component")
```

## Question 1 (Points: 12.0): **Step 2: Beam and Cable Components**

Bridges are built up of beams and cables, so we need to define the components that make up the bridge. A beam and a cable are both bridge components, so it makes sense to inherit from the BridgeComponent class.

We define **`Beam`** and **`Cable`**, each with a **compute_stress** function.


```python
# define the Beam class, make sure that it inherits from BridgeComponent
...
    """Represents a bridge beam that undergoes bending stress.
    
    A beam is a structural element that primarily resists loads applied laterally to its axis.
    It experiences bending stress due to these loads.
    
    Args:
        material (str): Material name of the beam
        length (float): Length of the beam in meters
        yield_strength (float): Yield strength of the material in Pa
        safety_factor (float): Safety factor to apply to yield strength
        moment_of_inertia (float): Second moment of area in m^4
        neutral_axis_distance (float): Distance from neutral axis to outer fiber in m
        uniform_load (float): Uniformly distributed load in N/m
    """

    # define the __init__ method, make sure that it calls the __init__ method of the BridgeComponent class, use super() to do this, it should take the arguments from the docstring
    ...

        # This is where you should call the __init__ method of the BridgeComponent class, use super() to do this, it should take the arguments from the docstring
        ...

        # This is where you should define the instance variables for the moment of inertia, neutral axis distance, and uniform load
        ...

    # define the compute_stress method
    # This should compute the maximum bending stress in the beam
    ...
        """Computes maximum bending stress in the beam.
        
        Calculates the maximum bending stress at the mid-span of the beam using
        the flexure formula σ = My/I, where M is the bending moment, y is the distance
        from neutral axis, and I is the moment of inertia.
        
        Returns:
            float: Maximum bending stress in Pa
        """
        # Compute the mid-span of the beam the length divided by 2
        ...

        # Compute the bending moment at the mid-span of the beam using the uniform load and the length of the beam
        # the formula is M = qx(L-x)/2, where q is the uniform load, x is the mid-span of the beam, and L is the length of the beam
        ...

        # Compute the maximum bending stress at the mid-span of the beam using the bending moment, the neutral axis distance, and the moment of inertia
        # the formula is σ = My/I, where M is the bending moment, y is the neutral axis distance, and I is the moment of inertia
        ...

    # define the plot_deflection method
    # This should plot the deflection curve of the beam using LaTeX equations
    ...
        """Plots the deflection curve of the beam using LaTeX equations.
        
        Creates a matplotlib figure showing the beam's deflection along its length.
        The deflection is calculated using the elastic curve equation for a uniformly
        loaded beam. The plot includes axis labels, title with the equation, and grid.
        
        Returns:
            None
        """
        # Compute the deflection of the beam using the elastic curve equation for a uniformly loaded beam
        # set x to be a linspaced array from 0 to the length of the beam with 100 points, use np.linspace
        # set y to be the deflection of the beam using the elastic curve equation for a uniformly loaded beam
        # the formula is y = (qx(L-x)/24EI) * (L^2 - x^2), where q is the uniform load, x is the position along the beam, L is the length of the beam, E is the elastic modulus, and I is the moment of inertia
        ...

        # plot the deflection of the beam using the elastic curve equation for a uniformly loaded beam
        # make a figure using the plt.figure() function, and assign it to the variable fig
        # make an axis using the fig.add_subplot() function, and assign it to the variable ax
        # plot the deflection of the beam using the ax.plot() function
        # set the xlabel of the axis to be "Position along bridge (m)"
        # set the ylabel of the axis to be "Deflection (m)"
        # set the title of the axis to be the LaTeX equation for the deflection of the beam the string should be r"Bridge Deflection: $y(x) = \frac{q x (L-x)}{24EI} (L^2 - x^2)$"
        # call the legend of the axis using the ax.legend() function
        # call the grid of the axis using the ax.grid() function
        # call the show of the figure using the plt.show() function
        ...

# define the Cable class, make sure that it inherits from BridgeComponent
...
    """Represents a bridge cable under tension.
    
    A cable is a structural element that carries only tensile forces.
    It is commonly used in suspension and cable-stayed bridges.
    
    Args:
        material (str): Material name of the cable
        length (float): Length of the cable in meters
        yield_strength (float): Yield strength of the material in Pa
        safety_factor (float): Safety factor to apply to yield strength
        tension_force (float): Axial tension force in the cable in N
        cross_section_area (float): Cross-sectional area of the cable in m^2
    """

    # define the __init__ method, make sure that it calls the __init__ method of the BridgeComponent class, use super() to do this, it should take the arguments from the docstring
    ...

        # This is where you should call the __init__ method of the BridgeComponent class, use super() to do this, it should take the arguments from the docstring
        ...

        # This is where you should define the instance variables for the tension force and the cross-sectional area
        ...

    # define the compute_stress method
    # This should compute the tensile stress in the cable using the formula σ = F/A, where F is the tension force and A is the cross-sectional area
    ...
```

```python
grader.check("q2-beam-cable")
```

## Question 2 (Points: 30.0): **Step 4: Bridge Class Using Composition**
Now, we define a **bridge object** that **contains beams and cables**.


```python
# define the Bridge class, make sure that it inherits from BridgeComponent
...
    """A class representing a bridge structure composed of multiple components.

    This class manages a collection of bridge components like beams and cables,
    and provides methods to add components and check their structural safety.

    Attributes:
        name (str): The name identifier for the bridge
        length (float): The total length of the bridge in meters
        elastic_modulus (float): Young's modulus of the bridge material in Pascals
        components (list): List of BridgeComponent objects that make up the bridge
    """

    # define the __init__ method, make sure that it takes the arguments from the docstring
    ...
        """Initialize a new Bridge instance.

        Args:
            name (str): The name identifier for the bridge
            length (float): The total length of the bridge in meters
            elastic_modulus (float): Young's modulus of the bridge material in Pascals
        """
        # This is where you should define the instance variables for the name, length, elastic modulus, and components
        # add an empty list to the instance variables for the components
        ...

    # define the add_component method
    # This should add a bridge component to the bridge structure
    ...
        """Adds a bridge component to the bridge structure.

        Args:
            component (BridgeComponent): A bridge component object (beam, cable, etc.)
        """
        # add the component to the components list
        ...

    # define the check_safety method
    # This should perform a safety analysis on all bridge components
    ...
    
        """Performs a safety analysis on all bridge components.

        Calculates stress for each component and checks if it's within safe limits.
        Prints a detailed safety report including stress values and safety status
        for each component, followed by an overall bridge safety verdict.

        Returns:
            None
        """
        # initialize a boolean variable called bridge_safe to be True
        ...
        # print the name of the bridge using the print() function
        # the string should be f"\n🔍 Safety Analysis for <name> Bridge:"
        # recall \n is a newline character and f is used for formatted strings, <name> should be the name of the bridge
        ...

        # loop through the components list, and assign the current component to the variable component
        ...
            # compute the stress of the current component using the compute_stress method
            ...
            # check if the stress is safe using the is_safe method
            # here we are using polymorphism to call the is_safe method of the current component
            ...
            # get the type of the current component using the __class__.__name__ attribute
            # this is a way to get the type of the current component using Python's magic methods
            # __class__ returns the class object of the instance
            # __name__ returns the string name of the class
            ...
            
            # print the type of the current component, the stress, and the safety status
            print(f"  - {component_type}: Stress = {stress:.2f} Pa | Safe: {'✅' if safe else '⚠️'}")
            
            # if the stress is not safe, set the bridge_safe variable to False
            ...

        # print the final verdict
        print("\nFinal Verdict:", "✅ Bridge is SAFE" if bridge_safe else "⚠️ WARNING: Bridge is UNSAFE")

```

```python
grader.check("Bridge-Class-Composition")
```

## Question 3 (Points: 16.0): **Step 3: Wind and Load Analysis**
We extend the bridge to handle **wind effects**.


```python
# define the WindLoadedBridge class, make sure that it inherits from Bridge
...
    """A bridge class that accounts for wind loading effects.
    
    This class extends the base Bridge class to include wind force calculations
    and analysis. It provides methods to analyze and visualize the combined effects
    of traffic and wind loads on the bridge structure.
    
    Args:
        name (str): Name of the bridge
        length (float): Length of the bridge in meters
        elastic_modulus (float): Elastic modulus of the bridge material in Pascals
        wind_load (float): Wind load force per unit length (N/m)
        
    Attributes:
        bridge (Bridge): Reference to the base bridge object
        wind_load (float): Wind load magnitude
    """

    # define the __init__ method, make sure that it calls the __init__ method of the Bridge class, use super() to do this, it should take the arguments from the docstring
    ...

        # This is where you should call the __init__ method of the Bridge class, use super() to do this, it should take the arguments from the docstring
        ...

    def plot_bending_moment(self):
        """Plots the bending moment distribution with LaTeX equations.
        
        Creates a matplotlib figure showing three curves:
        1. Bending moment due to traffic load
        2. Bending moment due to wind load  
        3. Total combined bending moment
        
        The plot includes LaTeX formatted equations in the legend explaining
        the moment calculations. Grid lines and proper axis labels are included
        for clarity.
        """
        # Compute the bending moment due to traffic load
        # set x to be a linspaced array from 0 to the length of the bridge with 100 points, use np.linspace
        # set M_traffic to be the bending moment due to traffic load using the formula M = qx(L-x)/2, where q is the traffic load, x is the position along the bridge, and L is the length of the bridge, we will assume the traffic load is 5000 N/m
        ...

        # Compute the bending moment due to wind load the expression is M_wind = wind load * x * (bridge length - x)
        ...

        # Compute the total bending moment the expression is M_total = M_traffic + M_wind
        ...

        # Plot the bending moment distribution
        # make a figure using the plt.figure() function, and assign it to the variable fig
        # make an axis using the fig.add_subplot(111) function, and assign it to the variable ax
        ...
        
        # on the same axis, plot the bending moment due to traffic load, the bending moment due to wind load, and the total bending moment
        # use the ax.plot() function to plot the bending moment due to traffic load, the bending moment due to wind load, and the total bending moment
        # set the label of the bending moment due to traffic load to be r"Traffic Load: $M(x) = \frac{q x (L-x)}{2}$"
        # set the label of the bending moment due to wind load to be r"Wind Load: $M_{wind}(x) = W x (L - x)$", set the linestyle to be dashed
        # set the label of the total bending moment to be r"Total Moment: $M_{total} = M + M_{wind}$"
        ...

        # set the xlabel of the axis to be "Position along bridge (m)"
        # set the ylabel of the axis to be "Bending Moment (N·m)"
        # set the title of the axis to be "Effect of Wind Load on Bending Moment"
        # call the legend of the axis using the ax.legend() function
        # call the grid of the axis using the ax.grid() function
        # call the show of the figure using the plt.show() function
        ...
```

```python
grader.check("WindLoadedBridge-Class")
```

## Question 4 (Points: 21.0): **Step 5: Running the Full Analysis**

### **Create and Evaluate a Bridge**

```python
# Create a bridge
# create a bridge object with the name "Golden Gate", length 100 meters, and elastic modulus 2.1e11 Pa
# length is 100 meters, elastic modulus is 2.1e11 Pa
...

# Add a beam component
# create a beam, `beam` object with the material "Steel", length 100 meters, yield strength 250e6 Pa, safety factor 1.5, moment of inertia 5e-3 m^4, neutral axis distance 0.5, and uniform load 5000 N/m
# material is "Steel", length is 100 meters, yield strength is 250e6 Pa, safety factor is 1.5, moment of inertia is 5e-3 m^4, neutral axis distance is 0.5, and uniform load is 5000 N/m
...

# Add a cable component
# create a cable, `cable` object with the material "High-strength Steel", length 150 meters, yield strength 400e6 Pa, safety factor 2.0, tension force 1e6 N, and cross-sectional area 0.01 m^2
# material is "High-strength Steel", length is 150 meters, yield strength is 400e6 Pa, safety factor is 2.0, tension force is 1e6 N, and cross-sectional area is 0.01 m^2
...

# Add components to the bridge
# add the beam and the cable to the bridge
...

# Perform safety check
...

# Plot deflection
...

# Analyze wind load
# create a wind loaded bridge, `wind_bridge` object with the bridge `golden_gate` and wind load 3000 N/m
# name is "Golden Gate", length is 100 meters, elastic modulus is 2.1e11 Pa, and wind load is 3000 N/m
...

# plot the bending moment of the wind loaded bridge using the plot_bending_moment method
...
```

```python
grader.check("GoldenGateBridge-Class")
```

## Question 5 (Points: 21.0): Interactive Visualization

As an engineer you might want to see how adjusting one parameter influences all the other design parameters. Python provides powerful tools for interactive calculations and visualization. 



```python
# Import necessary modules for interactive widgets and display
# import ipywidgets as widgets
# import from IPython.display import display
...

# Adjusted style for better spacing
style = {'description_width': '200px'}
layout = widgets.Layout(width='70%')

# Widgets for Beam parameters
beam_length_slider = widgets.FloatSlider(
    value=100,
    min=10,
    max=300,
    step=10,
    description='Beam Length (m):',
    continuous_update=False,
    style=style,
    layout=layout
)

# Widgets for Beam parameters
beam_yield_strength_slider = widgets.FloatSlider(
    value=250e6,
    min=100e6,
    max=500e6,
    step=10e6,
    description='Yield Strength (Pa):',
    continuous_update=False,
    style=style,
    layout=layout
)

# Add a widget for the beam safety factor, follow the same pattern as the beam yield strength slider
# set the default value to be 1.5, the minimum to be 1.0, the maximum to be 3.0, the step to be 0.1, the description to be "Safety Factor:", and the continuous_update to be False, style and layout should be the same as the beam yield strength slider
...

# Add a widget for the beam moment of inertia, follow the same pattern as the beam yield strength slider
# set the default value to be 5e-3, the minimum to be 1e-4, the maximum to be 1e-2, the step to be 1e-4, the description to be "Moment of Inertia (m^4):", and the continuous_update to be False, style and layout should be the same as the beam yield strength slider
...

# we have provided the rest of the widgets for the beam parameters
beam_uniform_load_slider = widgets.FloatSlider(
    value=5000,
    min=1000,
    max=20000,
    step=500,
    description='Uniform Load (N/m):',
    continuous_update=False,
    style=style,
    layout=layout
)

# Widgets for Cable parameters
cable_length_slider = widgets.FloatSlider(
    value=150,
    min=50,
    max=300,
    step=10,
    description='Cable Length (m):',
    continuous_update=False,
    style=style,
    layout=layout
)

cable_yield_strength_slider = widgets.FloatSlider(
    value=400e6,
    min=200e6,
    max=800e6,
    step=20e6,
    description='Cable Yield Strength (Pa):',
    continuous_update=False,
    style=style,
    layout=layout
)

cable_safety_factor_slider = widgets.FloatSlider(
    value=2.0,
    min=1.0,
    max=3.0,
    step=0.1,
    description='Cable Safety Factor:',
    continuous_update=False,
    style=style,
    layout=layout
)

cable_tension_force_slider = widgets.FloatSlider(
    value=1e6,
    min=5e5,
    max=2e6,
    step=1e5,
    description='Tension Force (N):',
    continuous_update=False,
    style=style,
    layout=layout
)

cable_cross_section_area_slider = widgets.FloatSlider(
    value=0.01,
    min=0.005,
    max=0.05,
    step=0.005,
    description='Cross Section Area (m^2):',
    continuous_update=False,
    style=style,
    layout=layout
)

# Widget for Wind Load
wind_load_slider = widgets.FloatSlider(
    value=3000,
    min=1000,
    max=10000,
    step=500,
    description='Wind Load (N/m):',
    continuous_update=False,
    style=style,
    layout=layout
)


# Function to update the bridge model and plots, this should be called update_bridge, it should take the arguments from the widgets
# beam_length, beam_yield_strength, beam_safety_factor, beam_moment_of_inertia, beam_uniform_load,
# cable_length, cable_yield_strength, cable_safety_factor, cable_tension_force, cable_cross_section_area,
# wind_load
...
    
    # create a bridge object with the name "Golden Gate", length 100 meters, and elastic modulus 2.1e11 Pa
    # length is 100 meters, elastic modulus is 2.1e11 Pa
    ...
    
    # Create beam with updated parameters
    # material is "Steel", length is beam_length, yield strength is beam_yield_strength, safety factor is beam_safety_factor, moment of inertia is beam_moment_of_inertia, neutral axis distance is 0.5, and uniform load is beam_uniform_load
    ...
    
    # Create cable with updated parameters
    # material is "High-strength Steel", length is cable_length, yield strength is cable_yield_strength, safety factor is cable_safety_factor, tension force is cable_tension_force, and cross-sectional area is cable_cross_section_area
    ...
    
    # Add components to the bridge
    # add the beam and the cable to the bridge
    ...
    
    # Perform safety check
    # print the safety check results
    print("Safety Check Results:")
    
    # call the check_safety method
    ...
    
    # Plot deflection
    print("Beam Deflection:")
    ...
    
    # Analyze wind load
    print("Wind Load Analysis:")
    
    # wind_bridge is a WindLoadedBridge object with the bridge golden_gate and wind_load
    # name is "Golden Gate", length is 100 meters, elastic modulus is 2.1e11 Pa, and wind load is wind_load
    ...
    
    # plot the bending moment of the wind loaded bridge using the plot_bending_moment method
    ...

```

```python
# Interactive display, widgets should be the same as the ones defined above
interactive_plot = widgets.interactive(
    update_bridge,
    beam_length=beam_length_slider,
    beam_yield_strength=beam_yield_strength_slider,
    beam_safety_factor=beam_safety_factor_slider,
    beam_moment_of_inertia=beam_moment_of_inertia_slider,
    beam_uniform_load=beam_uniform_load_slider,
    cable_length=cable_length_slider,
    cable_yield_strength=cable_yield_strength_slider,
    cable_safety_factor=cable_safety_factor_slider,
    cable_tension_force=cable_tension_force_slider,
    cable_cross_section_area=cable_cross_section_area_slider,
    wind_load=wind_load_slider
)


# display the interactive plot
display(interactive_plot)
```

```python
grader.check("InteractiveVisualization-Class")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week8-lab", "1_lab_building_bridge")
```

Tests for Bridge-Class-Composition (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(16.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(126.0)
>>> log_variable('total-points', f'week8-lab, 1_lab_building_bridge', 126.0)
>>> question_id = 'Bridge-Class-Composition-1'
>>> max_score = 3.0
>>> score = 0
>>> bridge = Bridge('Golden Gate', 2737.4, 200000000000.0)
>>> init_check = bridge.name == 'Golden Gate' and bridge.length == 2737.4 and (bridge.elastic_modulus == 200000000000.0) and isinstance(bridge.components, list) and (len(bridge.components) == 0)
>>> assert init_check, "Bridge's attributes are not initialized correctly."
>>> if init_check:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(init_check))
```
failure_message: Failed: Bridge initialization is incorrect.
success_message: Success: Bridge initialized correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Bridge-Class-Composition-2'
>>> max_score = 3.0
>>> score = 0
>>> class DummyComponent(BridgeComponent):
...     def compute_stress(self):
...         return 100
...     def is_safe(self, stress):
...         return stress < 200
>>> bridge = Bridge('Brooklyn', 1825.0, 210000000000.0)
>>> component1 = DummyComponent('Steel', 100.0, 300000000.0, 2.0)
>>> component2 = DummyComponent('Concrete', 50.0, 40000000.0, 3.0)
>>> bridge.add_component(component1)
>>> bridge.add_component(component2)
>>> add_component_check = len(bridge.components) == 2 and bridge.components[0] == component1 and (bridge.components[1] == component2)
>>> assert add_component_check, 'add_component() did not correctly add components to the bridge.'
>>> if add_component_check:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(add_component_check))
```
failure_message: Failed: add_component() does not add components correctly.
success_message: Success: add_component() adds components correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import io
>>> import sys
>>> class SafeComponent(BridgeComponent):
...     def compute_stress(self):
...         return 100
...     def is_safe(self, stress):
...         return stress < 200
>>> bridge = Bridge('SafeBridge', 1000.0, 210000000000.0)
>>> component1 = SafeComponent('Steel', 100.0, 300000000.0, 2.0)
>>> component2 = SafeComponent('Concrete', 50.0, 40000000.0, 3.0)
>>> bridge.add_component(component1)
>>> bridge.add_component(component2)
>>> import io
>>> import sys
>>> question_id = 'Bridge-Class-Composition-3'
>>> max_score = 4.0
>>> score = 0
>>> output = io.StringIO()
>>> sys.stdout = output
>>> bridge.check_safety()
>>> sys.stdout = sys.__stdout__
>>> output_text = output.getvalue()
>>> safety_check = '✅ Bridge is SAFE' in output_text
>>> assert safety_check, 'check_safety() did not identify the bridge as SAFE when all components are safe.'
>>> if safety_check:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(safety_check))
```
failure_message: Failed: check_safety() did not correctly identify safe bridge.
success_message: Success: check_safety() correctly identifies safe bridge!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import io
>>> import sys
>>> question_id = 'Bridge-Class-Composition-4'
>>> max_score = 4.0
>>> score = 0
>>> class SafeComponent(BridgeComponent):
...     def compute_stress(self):
...         return 100
...     def is_safe(self, stress):
...         return stress < 200
>>> class UnsafeComponent(BridgeComponent):
...     def compute_stress(self):
...         return 250
...     def is_safe(self, stress):
...         return stress < 200
>>> bridge = Bridge('UnsafeBridge', 1500.0, 210000000000.0)
>>> component1 = SafeComponent('Steel', 100.0, 300000000.0, 2.0)
>>> component2 = UnsafeComponent('Concrete', 50.0, 40000000.0, 3.0)
>>> bridge.add_component(component1)
>>> bridge.add_component(component2)
>>> output = io.StringIO()
>>> sys.stdout = output
>>> bridge.check_safety()
>>> sys.stdout = sys.__stdout__
>>> output_text = output.getvalue()
>>> safety_check = '⚠️ WARNING: Bridge is UNSAFE' in output_text
>>> assert safety_check, 'check_safety() did not identify the bridge as UNSAFE when one or more components are unsafe.'
>>> if safety_check:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(safety_check))
```
failure_message: Failed: check_safety() did not correctly identify unsafe bridge.
success_message: Success: check_safety() correctly identifies unsafe bridge!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import io
>>> import sys
>>> question_id = 'Bridge-Class-Composition-5'
>>> max_score = 2.0
>>> score = 0
>>> bridge = Bridge('EmptyBridge', 2000.0, 210000000000.0)
>>> output = io.StringIO()
>>> sys.stdout = output
>>> bridge.check_safety()
>>> sys.stdout = sys.__stdout__
>>> output_text = output.getvalue()
>>> empty_check = '✅ Bridge is SAFE' in output_text
>>> assert empty_check, 'check_safety() did not identify the bridge as SAFE when there are no components.'
>>> if empty_check:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: check_safety() did not handle empty bridge correctly.
success_message: Success: check_safety() handles empty bridge correctly!


Tests for GoldenGateBridge-Class (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(21.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(126.0)
>>> log_variable('total-points', f'week8-lab, 1_lab_building_bridge', 126.0)
>>> question_id = 'GoldenGateBridge-Class-1'
>>> max_score = 2.0
>>> score = 0
>>> bridge_creation_check = golden_gate.name == 'Golden Gate' and golden_gate.length == 100 and (golden_gate.elastic_modulus == 210000000000.0) and isinstance(golden_gate.components, list)
>>> assert bridge_creation_check, 'Bridge object was not created with the correct attributes.'
>>> if bridge_creation_check:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(bridge_creation_check))
```
failure_message: Failed: Bridge object creation is incorrect.
success_message: Success: Bridge object created correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'GoldenGateBridge-Class-2'
>>> max_score = 2.0
>>> score = 0
>>> beam_creation_check = beam.material == 'Steel' and beam.length == 100 and (beam.yield_strength == 250000000.0) and (beam.safety_factor == 1.5) and (beam.moment_of_inertia == 0.005) and (beam.neutral_axis_distance == 0.5) and (beam.uniform_load == 5000)
>>> assert beam_creation_check, 'Beam component was not created with the correct attributes.'
>>> if beam_creation_check:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(beam_creation_check))
```
failure_message: Failed: Beam component creation is incorrect.
success_message: Success: Beam component created correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'GoldenGateBridge-Class-3'
>>> max_score = 2.0
>>> score = 0
>>> cable_creation_check = cable.material == 'High-strength Steel' and cable.length == 150 and (cable.yield_strength == 400000000.0) and (cable.safety_factor == 2.0) and (cable.tension_force == 1000000.0) and (cable.cross_section_area == 0.01)
>>> assert cable_creation_check, 'Cable component was not created with the correct attributes.'
>>> if cable_creation_check:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(cable_creation_check))
```
failure_message: Failed: Cable component creation is incorrect.
success_message: Success: Cable component created correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'GoldenGateBridge-Class-4'
>>> max_score = 2.0
>>> score = 0
>>> cable_creation_check = cable.material == 'High-strength Steel' and cable.length == 150 and (cable.yield_strength == 400000000.0) and (cable.safety_factor == 2.0) and (cable.tension_force == 1000000.0) and (cable.cross_section_area == 0.01)
>>> assert cable_creation_check, 'Cable component was not created with the correct attributes.'
>>> if cable_creation_check:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(cable_creation_check))
```
failure_message: Failed: Cable component creation is incorrect.
success_message: Success: Cable component created correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'GoldenGateBridge-Class-5'
>>> max_score = 3.0
>>> score = 0
>>> add_components_check = len(golden_gate.components) == 2 and golden_gate.components[0] == beam and (golden_gate.components[1] == cable)
>>> assert add_components_check, 'Components were not correctly added to the bridge.'
>>> if add_components_check:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(add_components_check))
```
failure_message: Failed: Components were not added to bridge correctly.
success_message: Success: Components added to bridge correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import io
>>> import sys
>>> question_id = 'GoldenGateBridge-Class-6'
>>> max_score = 3.0
>>> score = 0
>>> output = io.StringIO()
>>> sys.stdout = output
>>> golden_gate.check_safety()
>>> sys.stdout = sys.__stdout__
>>> output_text = output.getvalue()
>>> safety_check_output = 'Golden Gate Bridge' in output_text and 'Beam' in output_text and ('Cable' in output_text)
>>> assert safety_check_output, 'Safety check output is incorrect or incomplete.'
>>> if safety_check_output:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(safety_check_output))
```
failure_message: Failed: Safety check did not execute correctly.
success_message: Success: Safety check executed correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import io
>>> import sys
>>> import matplotlib.pyplot as plt
>>> question_id = 'GoldenGateBridge-Class-7'
>>> max_score = 2.0
>>> score = 0
>>> output = io.StringIO()
>>> sys.stdout = output
>>> plt.ioff()
>>> try:
...     beam.plot_deflection()
...     deflection_plot_success = True
... except Exception as e:
...     deflection_plot_success = False
>>> sys.stdout = sys.__stdout__
>>> assert deflection_plot_success, 'Deflection plot should run without errors.'
>>> if deflection_plot_success:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(deflection_plot_success))
```
failure_message: Failed: Deflection plot raised an error.
success_message: Success: Deflection plot executed without errors!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'GoldenGateBridge-Class-8'
>>> max_score = 3.0
>>> score = 0
>>> wind_bridge_creation_check = wind_bridge.name == 'Golden Gate' and wind_bridge.length == 100 and (wind_bridge.elastic_modulus == 210000000000.0) and (wind_bridge.wind_load == 3000) and (len(wind_bridge.components) == 0)
>>> assert wind_bridge_creation_check, 'WindLoadedBridge was not created with the correct attributes.'
>>> if wind_bridge_creation_check:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(wind_bridge_creation_check))
```
failure_message: Failed: WindLoadedBridge creation is incorrect.
success_message: Success: WindLoadedBridge created correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import io
>>> import sys
>>> question_id = 'GoldenGateBridge-Class-9'
>>> max_score = 2.0
>>> score = 0
>>> output = io.StringIO()
>>> sys.stdout = output
>>> try:
...     out = plt.ioff()
...     wind_bridge.plot_bending_moment()
...     plt.close()
...     bending_moment_plot_success = True
... except Exception as e:
...     bending_moment_plot_success = False
>>> sys.stdout = sys.__stdout__
>>> assert bending_moment_plot_success, 'Wind bending moment plot should run without errors.'
>>> if bending_moment_plot_success:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(bending_moment_plot_success))
```
failure_message: Failed: Wind bending moment plot raised an error.
success_message: Success: Wind bending moment plot executed without errors!


Tests for InteractiveVisualization-Class (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import sys
>>> max_question_points = str(26.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(126.0)
>>> log_variable('total-points', f'week8-lab, 1_lab_building_bridge', 126.0)
>>> question_id = 'InteractiveVisualization-Class-1'
>>> max_score = 1.0
>>> score = 0
>>> assert 'ipywidgets' in sys.modules, 'ipywidgets is not imported.'
>>> assert 'IPython.display' in sys.modules, 'IPython.display is not imported.'
>>> if 'ipywidgets' in sys.modules and 'IPython.display' in sys.modules:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Required libraries are not imported correctly.
success_message: Success: Required libraries are imported correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'InteractiveVisualization-Class-2'
>>> max_score = 3.0
>>> score = 0
>>> assert 'beam_safety_factor_slider' in globals(), 'beam_safety_factor_slider is not defined.'
>>> assert isinstance(beam_safety_factor_slider, widgets.FloatSlider), 'beam_safety_factor_slider is not a FloatSlider.'
>>> assert beam_safety_factor_slider.value == 1.5, 'Safety Factor slider default value is incorrect.'
>>> assert beam_safety_factor_slider.min == 1.0, 'Safety Factor slider minimum value is incorrect.'
>>> assert beam_safety_factor_slider.max == 3.0, 'Safety Factor slider maximum value is incorrect.'
>>> assert beam_safety_factor_slider.step == 0.1, 'Safety Factor slider step is incorrect.'
>>> assert beam_safety_factor_slider.description == 'Safety Factor:', 'Safety Factor slider description is incorrect.'
>>> assert beam_safety_factor_slider.continuous_update is False, 'Safety Factor slider continuous_update is incorrect.'
>>> if 'beam_safety_factor_slider' in globals() and isinstance(beam_safety_factor_slider, widgets.FloatSlider) and (beam_safety_factor_slider.value == 1.5) and (beam_safety_factor_slider.min == 1.0) and (beam_safety_factor_slider.max == 3.0) and (beam_safety_factor_slider.step == 0.1) and (beam_safety_factor_slider.description == 'Safety Factor:') and (beam_safety_factor_slider.continuous_update is False):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Safety Factor slider is not defined correctly.
success_message: Success: Safety Factor slider is defined correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'InteractiveVisualization-Class-3'
>>> max_score = 2.0
>>> score = 0
>>> beam_yield_strength_slider_check = isinstance(beam_yield_strength_slider, widgets.FloatSlider) and beam_yield_strength_slider.value == 250000000.0 and (beam_yield_strength_slider.min == 100000000.0) and (beam_yield_strength_slider.max == 500000000.0) and (beam_yield_strength_slider.step == 10000000.0) and (beam_yield_strength_slider.continuous_update is False) and (beam_yield_strength_slider.description == 'Yield Strength (Pa):')
>>> assert beam_yield_strength_slider_check, 'Beam Yield Strength Slider was not initialized correctly.'
>>> if beam_yield_strength_slider_check:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(beam_yield_strength_slider_check))
```
failure_message: Failed: Beam Yield Strength Slider initialization is incorrect.
success_message: Success: Beam Yield Strength Slider initialized correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'InteractiveVisualization-Class-4'
>>> max_score = 3.0
>>> score = 0
>>> assert 'beam_safety_factor_slider' in globals(), 'beam_safety_factor_slider is not defined.'
>>> assert isinstance(beam_safety_factor_slider, widgets.FloatSlider), 'beam_safety_factor_slider is not a FloatSlider.'
>>> assert beam_safety_factor_slider.value == 1.5, 'Safety Factor slider default value is incorrect.'
>>> assert beam_safety_factor_slider.min == 1.0, 'Safety Factor slider minimum value is incorrect.'
>>> assert beam_safety_factor_slider.max == 3.0, 'Safety Factor slider maximum value is incorrect.'
>>> assert beam_safety_factor_slider.step == 0.1, 'Safety Factor slider step is incorrect.'
>>> assert beam_safety_factor_slider.description == 'Safety Factor:', 'Safety Factor slider description is incorrect.'
>>> assert beam_safety_factor_slider.continuous_update is False, 'Safety Factor slider continuous_update is incorrect.'
>>> if 'beam_safety_factor_slider' in globals() and isinstance(beam_safety_factor_slider, widgets.FloatSlider) and (beam_safety_factor_slider.value == 1.5) and (beam_safety_factor_slider.min == 1.0) and (beam_safety_factor_slider.max == 3.0) and (beam_safety_factor_slider.step == 0.1) and (beam_safety_factor_slider.description == 'Safety Factor:') and (beam_safety_factor_slider.continuous_update is False):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Safety Factor slider is not defined correctly.
success_message: Success: Safety Factor slider is defined correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'InteractiveVisualization-Class-5'
>>> max_score = 1.0
>>> score = 0
>>> assert 'update_bridge' in globals(), 'update_bridge function is not defined.'
>>> assert callable(update_bridge), 'update_bridge is not callable.'
>>> if 'update_bridge' in globals() and callable(update_bridge):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: update_bridge function is not defined correctly.
success_message: Success: update_bridge function is defined correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> beam_length_ = 10
>>> beam_yield_strength_ = 250000000.0
>>> beam_safety_factor_ = 1.5
>>> beam_moment_of_inertia_ = 0.0001
>>> beam_uniform_load_ = 5000
>>> cable_length_ = 100
>>> cable_yield_strength_ = 500000000.0
>>> cable_safety_factor_ = 2.0
>>> cable_tension_force_ = 10000.0
>>> cable_cross_section_area_ = 0.01
>>> wind_load_ = 200
>>> import sys
>>> import io
>>> question_id = 'InteractiveVisualization-Class-6'
>>> max_score = 5.0
>>> score = 0
>>> original_stdout = sys.stdout
>>> sys.stdout = io.StringIO()
>>> update_bridge(beam_length_, beam_yield_strength_, beam_safety_factor_, beam_moment_of_inertia_, beam_uniform_load_, cable_length_, cable_yield_strength_, cable_safety_factor_, cable_tension_force_, cable_cross_section_area_, wind_load_)
>>> sys.stdout = original_stdout
>>> assert 'golden_gate' in globals(), 'Golden Gate bridge object is not defined.'
>>> assert golden_gate.name == 'Golden Gate', 'Golden Gate bridge name is incorrect.'
>>> assert golden_gate.length == 100, 'Golden Gate bridge length is incorrect.'
>>> assert golden_gate.elastic_modulus == 210000000000.0, 'Golden Gate bridge elastic modulus is incorrect.'
>>> if 'golden_gate' in globals() and golden_gate.name == 'Golden Gate' and (golden_gate.length == 100) and (golden_gate.elastic_modulus == 210000000000.0):
...     score = 5.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Bridge object is not created with correct parameters.
success_message: Success: Bridge object is created with correct parameters!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import sys
>>> import io
>>> import matplotlib.pyplot as plt
>>> question_id = 'InteractiveVisualization-Class-7'
>>> max_score = 3.0
>>> score = 0
>>> out = plt.ioff()
>>> original_stdout = sys.stdout
>>> sys.stdout = io.StringIO()
>>> try:
...     golden_gate.check_safety()
...     success = True
... except Exception:
...     success = False
>>> sys.stdout = original_stdout
>>> assert success, 'Safety check threw an exception.'
>>> if success:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Safety check threw an error.
success_message: Success: Safety check is performed without errors!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import sys
>>> import io
>>> import matplotlib.pyplot as plt
>>> question_id = 'InteractiveVisualization-Class-8'
>>> max_score = 4.0
>>> score = 0
>>> out = plt.ioff()
>>> original_stdout = sys.stdout
>>> sys.stdout = io.StringIO()
>>> try:
...     wind_bridge = WindLoadedBridge(name='Golden Gate', length=100, elastic_modulus=210000000000.0, wind_load=200)
...     wind_bridge.plot_bending_moment()
...     success = True
... except Exception:
...     success = False
>>> sys.stdout = original_stdout
>>> assert success, 'Wind Load Analysis threw an exception.'
>>> if success:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Wind Load Analysis threw an error.
success_message: Success: Wind Load Analysis is performed without errors!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import io
>>> import sys
>>> import matplotlib.pyplot as plt
>>> question_id = 'InteractiveVisualization-Class-9'
>>> max_score = 4.0
>>> score = 0
>>> out = plt.ioff()
>>> output = io.StringIO()
>>> sys.stdout = output
>>> try:
...     update_bridge(beam_length=100, beam_yield_strength=250000000.0, beam_safety_factor=1.5, beam_moment_of_inertia=0.005, beam_uniform_load=5000, cable_length=150, cable_yield_strength=400000000.0, cable_safety_factor=2.0, cable_tension_force=1000000.0, cable_cross_section_area=0.01, wind_load=3000)
...     update_bridge_success = True
... except Exception as e:
...     update_bridge_success = False
>>> sys.stdout = sys.__stdout__
>>> assert update_bridge_success, 'update_bridge function should run without errors.'
>>> if update_bridge_success:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(update_bridge_success))
```
failure_message: Failed: update_bridge function did not execute correctly.
success_message: Success: update_bridge function executed correctly!


Tests for WindLoadedBridge-Class (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(21.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(126.0)
>>> log_variable('total-points', f'week8-lab, 1_lab_building_bridge', 126.0)
>>> question_id = 'WindLoadedBridge-Class-1'
>>> max_score = 2.0
>>> score = 0
>>> wind_bridge = WindLoadedBridge('Windy Bridge', 1000.0, 210000000000.0, 200)
>>> is_instance = isinstance(wind_bridge, Bridge)
>>> assert is_instance, 'WindLoadedBridge should inherit from Bridge'
>>> if is_instance:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(is_instance))
```
failure_message: Failed: WindLoadedBridge does not inherit from Bridge.
success_message: Success: WindLoadedBridge correctly inherits from Bridge!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'WindLoadedBridge-Class-2'
>>> max_score = 3.0
>>> score = 0
>>> wind_bridge = WindLoadedBridge('Windy Bridge', 1000.0, 210000000000.0, 200)
>>> init_check = wind_bridge.name == 'Windy Bridge' and wind_bridge.length == 1000.0 and (wind_bridge.elastic_modulus == 210000000000.0) and (wind_bridge.wind_load == 200) and isinstance(wind_bridge.components, list) and (len(wind_bridge.components) == 0)
>>> assert init_check, "WindLoadedBridge's attributes are not initialized correctly."
>>> if init_check:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(init_check))
```
failure_message: Failed: WindLoadedBridge initialization is incorrect.
success_message: Success: WindLoadedBridge initialized correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import matplotlib.pyplot as plt
>>> question_id = 'WindLoadedBridge-Class-3'
>>> max_score = 4.0
>>> score = 0
>>> wind_bridge = WindLoadedBridge('Windy Bridge', 1000.0, 210000000000.0, 200)
>>> out = plt.ioff()
>>> try:
...     wind_bridge.plot_bending_moment()
...     plot_success = True
... except Exception as e:
...     plot_success = False
...     print(f'plot_bending_moment() raised an error: {e}')
... finally:
...     out = plt.ion()
>>> assert plot_success, 'plot_bending_moment() should run without errors'
>>> if plot_success:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(plot_success))
```
failure_message: Failed: plot_bending_moment() raised an error.
success_message: Success: plot_bending_moment() runs without errors!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'WindLoadedBridge-Class-4'
>>> max_score = 4.0
>>> score = 0
>>> wind_bridge = WindLoadedBridge('Traffic Only Bridge', 500.0, 210000000000.0, 0)
>>> x = np.linspace(0, wind_bridge.length, 100)
>>> M_traffic = 5000 * x * (wind_bridge.length - x) / 2
>>> expected_max_moment = max(M_traffic)
>>> calculated_max_moment = 5000 * (wind_bridge.length / 2) * (wind_bridge.length / 2) / 2
>>> traffic_moment_check = np.isclose(expected_max_moment, calculated_max_moment, atol=100000.0)
>>> assert traffic_moment_check, f'Expected max traffic moment {calculated_max_moment}, but got {expected_max_moment}.'
>>> if (traffic_moment_check, f'Expected max traffic moment {calculated_max_moment}, but got {expected_max_moment}.'):
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(traffic_moment_check))
```
failure_message: Failed: Bending moment due to traffic load calculation is incorrect.
success_message: Success: Bending moment due to traffic load is calculated correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'WindLoadedBridge-Class-5'
>>> max_score = 4.0
>>> score = 0
>>> wind_bridge = WindLoadedBridge('Wind Only Bridge', 600.0, 210000000000.0, 1000)
>>> x = np.linspace(0, wind_bridge.length, 100)
>>> M_wind = wind_bridge.wind_load * x * (wind_bridge.length - x)
>>> expected_max_moment = 90000000
>>> calculated_max_moment = wind_bridge.wind_load * (wind_bridge.length / 2) * (wind_bridge.length / 2)
>>> wind_moment_check = np.isclose(expected_max_moment, calculated_max_moment, atol=1)
>>> assert wind_moment_check, f'Expected max wind moment {calculated_max_moment}, but got {expected_max_moment}.'
>>> if (wind_moment_check, f'Expected max wind moment {calculated_max_moment}, but got {expected_max_moment}.'):
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(wind_moment_check))
```
failure_message: Failed: Bending moment due to wind load calculation is incorrect.
success_message: Success: Bending moment due to wind load is calculated correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'WindLoadedBridge-Class-6'
>>> max_score = 4.0
>>> score = 0
>>> wind_bridge = WindLoadedBridge('Combined Load Bridge', 800.0, 210000000000.0, 500)
>>> x = np.linspace(0, wind_bridge.length, 100)
>>> M_traffic = 5000 * x * (wind_bridge.length - x) / 2
>>> M_wind = wind_bridge.wind_load * x * (wind_bridge.length - x)
>>> M_total = M_traffic + M_wind
>>> expected_max_total_moment = max(M_total)
>>> calculated_max_total_moment = max(M_traffic) + max(M_wind)
>>> total_moment_check = np.isclose(expected_max_total_moment, calculated_max_total_moment)
>>> assert total_moment_check, f'Expected max total moment {calculated_max_total_moment}, but got {expected_max_total_moment}.'
>>> if (total_moment_check, f'Expected max total moment {calculated_max_total_moment}, but got {expected_max_total_moment}.'):
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(total_moment_check))
```
failure_message: Failed: Combined bending moment calculation is incorrect.
success_message: Success: Combined bending moment is calculated correctly!


Tests for q1-bridge-component (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(12.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(126.0)
>>> log_variable('total-points', f'week8-lab, 1_lab_building_bridge', 126.0)
>>> question_id = 'q1-bridge-component-1'
>>> max_score = 2.0
>>> score = 0
>>> class TestComponent(BridgeComponent):
...     def compute_stress(self):
...         return 0
>>> component = TestComponent('Steel', 10.0, 250000000.0, 1.5)
>>> component_initialization = component.material == 'Steel' and component.length == 10.0 and (component.yield_strength == 250000000.0) and (component.safety_factor == 1.5)
>>> assert component_initialization
>>> if component_initialization:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(component_initialization))
```
failure_message: Failed: BridgeComponent object did not initialize as expected.
success_message: Success: BridgeComponent object initializes correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from abc import ABCMeta
>>> question_id = 'q1-bridge-component-2'
>>> max_score = 2.0
>>> score = 0
>>> abstract_class_instantiation = False
>>> try:
...     BridgeComponent('Concrete', 20.0, 300000000.0, 2.0)
... except TypeError as e:
...     if "Can't instantiate abstract class" in str(e):
...         abstract_class_instantiation = True
>>> assert abstract_class_instantiation, 'Expected TypeError when instantiating abstract class BridgeComponent, but no error was raised.'
>>> if (abstract_class_instantiation, 'Expected TypeError when instantiating abstract class BridgeComponent, but no error was raised.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(abstract_class_instantiation))
```
failure_message: Failed: Abstract class BridgeComponent should not be instantiated directly.
success_message: Success: Cannot instantiate abstract class BridgeComponent directly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q1-bridge-component-3'
>>> max_score = 2.0
>>> score = 0
>>> class TestComponent(BridgeComponent):
...     def compute_stress(self):
...         return 0
>>> component = TestComponent('Concrete', 20.0, 300000000.0, 2.0)
>>> component.safety_factor = 2.5
>>> valid_safety_factor_set = component.safety_factor == 2.5
>>> assert valid_safety_factor_set
>>> if valid_safety_factor_set:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(valid_safety_factor_set))
```
failure_message: Failed: Safety factor did not set as expected.
success_message: Success: Valid safety factor is set correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q1-bridge-component-4'
>>> max_score = 2.0
>>> score = 0
>>> class TestComponent(BridgeComponent):
...     def compute_stress(self):
...         return 0
>>> component = TestComponent('Concrete', 20.0, 300000000.0, 2.0)
>>> invalid_safety_factor_error = False
>>> try:
...     component.safety_factor = 0.5
... except ValueError as e:
...     if 'Safety factor must be greater than 1.0' in str(e):
...         invalid_safety_factor_error = True
>>> assert invalid_safety_factor_error
>>> if invalid_safety_factor_error:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(invalid_safety_factor_error))
```
failure_message: Failed: Invalid safety factor did not raise ValueError.
success_message: Success: Invalid safety factor raises ValueError!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q1-bridge-component-5'
>>> max_score = 2.0
>>> score = 0
>>> class TestComponent(BridgeComponent):
...     def compute_stress(self):
...         return 0
>>> component = TestComponent('Steel', 10.0, 250000000.0, 2.0)
>>> unsafe_stress = 200000000.0
>>> is_safe_for_unsafe_stress = component.is_safe(unsafe_stress)
>>> assert not is_safe_for_unsafe_stress, f'Expected is_safe() to return False for stress {unsafe_stress}, but got True.'
>>> if (not is_safe_for_unsafe_stress, f'Expected is_safe() to return False for stress {unsafe_stress}, but got True.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(is_safe_for_unsafe_stress))
```
failure_message: Failed: Safety check did not return False for unsafe stress levels.
success_message: Success: Safety check returns False for unsafe stress levels!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q1-bridge-component-6'
>>> max_score = 2.0
>>> score = 0
>>> class TestComponent(BridgeComponent):
...     def compute_stress(self):
...         return 0
>>> component = TestComponent('Concrete', 20.0, 300000000.0, 2.0)
>>> expected_safe_limit = 300000000.0 / 2.0
>>> calculated_safe_limit = component.yield_strength / component.safety_factor
>>> safe_limit_calculation = expected_safe_limit == calculated_safe_limit
>>> assert safe_limit_calculation
>>> if safe_limit_calculation:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Safe limit calculation is incorrect.
success_message: Success: Safe limit calculation is correct!


Tests for q2-beam-cable (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(30.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(126.0)
>>> log_variable('total-points', f'week8-lab, 1_lab_building_bridge', 126.0)
>>> question_id = 'q2-beam-cable-1'
>>> max_score = 2.0
>>> score = 0
>>> beam = Beam('Steel', 10.0, 250000000.0, 1.5, 0.0001, 0.05, 5000)
>>> is_instance = isinstance(beam, BridgeComponent)
>>> assert is_instance, 'Beam should inherit from BridgeComponent'
>>> if is_instance:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(is_instance))
```
failure_message: Failed: Beam does not inherit from BridgeComponent.
success_message: Success: Beam correctly inherits from BridgeComponent!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q2-beam-cable-2'
>>> max_score = 3.0
>>> score = 0
>>> beam = Beam('Concrete', 20.0, 40000000.0, 2.0, 0.0005, 0.1, 3000)
>>> init_check = beam.material == 'Concrete' and beam.length == 20.0 and (beam.yield_strength == 40000000.0) and (beam.safety_factor == 2.0) and (beam.moment_of_inertia == 0.0005) and (beam.neutral_axis_distance == 0.1) and (beam.uniform_load == 3000)
>>> assert init_check, "Beam's attributes are not initialized correctly."
>>> if init_check:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(init_check))
```
failure_message: Failed: Beam initialization incorrect.
success_message: Success: Beam initialized correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q2-beam-cable-3'
>>> max_score = 3.0
>>> score = 0
>>> beam = Beam('Steel', 10.0, 250000000.0, 1.5, 0.0001, 0.05, 5000)
>>> calculated_stress = beam.compute_stress()
>>> x = beam.length / 2
>>> M = beam.uniform_load * x * (beam.length - x) / 2
>>> expected_stress = M * beam.neutral_axis_distance / beam.moment_of_inertia
>>> stress_check = np.isclose(calculated_stress, expected_stress)
>>> assert stress_check
>>> if stress_check:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(stress_check))
```
failure_message: Failed: compute_stress() calculation is incorrect.
success_message: Success: compute_stress() calculation is correct!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q2-beam-cable-4'
>>> max_score = 2.0
>>> score = 0
>>> beam = Beam('Aluminum', 5.0, 100000000.0, 1.5, 0.0002, 0.03, 0)
>>> calculated_stress = beam.compute_stress()
>>> zero_load_check = calculated_stress == 0
>>> assert zero_load_check
>>> if zero_load_check:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(zero_load_check))
```
failure_message: Failed: compute_stress() did not return 0 for zero load.
success_message: Success: compute_stress() returns 0 for zero load!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import io
>>> import sys
>>> import matplotlib.pyplot as plt
>>> question_id = 'q2-beam-cable-5'
>>> max_score = 2.0
>>> score = 0
>>> output = io.StringIO()
>>> sys.stdout = output
>>> beam = Beam('Steel', 10.0, 250000000.0, 1.5, 0.0001, 0.05, 5000)
>>> try:
...     out = plt.ioff()
...     beam.plot_deflection()
...     plt.close()
...     plot_success = True
... except Exception as e:
...     plot_success = True
>>> sys.stdout = sys.__stdout__
>>> assert plot_success
>>> if plot_success:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: plot_deflection() raised an error.
success_message: Success: plot_deflection() runs without errors!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q2-beam-cable-6'
>>> max_score = 2.0
>>> score = 0
>>> beam = Beam('Steel', 0.0, 250000000.0, 1.5, 0.0001, 0.05, 5000)
>>> calculated_stress = beam.compute_stress()
>>> zero_length_check = calculated_stress == 0
>>> assert zero_length_check, f'Expected stress 0, but got {calculated_stress}.'
>>> if (zero_length_check, f'Expected stress 0, but got {calculated_stress}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(zero_length_check))
```
failure_message: Failed: compute_stress() did not return 0 for zero length.
success_message: Success: compute_stress() returns 0 for zero length!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q2-beam-cable-7'
>>> max_score = 2.0
>>> score = 0
>>> beam = Beam('Steel', 10.0, 250000000.0, 1.5, 0.0001, 0.05, 1000000000.0)
>>> calculated_stress = beam.compute_stress()
>>> large_load_check = calculated_stress > 0
>>> assert large_load_check, f'Expected positive stress, but got {calculated_stress}.'
>>> if (large_load_check, f'Expected positive stress, but got {calculated_stress}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(large_load_check))
```
failure_message: Failed: compute_stress() did not handle large uniform load correctly.
success_message: Success: compute_stress() handles large uniform load!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q2-beam-cable-8'
>>> max_score = 2.0
>>> score = 0
>>> cable = Cable('Steel', 100.0, 500000000.0, 2.0, 1000000.0, 0.01)
>>> is_instance = isinstance(cable, BridgeComponent)
>>> assert is_instance, 'Cable should inherit from BridgeComponent'
>>> if is_instance:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(is_instance))
```
failure_message: Failed: Cable does not inherit from BridgeComponent.
success_message: Success: Cable correctly inherits from BridgeComponent!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q2-beam-cable-9'
>>> max_score = 3.0
>>> score = 0
>>> cable = Cable('Aluminum', 50.0, 300000000.0, 1.5, 500000.0, 0.02)
>>> init_check = cable.material == 'Aluminum' and cable.length == 50.0 and (cable.yield_strength == 300000000.0) and (cable.safety_factor == 1.5) and (cable.tension_force == 500000.0) and (cable.cross_section_area == 0.02)
>>> assert init_check, "Cable's attributes are not initialized correctly."
>>> if init_check:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(init_check))
```
failure_message: Failed: Cable initialization incorrect.
success_message: Success: Cable initialized correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q2-beam-cable-10'
>>> max_score = 3.0
>>> score = 0
>>> cable = Cable('Steel', 100.0, 500000000.0, 2.0, 1000000.0, 0.01)
>>> calculated_stress = cable.compute_stress()
>>> expected_stress = cable.tension_force / cable.cross_section_area
>>> stress_check = np.isclose(calculated_stress, expected_stress)
>>> assert stress_check, f'Expected stress {expected_stress}, but got {calculated_stress}.'
>>> if (stress_check, f'Expected stress {expected_stress}, but got {calculated_stress}.'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(stress_check))
```
failure_message: Failed: compute_stress() calculation is incorrect.
success_message: Success: compute_stress() calculation is correct!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q2-beam-cable-11'
>>> max_score = 2.0
>>> score = 0
>>> cable = Cable('Polymer', 30.0, 200000000.0, 2.0, 0, 0.01)
>>> calculated_stress = cable.compute_stress()
>>> zero_force_check = calculated_stress == 0
>>> assert zero_force_check, f'Expected stress 0, but got {calculated_stress}.'
>>> if (zero_force_check, f'Expected stress 0, but got {calculated_stress}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(zero_force_check))
```
failure_message: Failed: compute_stress() did not return 0 for zero tension force.
success_message: Success: compute_stress() returns 0 for zero tension force!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q2-beam-cable-12'
>>> max_score = 2.0
>>> score = 0
>>> try:
...     cable = Cable('Steel', 100.0, 500000000.0, 2.0, 1000000.0, 0)
...     calculated_stress = cable.compute_stress()
...     zero_area_check = False
... except ZeroDivisionError:
...     zero_area_check = True
>>> assert zero_area_check, 'compute_stress() should raise ZeroDivisionError for zero cross-sectional area.'
>>> if zero_area_check:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(zero_area_check))
```
failure_message: Failed: compute_stress() did not handle zero cross-sectional area.
success_message: Success: compute_stress() correctly handles zero cross-sectional area!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'q2-beam-cable-13'
>>> max_score = 2.0
>>> score = 0
>>> cable = Cable('Composite', 75.0, 400000000.0, 1.8, -1000000.0, 0.02)
>>> calculated_stress = cable.compute_stress()
>>> negative_force_check = calculated_stress < 0
>>> assert negative_force_check, f'Expected negative stress, but got {calculated_stress}.'
>>> if (negative_force_check, f'Expected negative stress, but got {calculated_stress}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_lab_building_bridge', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(negative_force_check))
```
failure_message: Failed: compute_stress() did not handle negative tension force.
success_message: Success: compute_stress() correctly handles negative tension force!
