```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_practicequiz_q", "week_4", "practicequiz", assignment_points = 15.0, assignment_tag = 'week4-practicequiz')

# Initialize Otter
import otter
grader = otter.Notebook("1_practicequiz_q.ipynb")
```

 # ❓ Practice Quiz: Control Structures for Renewable Energy Applications

![](./assets/figures/wind-turbine-gif-quiz.gif)

 Test your understanding of conditional statements (`if-else`) and loops in Python by solving problems relevant to renewable energy systems, environmental monitoring, and automation tasks.

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_practicequiz_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_practicequiz_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_practicequiz_q import Question3
Question3().show()

```

 ## Renewable Energy Monitoring System



 Renewable energy systems, such as solar panels and wind turbines, generate data that needs to be processed to ensure efficiency and safety. In this task, you will implement a program to monitor energy production and analyze performance.



 The program should:



 1. Iterate through a list of energy production measurements (in kilowatt-hours, kWh).

 2. Print a warning message if the energy production drops below a threshold of 50 kWh.

 3. Calculate the average energy production and display it.

 4. Determine if the average energy production meets the optimal threshold (≥ 100 kWh) and display a summary message.



 ### Implementation Details



 1. Define a list of energy production measurements: `energy_readings = [120, 95, 85, 45, 150, 90]`.

 2. Set the optimal energy threshold: `optimal_threshold = 100`.

 3. Loop through all energy readings and:

    - Print a warning message if the energy production drops below the threshold.

 4. Calculate the average energy production. Save to the variable `average_energy`.

    - Print the average energy production value. We have provided the code for you to calculate the average energy.

 5. Check if the average energy production meets the optimal threshold:

    - If the average energy production is greater than or equal to the optimal threshold, print a message indicating optimal performance.

    - If the average energy production is below the optimal threshold, print a message indicating sub-optimal performance.

    We have provided the print statements for you.

```python
# Define the energy production measurements
...

# Define the optimal energy threshold
...

# Iterate through energy readings
...

    # Check if the energy production is below the threshold
    ...
        print(f"Warning: Energy production dropped below threshold at {energy} kWh!")

# Calculate the average energy production
...

# Print the average energy production
print(f"Average energy production: {average_energy:.2f} kWh")

# Check if the average energy production is optimal
...
    print("The average energy production meets the optimal performance threshold.")
...
    print("The average energy production is below the optimal performance threshold!")
```

```python
grader.check("Renewable-Energy-Monitoring")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week4-practicequiz", "1_practicequiz_q")
```

Tests for Renewable-Energy-Monitoring (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(6.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(6.0)
>>> log_variable('total-points', f'Reading-Week-X, 1_practicequiz_q', 6.0)
>>> question_id = 'Renewable-Energy-Monitoring-1'
>>> max_score = 1.0
>>> score = 0
>>> condition = isinstance(energy_readings, list) and len(energy_readings) > 0 and all((isinstance(e, (int, float)) for e in energy_readings))
>>> assert condition, 'Energy readings must be a list of numerical values.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicequiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(energy_readings))
```
failure_message: Failed: Energy readings are not defined correctly.
success_message: Success: Energy readings are defined correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Renewable-Energy-Monitoring-2'
>>> max_score = 1.0
>>> score = 0
>>> condition = isinstance(optimal_threshold, (int, float)) and optimal_threshold > 0
>>> assert condition, 'Optimal threshold must be a positive numerical value.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicequiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(optimal_threshold))
```
failure_message: Failed: The optimal energy threshold is not defined correctly.
success_message: Success: The optimal energy threshold is defined correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Renewable-Energy-Monitoring-3'
>>> max_score = 1.0
>>> score = 0
>>> iteration_count = 0
>>> for energy in energy_readings:
...     iteration_count += 1
>>> expected_count = len(energy_readings)
>>> condition = iteration_count == expected_count
>>> assert condition, f'Expected iteration over {expected_count} readings, but iterated {iteration_count} times.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicequiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(iteration_count))
>>> responses = update_responses(question_id, str(expected_count))
```
failure_message: Failed: The code does not correctly iterate through energy readings.
success_message: Success: The code correctly iterates through energy readings!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Renewable-Energy-Monitoring-4'
>>> max_score = 1.0
>>> score = 0
>>> warnings = []
>>> for energy in energy_readings:
...     if energy < 50:
...         warnings.append(energy)
>>> expected_warnings = [e for e in energy_readings if e < 50]
>>> condition = warnings == expected_warnings
>>> assert condition, f'Expected warnings: {expected_warnings}, but got: {warnings}.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicequiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(warnings))
>>> responses = update_responses(question_id, str(expected_warnings))
```
failure_message: Failed: The code does not correctly check energy production against the threshold.
success_message: Success: The code correctly checks energy production against the threshold!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Renewable-Energy-Monitoring-5'
>>> max_score = 1.0
>>> score = 0
>>> expected_average = sum(energy_readings) / len(energy_readings)
>>> condition = abs(average_energy - expected_average) < 1e-06
>>> assert condition, f'Expected average energy: {expected_average}, but got: {average_energy}.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicequiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(average_energy))
>>> responses = update_responses(question_id, str(expected_average))
```
failure_message: Failed: The code does not correctly calculate the average energy production.
success_message: Success: The code correctly calculates the average energy production!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Renewable-Energy-Monitoring-6'
>>> max_score = 1.0
>>> score = 0
>>> is_optimal = average_energy >= optimal_threshold
>>> expected_result = 'The average energy production meets the optimal performance threshold.' if is_optimal else 'The average energy production is below the optimal performance threshold!'
>>> condition = is_optimal and average_energy >= optimal_threshold or (not is_optimal and average_energy < optimal_threshold)
>>> assert condition, expected_result
>>> if (condition, expected_result):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_practicequiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(average_energy))
>>> responses = update_responses(question_id, str(optimal_threshold))
>>> responses = update_responses(question_id, str(is_optimal))
```
failure_message: Failed: The code does not correctly check if the average energy production is optimal.
success_message: Success: The code correctly checks if the average energy production is optimal!
