```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_homework", "week_2", "homework", assignment_points = 57.75, assignment_tag = 'week2-homework')

# Initialize Otter
import otter
grader = otter.Notebook("1_homework.ipynb")
```

# 🏠 Python Programming: Explore the Nutshell Studies 🕵️‍♀️🔍

<iframe width="1536" height="864" src="https://www.youtube.com/embed/2-6ndwvK3UI" title="Murder Is Her Hobby: Frances Glessner Lee and The Nutshell Studies of Unexplained Death" frameborder="0" allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share" referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>

The Nutshell Studies of Unexplained Death were created by Frances Glessner Lee as miniature dioramas replicating real crime scenes. These models are used to train investigators in observing details, solving mysteries, and analyzing evidence.  

![](./assets/figures/francis.png) 

In these Python exercises, you'll recreate parts of this fascinating process by solving problems involving data types, numerical precision, lists, dictionaries, and basic calculations.


### Question 1: Scaling the Mystery 🧮📏

![](./assets/figures/nutshell-1.jpg)

"Everything must be scaled to perfection!"

The Nutshell Studies require precise scaling to recreate scenes. You are given:
- The real room width is 15.5 feet.
- The model room width is given as a string: `"4.25"` inches.

#### Tasks:
1. Convert the room width and length to inches. Save the results in variables `room_width_inches` and `room_length_inches`.
2. Convert the variable `model_width` to a float, replace the original variable. To do this you should use the `float()` function.
3. Calculate the scale ratio by dividing the model width by the room width. Save the result in a variable `scale_ratio`.
4. Calculate the real-world area:
   - Compute the real-world area of the rectangular room using the formula:  
     $$
     \text{Real-world area} = \text{width (in feet)} \times \text{length (in feet)}
     $$
   - This will give the area in square feet.
   - Save the result in a variable `real_area_sqft`.

5. Calculate the scaled diorama area:
   - First, calculate the area of the real room in square inches by using the width and length in inches and multiplying them.
   - Scale this area down using the square of the scale ratio:  
     
     $$
     \text{Scaled diorama area} = \frac{\text{Real area (in square inches)}}{(\text{scale ratio})^2}
     $$
     
  - Save the result in a variable `scaled_area_sqin`.

```python
def question_1():
    # Given information
    # Room dimensions
    room_width_feet = 15.5  # Real room width in feet
    room_length_feet = 20   # Real room length in feet
    model_width = "4.25"  # Model width as a string

    # Convert dimensions to inches for consistency
    ...

    # Convert the model width to a float
    ...

    # Calculate the scale ratio
    ...

    # Real-world area (in square feet) and scaled diorama area (in square inches)
    ...

    return room_width_inches, room_length_inches, model_width, scale_ratio, real_area_sqft, scaled_area_sqin


# We have provided this code for you to test your implementation
room_width_inches, room_length_inches, model_width, scale_ratio, real_area_sqft, scaled_area_sqin = question_1()

# Print results
print(f"Scale of the model: {scale_ratio:.6f}")
print(f"Real-world area: {real_area_sqft:.2f} square feet")
print(f"Scaled diorama area: {scaled_area_sqin:.2f} square inches")
```

```python
grader.check("scaling-the-mystery")
```

### Question 2: Cataloging Crime Scene Objects 📝🔪

![](./assets/figures/nutshell-2.jpg)

"What’s in the crime scene?"

A crime scene contains various objects: `"knife"`, `"glass"`, `"rug"`, `"chair"`, and `"lamp"`.  
To investigate, you'll need to create and modify a list.

#### Tasks:
1. Build a list of object, `objects`, containing the objects mentioned above. 
    - make sure to add a comment saying `List of objects found at the crime scene`.
2. Add three more objects: `"hat"`, `"pen"`, `"rope"`.
    - There are several ways to add items to a list. You can use the `append()` method or the `+` operator.


3. Calculate the number of objects at the crime scene and save the result in a variable `total_objects`. To do this, use the `len()` function.

3. Retrieve and print the 3rd, 5th, and last objects from the list.
    - Retrieve the third object and assign it to the variable `third_object`.
    - Retrieve the fifth object and assign it to the variable `fifth_object`.
    - Retrieve the last object and assign it to the variable `last_object`. Make sure this will always get the last object, regardless of the number of objects in the list, by using the index `-1`.

4. Write the code to print information: 
    - The updated list of objects, "Updated list of objects: [list of objects]".
    - The total number of objects, "Total number of objects: [total_objects]".
    - The third object, "3rd object: [third_object]".
    - The fifth object, "5th object: [fifth_object]".
    - The last object, "Last object: [last_object]".

```python
def question_2():
    # List of objects in the crime scene
    # Build a list of object, `objects`, with the following objects: "knife", "glass", "rug", "chair", "lamp"
    ...

    # Add three more objects: `"hat"`, `"pen"`, `"rope"`.
    ...

    # Calculate the total number of objects using the `len()` function
    ...

    # Retrieve specific objects

    # Retrieve the third object and assign it to the variable `third_object`.
    ...

    # Retrieve the fifth object and assign it to the variable `fifth_object`.
    ...

    # Retrieve the last object and assign it to the variable `last_object`.
    ...
        
    # The updated list of objects, "Updated list of objects: [list of objects]". 
    ...

    # The total number of objects, "Total number of objects: [total_objects]".
    ...

    # The third object, "3rd object: [third_object]".
    ...

    # The fifth object, "5th object: [fifth_object]".
    ...

    # The last object, "Last object: [last_object]".
    ...
    
    return objects, total_objects, third_object, fifth_object, last_object

objects, total_objects, third_object, fifth_object, last_object = question_2()
```

```python
grader.check("cataloging-crime-scenes")
```

## Question 3: Object Details Dictionary 🗂️🔍

![](./assets/figures/nutshell-3.jpg)

"Track every detail!"

Each object in the Nutshell diorama has attributes like its position, whether it’s a clue, and its priority for solving the case.

### Tasks:
1. Create a dictionary `objects_details` with the following details (this should be a nested dictionary):
   - knife: position = `"near victim's hand"`, is_clue = `True`, priority = `8`
   - glass: position = `"on the table"`, is_clue = `False`, priority = `3`
   - rug: position = `"center of the room"`, is_clue = `False`, priority = `2`
   - lamp: position = `"overturned near the corner"`, is_clue = `True`, priority = `6`
2. Update the priority of `"glass"` to `7`.  
3. Retrieve and print the position and priority of the `"knife"` and `"rug"`. Save the knife position and priority in variables `knife_position` and `knife_priority`, and the rug position and priority in variables `rug_position` and `rug_priority`. 

```python
import copy

def question_3():
    # Dictionary of objects with details
    #     1. Create a dictionary with the following details (this should be a nested dictionary):
    #    - **knife**: position = `"near victim's hand"`, is_clue = `True`, priority = `8`
    #    - **glass**: position = `"on the table"`, is_clue = `False`, priority = `3`
    #    - **rug**: position = `"center of the room"`, is_clue = `False`, priority = `2`
    #    - **lamp**: position = `"overturned near the corner"`, is_clue = `True`, priority = `6`
    ...

    # Do not modify this code, it is used for grading
    original_details = copy.deepcopy(objects_details)

    # Update the priority of the "glass"
    # 2. Update the **priority** of `"glass"` to `7`.
    ...

    # Retrieve details for "knife" and "rug"
    # 3. Retrieve and print the **position** and **priority** of the `"knife"` and `"rug"`.
    # Save the knife position and priority in variables `knife_position` and `knife_priority`,
    # and the rug position and priority in variables `rug_position` and `rug_priority`.
    ...
    
    return (
        original_details, 
        objects_details,
        knife_position,
        knife_priority,
        rug_position,
        rug_priority,
    )

original_details, objects_details, knife_position, knife_priority, rug_position, rug_priority = question_3()

# Print results
print("Updated dictionary:", objects_details)
print("\nDetails of knife:")
print("Position:", knife_position)
print("Priority:", knife_priority)
print("\nDetails of rug:")
print("Position:", rug_position)
print("Priority:", rug_priority)
```

```python
grader.check("Object-Details-Dictionary")
```

## Bonus Challenge: Handling Missing Objects 🕵️‍♂️❓

![](./assets/figures/nutshell-4.webp)

The investigators occasionally encounter missing data for certain objects in the crime scene. Use dictionary methods like `.get()` to safely handle cases where some objects might not exist in the dictionary.

### Scenario:
The following objects and their priority scores are provided:
- knife: priority = 8
- glass: priority = 7
- rug: priority = 2
- lamp: priority = 6

An investigator asks about the priority of the following objects:
1. knife
2. hat (not in the dictionary)
3. rug
4. book (not in the dictionary)

### Tasks:
1. Use the `.get()` method to retrieve the priorities of the requested objects.  
2. Assign a default priority of `0` for any objects not found in the dictionary.  

```python
def bonus():

    # Dictionary of objects with their priority scores
    priority_scores = {
        "knife": 8,
        "glass": 7,
        "rug": 2,
        "lamp": 6
    }

    # Use the `.get()` method to safely retrieve priorities, with a default value of 0 for missing objects
    # Save the retrieved priorities in variables knife_priority, hat_priority, rug_priority, and book_priority
    ...
    
    return knife_priority, hat_priority, rug_priority, book_priority

knife_priority, hat_priority, rug_priority, book_priority = bonus()

# Print the retrieved priorities
print("Retrieved priority scores:")
print(f"Knife: {knife_priority}")
print(f"Hat: {hat_priority}")
print(f"Rug: {rug_priority}")
print(f"Book: {book_priority}")
```

```python
grader.check("handeling-missing-objects")
```

```python
from IPython.display import IFrame

IFrame(
    "https://americanart.si.edu/exhibitions/nutshells/inside",
    width=800,
    height=800,
)
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week2-homework", "1_homework")
```

Tests for Object-Details-Dictionary (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(14.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(57.75)
>>> log_variable('total-points', f'Reading-Week-X, 1_homework', 57.75)
>>> question_id = 'Object-Details-Dictionary-1'
>>> max_score = 2.0
>>> score = 0
>>> responses = update_responses(question_id, original_details)
>>> original_details, _, _, _, _, _ = question_3()
>>> expected_keys = {'knife', 'glass', 'rug', 'lamp'}
>>> actual_keys = {key.strip().lower() for key in original_details.keys()}
>>> assert expected_keys == actual_keys, f'Expected dictionary keys: {expected_keys}, but got: {actual_keys}.'
>>> expected_original_details = {'knife': {'position': "near victim's hand", 'is_clue': True, 'priority': 8}, 'glass': {'position': 'on the table', 'is_clue': False, 'priority': 3}, 'rug': {'position': 'center of the room', 'is_clue': False, 'priority': 2}, 'lamp': {'position': 'overturned near the corner', 'is_clue': True, 'priority': 6}}
>>> assert original_details == expected_original_details, 'The original dictionary is incorrect.'
>>> if expected_keys == actual_keys and original_details == expected_original_details:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The dictionary structure is incorrect or missing required keys.
success_message: The initial dictionary is correctly structured and initialized.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Object-Details-Dictionary-2'
>>> max_score = 2.0
>>> score = 0
>>> responses = update_responses(question_id, objects_details['glass']['priority'])
>>> _, objects_details, _, _, _, _ = question_3()
>>> glass_priority = objects_details['glass']['priority']
>>> assert glass_priority == 7, "The priority of 'glass' should be updated to 7."
>>> if glass_priority == 7:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The priority of "glass" was not correctly updated to 7.
success_message: The priority of "glass" is correctly updated to 7.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Object-Details-Dictionary-3'
>>> max_score = 2.0
>>> score = 0
>>> responses = update_responses(question_id, knife_position)
>>> _, _, knife_position, _, _, _ = question_3()
>>> expected_knife_position = "near victim's hand"
>>> assert knife_position.strip().lower() == expected_knife_position.strip().lower(), "The position of 'knife' is incorrect."
>>> if knife_position.strip().lower() == expected_knife_position.strip().lower():
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The position of "knife" is incorrect.
success_message: The position of "knife" is correctly retrieved.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Object-Details-Dictionary-4'
>>> max_score = 2.0
>>> score = 0
>>> responses = update_responses(question_id, knife_priority)
>>> _, _, _, knife_priority, _, _ = question_3()
>>> expected_knife_priority = 8
>>> assert knife_priority == expected_knife_priority, "The priority of 'knife' is incorrect."
>>> if knife_priority == expected_knife_priority:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The priority of "knife" is incorrect.
success_message: The priority of "knife" is correctly retrieved.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Object-Details-Dictionary-5'
>>> max_score = 1.5
>>> score = 0
>>> responses = update_responses(question_id, rug_position)
>>> _, _, _, _, rug_position, _ = question_3()
>>> expected_rug_position = 'center of the room'
>>> assert rug_position.strip().lower() == expected_rug_position.strip().lower(), "The position of 'rug' is incorrect."
>>> if rug_position.strip().lower() == expected_rug_position.strip().lower():
...     score = 1.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The position of "rug" is incorrect.
success_message: The position of "rug" is correctly retrieved.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Object-Details-Dictionary-6'
>>> max_score = 1.5
>>> score = 0
>>> responses = update_responses(question_id, rug_priority)
>>> _, _, _, _, _, rug_priority = question_3()
>>> expected_rug_priority = 2
>>> assert rug_priority == expected_rug_priority, "The priority of 'rug' is incorrect."
>>> if rug_priority == expected_rug_priority:
...     score = 1.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The priority of "rug" is incorrect.
success_message: The priority of "rug" is correctly retrieved.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Object-Details-Dictionary-7'
>>> max_score = 3.0
>>> score = 0
>>> responses = update_responses(question_id, objects_details)
>>> _, objects_details, _, _, _, _ = question_3()
>>> expected_final_objects_details = {'knife': {'position': "near victim's hand", 'is_clue': True, 'priority': 8}, 'glass': {'position': 'on the table', 'is_clue': False, 'priority': 7}, 'rug': {'position': 'center of the room', 'is_clue': False, 'priority': 2}, 'lamp': {'position': 'overturned near the corner', 'is_clue': True, 'priority': 6}}
>>> assert objects_details == expected_final_objects_details, 'The final dictionary structure is incorrect after updates.'
>>> if objects_details == expected_final_objects_details:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The final dictionary structure is incorrect after updates.
success_message: The final dictionary structure is correct after updates.


Tests for cataloging-crime-scenes (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> import inspect
>>> max_question_points = str(25.25)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(57.75)
>>> log_variable('total-points', f'Reading-Week-X, 1_homework', 57.75)
>>> question_id = 'cataloging-crime-scenes-1'
>>> max_score = 3.5
>>> score = 0
>>> function_source = inspect.getsource(question_2)
>>> condition = 'len(' in function_source
>>> assert condition, 'The `len()` function is not used to calculate the total number of objects.'
>>> condition_2 = '[-1]' in function_source
>>> assert condition_2, 'The `[-1]` indexing is not used to retrieve the last object.'
>>> if condition and condition_2:
...     score = 3.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Required methods (len() and [-1]) are missing in the implementation.
success_message: The required methods (len() and [-1]) are used in the implementation.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import sys
>>> import io
>>> question_id = 'cataloging-crime-scenes-2'
>>> max_score = 1.5
>>> score = 0
>>> output = io.StringIO()
>>> sys.stdout = output
>>> objects, _, _, _, _ = question_2()
>>> sys.output = sys.__stdout__
>>> condition = 'knife' in objects and 'glass' in objects and ('rug' in objects) and ('chair' in objects) and ('lamp' in objects)
>>> assert condition, 'The initial list is incorrect.'
>>> if condition:
...     score = 1.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The initial list of objects is missing or incorrect.
success_message: The initial list of objects is created correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import sys
>>> import io
>>> question_id = 'cataloging-crime-scenes-3'
>>> max_score = 1.5
>>> score = 0
>>> output = io.StringIO()
>>> sys.stdout = output
>>> objects, _, _, _, _ = question_2()
>>> sys.stdout = sys.__stdout__
>>> condition = 'hat' in objects and 'pen' in objects and ('rope' in objects)
>>> assert condition, 'New objects were not added correctly.'
>>> if condition:
...     score = 1.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The new objects are not added correctly.
success_message: New objects are added correctly to the list.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import sys
>>> import io
>>> question_id = 'cataloging-crime-scenes-4'
>>> max_score = 1.5
>>> score = 0
>>> output = io.StringIO()
>>> sys.stdout = output
>>> _, total_objects, _, _, _ = question_2()
>>> sys.stdout = sys.__stdout__
>>> assert total_objects == 8, f'Total objects should be 8, but got {total_objects}.'
>>> if total_objects == 8:
...     score = 1.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The total number of objects is incorrect.
success_message: The total number of objects is calculated correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import sys
>>> import io
>>> question_id = 'cataloging-crime-scenes-5'
>>> max_score = 3.5
>>> score = 0
>>> output = io.StringIO()
>>> sys.stdout = output
>>> _, _, third_object, fifth_object, last_object = question_2()
>>> sys.stdout = sys.__stdout__
>>> assert third_object == 'rug', f"The third object should be 'rug', but got {third_object}."
>>> assert fifth_object == 'lamp', f"The fifth object should be 'lamp', but got {fifth_object}."
>>> assert last_object == 'rope', f"The last object should be 'rope', but got {last_object}."
>>> if third_object == 'rug' and fifth_object == 'lamp' and (last_object == 'rope'):
...     score = 3.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Retrieval of specific objects is incorrect.
success_message: Retrieval of specific objects is correct.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import io
>>> import sys
>>> question_id = 'cataloging-crime-scenes-6'
>>> max_score = 1.5
>>> score = 0
>>> output = io.StringIO()
>>> sys.stdout = output
>>> question_2()
>>> sys.stdout = sys.__stdout__
>>> output_str = output.getvalue()
>>> condition1 = 'Updated list of objects:' in output_str
>>> condition2 = 'Total number of objects: 8' in output_str
>>> condition3 = '3rd object: rug' in output_str
>>> condition4 = '5th object: lamp' in output_str
>>> condition5 = 'Last object: rope' in output_str
>>> assert condition1, 'The updated list of objects is not printed correctly.'
>>> assert condition2, 'The total number of objects is not printed correctly.'
>>> assert condition3, 'The third object is not printed correctly.'
>>> assert condition4, 'The fifth object is not printed correctly.'
>>> assert condition5, 'The last object is not printed correctly.'
>>> if condition1 and condition2 and condition3 and condition4 and condition5:
...     score = 1.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The printed outputs are incorrect or missing.
success_message: The printed outputs are formatted correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> import inspect
>>> question_id = 'cataloging-crime-scenes-7'
>>> max_score = 1.5
>>> score = 0
>>> function_source = inspect.getsource(question_2)
>>> condition = 'append(' in function_source or '+' in function_source
>>> assert condition, 'The `append()` method or `+` operator is not used to add items to the list.'
>>> if condition:
...     score = 1.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The method for adding items (`append()` or `+`) is missing or incorrect.
success_message: The required methods for adding items (`append()` or `+`) are used correctly.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> import inspect
>>> question_id = 'cataloging-crime-scenes-8'
>>> max_score = 9.25
>>> score = 0
>>> function_source = inspect.getsource(question_2)
>>> updated_list_regex = 'print\\(\\s*f?[\'\\"]Updated list of objects[:]?.*?(\\{?objects\\}?)?[\'\\"]\\s*(,\\s*objects\\s*)?\\)'
>>> total_objects_regex = 'print\\(\\s*f?[\'\\"]Total number of objects[:]?.*?(\\{?total_objects\\}?)?[\'\\"]\\s*(,\\s*total_objects\\s*)?\\)'
>>> third_object_regex = 'print\\(\\s*f?[\'\\"]3rd object[:]?.*?(\\{?third_object\\}?)?[\'\\"]\\s*(,\\s*third_object\\s*)?\\)'
>>> fifth_object_regex = 'print\\(\\s*f?[\'\\"]5th object[:]?.*?(\\{?fifth_object\\}?)?[\'\\"]\\s*(,\\s*fifth_object\\s*)?\\)'
>>> last_object_regex = 'print\\(\\s*f?[\'\\"]Last object[:]?.*?(\\{?last_object\\}?)?[\'\\"]\\s*(,\\s*last_object\\s*)?\\)'
>>> condition_1 = re.search(updated_list_regex, function_source)
>>> condition_2 = re.search(total_objects_regex, function_source)
>>> condition_3 = re.search(third_object_regex, function_source)
>>> condition_4 = re.search(fifth_object_regex, function_source)
>>> condition_5 = re.search(last_object_regex, function_source)
>>> assert condition_1, 'The print statement for the updated list is hard-coded.'
>>> assert condition_2, 'The print statement for the total number of objects is hard-coded.'
>>> assert condition_3, 'The print statement for the third object is hard-coded.'
>>> assert condition_4, 'The print statement for the fifth object is hard-coded.'
>>> assert condition_5, 'The print statement for the last object is hard-coded.'
>>> if condition_1 and condition_2 and condition_3 and condition_4 and condition_5:
...     score = 9.25
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The print statements are hard-coded instead of dynamically generated.
success_message: The print statements are dynamically generated and not hard-coded.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> import inspect
>>> question_id = 'cataloging-crime-scenes-9'
>>> max_score = 1.5
>>> score = 0
>>> function_source = inspect.getsource(question_2)
>>> hard_coded_list_regex = 'objects\\s*=\\s*\\[\\s*[\'\\"]knife[\'\\"],\\s*[\'\\"]glass[\'\\"],\\s*[\'\\"]rug[\'\\"],\\s*[\'\\"]chair[\'\\"],\\s*[\'\\"]lamp[\'\\"],\\s*[\'\\"]hat[\'\\"],\\s*[\'\\"]pen[\'\\"],\\s*[\'\\"]rope[\'\\"]\\s*\\]'
>>> condition = not re.search(hard_coded_list_regex, function_source)
>>> assert condition, 'The list of objects is hard-coded in the function. Please ensure the list is dynamically constructed using methods like concatenation (`+`) or `append()`.'
>>> if condition:
...     score = 1.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The list values are hard-coded in the function instead of being dynamically generated.
success_message: The full list is dynamically generated and not hard-coded.


Tests for handeling-missing-objects (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(5.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(57.75)
>>> log_variable('total-points', f'Reading-Week-X, 1_homework', 57.75)
>>> question_id = 'handeling-missing-objects-1'
>>> max_score = 1.0
>>> score = 0
>>> responses = update_responses(question_id, knife_priority)
>>> knife_priority, _, _, _ = bonus()
>>> assert knife_priority == 8, "The priority of 'knife' should be 8."
>>> if knife_priority == 8:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The priority of "knife" is incorrect.
success_message: The priority of "knife" is correctly retrieved.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'handeling-missing-objects-2'
>>> max_score = 1.0
>>> score = 0
>>> responses = update_responses(question_id, hat_priority)
>>> _, hat_priority, _, _ = bonus()
>>> assert hat_priority == 0, "The priority of 'hat' should default to 0."
>>> if hat_priority == 0:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The priority of "hat" should default to 0 when missing.
success_message: The default priority of "hat" is correctly set to 0.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'handeling-missing-objects-3'
>>> max_score = 1.0
>>> score = 0
>>> responses = update_responses(question_id, rug_priority)
>>> _, _, rug_priority, _ = bonus()
>>> assert rug_priority == 2, "The priority of 'rug' should be 2."
>>> if rug_priority == 2:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The priority of "rug" is incorrect.
success_message: The priority of "rug" is correctly retrieved.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'handeling-missing-objects-4'
>>> max_score = 1.0
>>> score = 0
>>> responses = update_responses(question_id, book_priority)
>>> _, _, _, book_priority = bonus()
>>> assert book_priority == 0, "The priority of 'book' should default to 0."
>>> if book_priority == 0:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The priority of "book" should default to 0 when missing.
success_message: The default priority of "book" is correctly set to 0.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> import inspect
>>> question_id = 'handeling-missing-objects-5'
>>> max_score = 1.0
>>> score = 0
>>> function_source = inspect.getsource(bonus)
>>> get_usage = re.findall('priority_scores\\.get\\(', function_source)
>>> res_get_usage = len(get_usage) >= 4
>>> assert res_get_usage, 'The function should use the `.get()` method for each object, with a default value for missing objects.'
>>> if res_get_usage:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The function does not correctly use the `.get()` method for safe retrieval.
success_message: The function uses the `.get()` method correctly to handle missing keys.


Tests for scaling-the-mystery (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(13.5)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(57.75)
>>> log_variable('total-points', f'Reading-Week-X, 1_homework', 57.75)
>>> question_id = 'scaling-the-mystery-1'
>>> max_score = 1.5
>>> score = 0
>>> responses = update_responses(question_id, str(type(model_width)))
>>> responses = update_responses(question_id, model_width)
>>> room_width_inches, room_length_inches, model_width, scale_ratio, real_area_sqft, scaled_area_sqin = question_1()
>>> modal_width_float = isinstance(model_width, float)
>>> assert modal_width_float, 'model_width must be a float'
>>> assert model_width == 4.25, 'model_width value should be 4.25 after conversion'
>>> if modal_width_float and model_width == 4.25:
...     score = 1.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The model width string must be converted to a float.
success_message: The model width string is correctly converted to a float.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'scaling-the-mystery-2'
>>> max_score = 1.5
>>> score = 0
>>> responses = update_responses(question_id, room_width_inches)
>>> room_width_inches, room_length_inches, model_width, scale_ratio, real_area_sqft, scaled_area_sqin = question_1()
>>> expected_room_width_inches = 15.5 * 12
>>> assert room_width_inches == expected_room_width_inches, 'room_width_inches value is incorrect'
>>> if room_width_inches == expected_room_width_inches:
...     score = 1.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Real room width conversion to inches is incorrect.
success_message: Real room width is correctly converted to inches.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> import inspect
>>> question_id = 'scaling-the-mystery-3'
>>> max_score = 1.5
>>> score = 0
>>> responses = update_responses(question_id, room_width_inches)
>>> function_source = inspect.getsource(question_1)
>>> condition = re.findall('.*float\\(model_width\\).*', function_source)
>>> assert condition, 'You must convert the model width to a float, using the float() function.'
>>> if condition:
...     score = 1.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The model width string must be converted to a float.
success_message: The model width string is correctly converted to a float.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'scaling-the-mystery-4'
>>> max_score = 1.5
>>> score = 0
>>> responses = update_responses(question_id, room_length_inches)
>>> room_width_inches, room_length_inches, model_width, scale_ratio, real_area_sqft, scaled_area_sqin = question_1()
>>> expected_room_length_inches = 20 * 12
>>> assert room_length_inches == expected_room_length_inches, 'room_length_inches value is incorrect'
>>> if room_length_inches == expected_room_length_inches:
...     score = 1.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Real room length conversion to inches is incorrect.
success_message: Real room length is correctly converted to inches.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'scaling-the-mystery-5'
>>> max_score = 1.5
>>> score = 0
>>> responses = update_responses(question_id, scale_ratio)
>>> room_width_inches, room_length_inches, model_width, scale_ratio, real_area_sqft, scaled_area_sqin = question_1()
>>> expected_scale_ratio = 4.25 / (15.5 * 12)
>>> assert scale_ratio == expected_scale_ratio, 'scale_ratio value is incorrect'
>>> if scale_ratio == expected_scale_ratio:
...     score = 1.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The scale ratio calculation is incorrect.
success_message: The scale ratio is correctly calculated.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'scaling-the-mystery-6'
>>> max_score = 1.5
>>> score = 0
>>> responses = update_responses(question_id, real_area_sqft)
>>> room_width_inches, room_length_inches, model_width, scale_ratio, real_area_sqft, scaled_area_sqin = question_1()
>>> expected_real_area_sqft = 15.5 * 20
>>> assert real_area_sqft == expected_real_area_sqft, 'real_area_sqft value is incorrect'
>>> if real_area_sqft == expected_real_area_sqft:
...     score = 1.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Real-world area calculation in square feet is incorrect.
success_message: The real-world area in square feet is correctly calculated.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import numpy as np
>>> question_id = 'scaling-the-mystery-7'
>>> max_score = 1.5
>>> score = 0
>>> responses = update_responses(question_id, scaled_area_sqin)
>>> room_width_inches, room_length_inches, model_width, scale_ratio, real_area_sqft, scaled_area_sqin = question_1()
>>> expected_real_area_sqin = 85501200.83
>>> res_scaled_area_sqin = np.isclose(scaled_area_sqin, expected_real_area_sqin, atol=0.001)
>>> assert res_scaled_area_sqin, 'scaled_area_sqin value is incorrect'
>>> if res_scaled_area_sqin:
...     score = 1.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Real-world area calculation in square inches is incorrect.
success_message: The real-world area in square inches is correctly calculated.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'scaling-the-mystery-8'
>>> max_score = 1.5
>>> score = 0
>>> responses = update_responses(question_id, scale_ratio)
>>> room_width_inches, room_length_inches, model_width, scale_ratio, real_area_sqft, scaled_area_sqin = question_1()
>>> expected_scale_ratio = 0.022849462
>>> res_ratio_output = abs(scale_ratio - expected_scale_ratio) < 1e-06
>>> assert res_ratio_output, 'Scale ratio output is incorrect'
>>> if res_ratio_output:
...     score = 1.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The scale ratio output is incorrect.
success_message: The scale ratio is correctly output.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'scaling-the-mystery-9'
>>> max_score = 1.5
>>> score = 0
>>> responses = update_responses(question_id, scaled_area_sqin)
>>> room_width_inches, room_length_inches, model_width, scale_ratio, real_area_sqft, scaled_area_sqin = question_1()
>>> res_scaled_area_sqin = abs(scaled_area_sqin - 85501200.83044983) < 0.01
>>> assert res_scaled_area_sqin, 'Scaled diorama area output is incorrect'
>>> if res_scaled_area_sqin:
...     score = 1.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_homework', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The scaled diorama area output is incorrect.
success_message: The scaled diorama area is correctly output.
